---
id: 66fbcf750a62784cf11f5633
title: Étape 9
challengeType: 1
dashedName: step-9
---

# --description--

Crée une fonction `getDifficultyLevel` qui prend en paramètre un nombre indiquant le temps de cuisson.

Si le temps de cuisson est inférieur ou égal à `30`, la fonction doit retourner `"easy"`. S'il est inférieur ou égal à `60`, la fonction doit retourner `"medium"`. Sinon, la fonction doit retourner `"hard"`.

# --hints--

Tu dois créer une fonction `getDifficultyLevel`.

```js
assert.isFunction(getDifficultyLevel);
```

Ta fonction `getDifficultyLevel` doit avoir un seul paramètre.

```js
assert.lengthOf(getDifficultyLevel, 1);
```

Ta fonction `getDifficultyLevel` doit retourner `"easy"` lorsque le temps de cuisson est inférieur ou égal à `30`.

```js
assert.strictEqual(getDifficultyLevel(10), "easy");
assert.strictEqual(getDifficultyLevel(20), "easy");
assert.strictEqual(getDifficultyLevel(29), "easy");
assert.strictEqual(getDifficultyLevel(30), "easy");
```

Ta fonction `getDifficultyLevel` doit retourner `"medium"` lorsque le temps de cuisson est supérieur à `31` et inférieur ou égal à `60`.

```js
assert.strictEqual(getDifficultyLevel(31), "medium");
assert.strictEqual(getDifficultyLevel(40), "medium");
assert.strictEqual(getDifficultyLevel(50), "medium");
assert.strictEqual(getDifficultyLevel(60), "medium");
```

Ta fonction `getDifficultyLevel` doit retourner `"hard"` lorsque le temps de cuisson est supérieur à `60`.

```js
assert.strictEqual(getDifficultyLevel(61), "hard");
assert.strictEqual(getDifficultyLevel(65), "hard");
assert.strictEqual(getDifficultyLevel(70), "hard");
assert.strictEqual(getDifficultyLevel(75), "hard");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

--fcc-editable-region--

--fcc-editable-region--
```
