---
id: 66fbcf750a62784cf11f5630
title: Étape 7
challengeType: 1
dashedName: step-7
---

# --description--

Tu dois maintenant ajouter les trois objets au tableau `recipes`. Pour ce faire, tu peux utiliser la méthode `push()`.

Utilise la méthode `push()` pour ajouter tous les objets recette au tableau `recipes`. Veille à ajouter `recipe1`, `recipe2` et `recipe3` dans cet ordre.

Supprime aussi les variables `recipe1Name`, `recipe2CookingTime` et `recipe3Ingredients`, ainsi que les instructions `console.log` qui affichent ces variables.

# --hints--

Tu dois ajouter tous les objets recette au tableau `recipes`. Veille à les ajouter dans l'ordre où ils sont déclarés.

```js
assert.deepEqual(recipes, [recipe1, recipe2, recipe3]);
```

Tu dois supprimer la variable `recipe1Name`.

```js
assert.throws(() => recipe1Name, ReferenceError);
```

Tu dois supprimer l'instruction `console.log(recipe1Name);`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe1Name\s*\)\s*;?/);
```

Tu dois supprimer la variable `recipe2CookingTime`.

```js
assert.throws(() => recipe2CookingTime, ReferenceError);
```

Tu dois supprimer l'instruction `console.log(recipe2CookingTime);`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe2CookingTime\s*\)\s*;?/);
```

Tu dois supprimer la variable `recipe3Ingredients`.

```js
assert.throws(() => recipe3Ingredients, ReferenceError);
```

Tu dois supprimer l'instruction `console.log(recipe3Ingredients);`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe3Ingredients\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--
const recipe1Name = recipe1.name;
console.log(recipe1Name);

const recipe2CookingTime = recipe2.cookingTime;
console.log(recipe2CookingTime);

const recipe3Ingredients = recipe3.ingredients;
console.log(recipe3Ingredients);
--fcc-editable-region--
```
