---
id: 66fbcf750a62784cf11f5634
title: Étape 10
challengeType: 1
dashedName: step-10
---

# --description--

Il est temps de tester chacune des fonctions. Tu peux utiliser n'importe laquelle des recettes pour cela, mais ce tutoriel commencera avec `recipe1`.

Crée deux nouvelles variables : `recipe1TotalIngredients` et `recipe1DifficultyLevel`. Affecte-leur les valeurs en appelant la fonction correspondante pour chaque variable et en passant la propriété `recipe1` appropriée.

Enfin, affiche chaque variable dans la console pour voir les résultats.

# --hints--

Tu dois créer une variable `recipe1TotalIngredients`.

```js
assert.isNotNull(recipe1TotalIngredients);
```

Ta variable `recipe1TotalIngredients` doit être définie sur `getTotalIngredients(recipe1.ingredients)`.

```js
assert.equal(recipe1TotalIngredients, getTotalIngredients(recipe1.ingredients));
```

Tu dois afficher `recipe1TotalIngredients` dans la console.

```js
assert.match(code, /console\.log\(recipe1TotalIngredients\);?/);
```

Tu dois créer une variable `recipe1DifficultyLevel`.

```js
assert.isNotNull(recipe1DifficultyLevel);
```

Ta variable `recipe1DifficultyLevel` doit être définie sur `getDifficultyLevel(recipe1.cookingTime)`.

```js
assert.equal(recipe1DifficultyLevel, getDifficultyLevel(recipe1.cookingTime));
```

Tu dois afficher `recipe1DifficultyLevel` dans la console.

```js
assert.match(code, /console\.log\(recipe1DifficultyLevel\);?/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

function getDifficultyLevel(cookingTime) {
  if (cookingTime <= 30) {
    return "easy";
  } else if (cookingTime <= 60) {
    return "medium";
  } else {
    return "hard";
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
