---
id: 66fbcf750a62784cf11f5632
title: Étape 8
challengeType: 1
dashedName: step-8
---

# --description--

Crée une fonction `getTotalIngredients` qui prend un seul argument, représentant un tableau d'ingrédients, et retourne le nombre d'ingrédients du tableau passé à la fonction.

# --hints--

Tu dois créer une fonction `getTotalIngredients`.

```js
assert.isFunction(getTotalIngredients);
```

Ta fonction `getTotalIngredients` doit avoir un seul paramètre.

```js
assert.lengthOf(getTotalIngredients, 1);
```

Ta fonction `getTotalIngredients` doit retourner un nombre.

```js
assert.isNumber(getTotalIngredients(recipe1.ingredients));
```

`getTotalIngredients(recipe1.ingredients)` doit retourner `4`.

```js
assert.strictEqual(getTotalIngredients(recipe1.ingredients), 4);
```

`getTotalIngredients(recipe2.ingredients)` doit retourner `5`.

```js
assert.strictEqual(getTotalIngredients(recipe2.ingredients), 5);
```

`getTotalIngredients(recipe3.ingredients)` doit retourner `3`.

```js
assert.strictEqual(getTotalIngredients(recipe3.ingredients), 3);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

--fcc-editable-region--

--fcc-editable-region--
```
