---
id: 66fbcf750a62784cf11f562d
title: Étape 2
challengeType: 1
dashedName: step-2
---

# --description--

Crée un objet nommé `recipe1`. À l'intérieur de l'objet `recipe1`, crée une propriété `name` avec la valeur `"Spaghetti Carbonara"`.

Toujours à l'intérieur de l'objet `recipe1`, crée une propriété `ingredients` avec un tableau comme valeur. Le tableau doit contenir `spaghetti`, `Parmesan cheese`, `pancetta` et `black pepper`.

# --hints--

Tu dois créer un objet nommé `recipe1`.

```js
assert.isObject(recipe1);
```

Ton objet `recipe1` doit avoir une propriété `name`.

```js
assert.property(recipe1, 'name');
```

Ta propriété `name` doit être définie à `Spaghetti Carbonara`.

```js
assert.equal(recipe1.name, 'Spaghetti Carbonara');
```

Ton objet `recipe1` doit avoir une propriété `ingredients`.

```js
assert.property(recipe1, 'ingredients');
```

Ta propriété `ingredients` doit avoir une valeur de type tableau.

```js
assert.isArray(recipe1.ingredients);
```

Le tableau de ta propriété `ingredients` doit contenir `spaghetti`, `Parmesan cheese`, `pancetta` et `black pepper`.

```js
assert.deepEqual(recipe1.ingredients, ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"]);
```

# --seed--

## --seed-contents--

```js
const recipes = [];
--fcc-editable-region--

--fcc-editable-region--
```
