---
id: 66fbcdfed6848e48947dbbc9
title: Étape 1
challengeType: 1
dashedName: step-1
---

# --description--

Dans ce workshop, tu vas créer un traqueur de recettes en utilisant les objets JavaScript.

Commence par créer un tableau vide nommé `recipes`. C'est le tableau dans lequel tu ajouteras les objets recette plus tard.

# --hints--

Tu dois créer un tableau nommé `recipes`.

```js
assert.isArray(recipes);
```

Ton tableau `recipes` doit être vide.

```js
assert.isEmpty(recipes);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
