---
id: 66fbcf750a62784cf11f562e
title: Étape 4
challengeType: 1
dashedName: step-4
---

# --description--

Crée un objet `recipe2` avec les propriétés et valeurs suivantes :

| Clé | Valeur   |
| ----------- | ------- |
| `name`       | `Chicken Curry` |
| `ingredients`       | `["chicken breast", "coconut milk", "curry powder", "onion", "garlic"]` |
| `cookingTime`       | `42` |
| `totalIngredients`   | `null`   |
| `difficultyLevel`   | `""`   |

# --hints--

Tu dois créer un objet nommé `recipe2`.

```js
assert.isObject(recipe2);
```

Ton objet `recipe2` doit avoir une propriété `name`.

```js
assert.property(recipe2, "name");
```

Ta propriété `name` doit être définie à `Chicken Curry`.

```js
assert.equal(recipe2.name, "Chicken Curry");
```

Ton objet `recipe2` doit avoir une propriété `ingredients`.

```js
assert.property(recipe2, "ingredients");
```

Ta propriété `ingredients` doit avoir une valeur de type tableau.

```js
assert.isArray(recipe2.ingredients);
```

Le tableau de ta propriété `ingredients` doit contenir `chicken breast`, `coconut milk`, `curry powder`, `onion` et `garlic`.

```js
assert.deepEqual(recipe2.ingredients, ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"]);
```

Ton objet `recipe2` doit avoir une propriété `cookingTime`.

```js
assert.property(recipe2, "cookingTime");
```

La valeur de ta propriété `cookingTime` doit être un nombre.

```js
assert.isNumber(recipe2.cookingTime);
```

Ta propriété `cookingTime` doit être définie à `42`.

```js
assert.equal(recipe2.cookingTime, 42);
```

Ton objet `recipe2` doit avoir une propriété `totalIngredients`.

```js
assert.property(recipe2, "totalIngredients");
```

Ta propriété `totalIngredients` doit être définie à `null`.

```js
assert.isNull(recipe2.totalIngredients);
```

Ton objet `recipe2` doit avoir une propriété `difficultyLevel`.

```js
assert.property(recipe2, "difficultyLevel");
```

Ta propriété `difficultyLevel` doit être définie à une chaîne vide.

```js
assert.deepEqual(recipe2.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
  };

--fcc-editable-region--

--fcc-editable-region--
```
