---
id: 66fbcf750a62784cf11f5635
title: Étape 11
challengeType: 1
dashedName: step-11
---

# --description--

Tu peux maintenant remplir chaque élément du tableau `recipes` avec des valeurs pour les propriétés `totalIngredients` et `difficultyLevel`.

Pour l'instant, accède aux propriétés `totalIngredients` et `difficultyLevel` de `recipe1` et définis-les sur les résultats appropriés des appels de fonction et des arguments.

# --hints--

Tu dois accéder à la propriété `totalIngredients` de `recipe1`.

```js
assert.isNotNull(recipe1.totalIngredients);
```

Tu dois affecter le résultat de l'appel de `getTotalIngredients` avec `recipe1.ingredients` à la propriété `totalIngredients` de `recipe1`.

```js
assert.strictEqual(recipe1.totalIngredients, getTotalIngredients(recipe1.ingredients));
```

Tu dois accéder à la propriété `difficultyLevel` de `recipe1`.

```js
assert.isNotEmpty(recipe1.difficultyLevel);
```

Tu dois affecter le résultat de l'appel de `getDifficultyLevel` avec `recipe1.cookingTime` à la propriété `cookingTime` de `recipe1`.

```js
assert.strictEqual(recipe1.difficultyLevel, getDifficultyLevel(recipe1.cookingTime));
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

function getDifficultyLevel(cookingTime) {
  if (cookingTime <= 30) {
    return "easy";
  } else if (cookingTime <= 60) {
    return "medium";
  } else {
    return "hard";
  }
}

const recipe1TotalIngredients = getTotalIngredients(recipe1.ingredients);
console.log(recipe1TotalIngredients);

const recipe1DifficultyLevel = getDifficultyLevel(recipe1.cookingTime);
console.log(recipe1DifficultyLevel);

--fcc-editable-region--

--fcc-editable-region--
```
