---
id: 66ff02792b3042a2b811f475
title: Étape 3
challengeType: 1
dashedName: step-3
---

# --description--

Ajoute les propriétés suivantes à l'objet `recipe1` :

| Clé | Valeur   |
| ----------- | ------- |
|`cookingTime`|`22`|
|`totalIngredients`|`null`|
|`difficultyLevel`|`""`|

Les propriétés ayant pour valeurs `null` et une chaîne vide seront mises à jour plus tard, après que tu les auras calculées.

# --hints--

Ton objet `recipe1` doit avoir une propriété `cookingTime`.

```js
assert.property(recipe1, "cookingTime");
```

Ta propriété `cookingTime` doit être définie à `22`.

```js
assert.equal(recipe1.cookingTime, 22);
```

Ton objet `recipe1` doit avoir une propriété `totalIngredients`.

```js
assert.property(recipe1, "totalIngredients");
```

Ta propriété `totalIngredients` doit être définie à `null`.

```js
assert.isNull(recipe1.totalIngredients);
```

Ton objet `recipe1` doit avoir une propriété `difficultyLevel`.

```js
assert.property(recipe1, "difficultyLevel");
```

Ta propriété `difficultyLevel` doit être définie à une chaîne vide.

```js
assert.deepEqual(recipe1.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
--fcc-editable-region--

--fcc-editable-region--
};
```
