---
id: 670e4f45f7116c0f216a5177
title: Étape 6
challengeType: 1
dashedName: step-6
---

# --description--

Avant de continuer, tu dois t'entraîner à accéder aux propriétés d'un objet.

Tu peux utiliser soit la notation par point (`.`), soit la notation par crochets (`[]`) pour ce faire. Voici un exemple :

```js
const person = {
  name: "John",
  age: 30,
  job: "Software Engineer"
};

console.log(person.name); // John
console.log(person['age']);  // 30
```

Accède à la propriété `name` de `recipe1`, et affecte-la à la variable `recipe1Name`.

Ensuite, accède à la propriété `cookingTime` de `recipe2` et affecte-la à la variable `recipe2CookingTime`.

Enfin, accède à la propriété `ingredients` de `recipe3` et affecte-la à la variable `recipe3Ingredients`.

Assure-toi que toutes les variables que tu as créées sont affichées dans la console.

# --hints--

Tu dois créer une variable `recipe1Name`.

```js
assert.isDefined(recipe1Name);
```

Tu dois affecter la valeur de la propriété `name` de `recipe1` à ta variable `recipe1name`.

```js
assert.strictEqual(recipe1Name, recipe1.name);
```

Tu dois afficher `recipe1Name` dans la console.

```js
assert.match(code, /console\.log\(\s*recipe1Name\s*\)/);
```

Tu dois créer une variable `recipe2CookingTime`.

```js
assert.isDefined(recipe2CookingTime);
```

Tu dois affecter la valeur de la propriété `cookingTime` de `recipe2` à ta variable `recipe2CookingTime`.

```js
assert.strictEqual(recipe2CookingTime, recipe2.cookingTime);
```

Tu dois afficher `recipe2CookingTime` dans la console.

```js
assert.match(code, /console\.log\(recipe2CookingTime\)/);
```

Tu dois créer une variable `recipe3Ingredients`.

```js
assert.isDefined(recipe3Ingredients);
```

Tu dois affecter la valeur de la propriété `ingredients` de `recipe3` à ta variable `recipe3Ingredients`.

```js
assert.strictEqual(recipe3Ingredients, recipe3.ingredients);
```

Tu dois afficher `recipe3Ingredients` dans la console.

```js
assert.match(code, /console\.log\(recipe3Ingredients\)/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--

--fcc-editable-region--
```
