---
id: 69161f4c0c9c8c48b3c843c5
title: Étape 5
challengeType: 1
dashedName: step-5
---

# --description--

Crée un objet `recipe3` avec les propriétés et valeurs suivantes :

|Clé|Valeur|
|-----------|-------|
|`name`|`Vegetable Stir Fry`|
|`ingredients`|`["broccoli", "carrot", "bell pepper"]`|
|`cookingTime`|`15`|
|`totalIngredients`|`null`|
|`difficultyLevel`|`""`|

# --hints--

Tu dois créer un objet nommé `recipe3`.

```js
assert.isObject(recipe3);
```

Ton objet `recipe3` doit avoir une propriété `name`.

```js
assert.property(recipe3, "name");
```

Ta propriété `name` doit être définie à `Vegetable Stir Fry`.

```js
assert.equal(recipe3.name, "Vegetable Stir Fry");
```

Ton objet `recipe3` doit avoir une propriété `ingredients`.

```js
assert.property(recipe3, "ingredients");
```

Ta propriété `ingredients` doit avoir une valeur de type tableau.

```js
assert.isArray(recipe3.ingredients);
```

Le tableau de ta propriété `ingredients` doit contenir `broccoli`, `carrot` et `bell pepper`.

```js
assert.deepEqual(recipe3.ingredients, ["broccoli", "carrot", "bell pepper"]);
```

Ton objet `recipe3` doit avoir une propriété `cookingTime`.

```js
assert.property(recipe3, "cookingTime");
```

La valeur de ta propriété `cookingTime` doit être un nombre.

```js
assert.isNumber(recipe3.cookingTime);
```

Ta propriété `cookingTime` doit être définie à `15`.

```js
assert.equal(recipe3.cookingTime, 15);
```

Ton objet `recipe3` doit avoir une propriété `totalIngredients`.

```js
assert.property(recipe3, "totalIngredients");
```

Ta propriété `totalIngredients` doit être définie à `null`.

```js
assert.isNull(recipe3.totalIngredients);
```

Ton objet `recipe3` doit avoir une propriété `difficultyLevel`.

```js
assert.property(recipe3, "difficultyLevel");
```

Ta propriété `difficultyLevel` doit être définie à une chaîne vide.

```js
assert.deepEqual(recipe3.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--

--fcc-editable-region--
```
