---
id: 6140cdebd39d6a101e747529
title: Étape 6
challengeType: 0
dashedName: step-6
---

# --description--

Crée un sélecteur pour tes éléments `.line`. Commence par définir `background-color` sur `black`, `width` sur `50%` et `height` sur `2px`.

# --hints--

Tu devrais avoir un sélecteur `.line`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line'));
```

Ton sélecteur `.line` devrait avoir une propriété `background-color` définie sur `black`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line')?.backgroundColor, "black");
```

Ton sélecteur `.line` devrait avoir une propriété `width` définie sur `50%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line')?.width, "50%");
```

Ton sélecteur `.line` devrait avoir une propriété `height` définie sur `2px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line')?.height, "2px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--

--fcc-editable-region--
```
