---
id: 6140cbeec34e970dfe75e710
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

Crée un sélecteur pour ton élément `.wheel`. Commence par définir `border` sur `2px solid black`, `border-radius` sur `50%` et `margin-left` sur `50px`.

# --hints--

Tu devrais avoir un sélecteur `.wheel`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.wheel'));
```

Ton sélecteur `.wheel` devrait avoir une propriété `border` définie sur `2px solid black`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.border, '2px solid black');
```

Ton sélecteur `.wheel` devrait avoir une propriété `border-radius` définie sur `50%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.borderRadius, '50%');
```

Ton sélecteur `.wheel` devrait avoir une propriété `margin-left` définie sur `50px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.marginLeft, '50px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
