---
id: 6140d4bc9db3c81c51a09ab7
title: Étape 15
challengeType: 0
dashedName: step-15
---

# --description--

En continuant le modèle, sélectionne les éléments `.cabin` suivants et applique-leur les règles spécifiques :

- Le deuxième `.cabin` devrait avoir la propriété `right` définie sur `17%` et la propriété `top` définie sur `93.5%`.
- Le troisième `.cabin` devrait avoir la propriété `right` définie sur `67%` et la propriété `top` définie sur `93.5%`.
- Le quatrième `.cabin` devrait avoir la propriété `left` définie sur `-8.5%` et la propriété `top` définie sur `50%`.
- Le cinquième `.cabin` devrait avoir la propriété `left` définie sur `17%` et la propriété `top` définie sur `7%`.
- Le sixième `.cabin` devrait avoir la propriété `right` définie sur `17%` et la propriété `top` définie sur `7%`.

# --hints--

Tu devrais avoir un sélecteur `.cabin:nth-of-type(2)`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(2)'));
```

Ton sélecteur `.cabin:nth-of-type(2)` devrait avoir une propriété `right` définie sur `17%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(2)')?.right, '17%');
```

Ton sélecteur `.cabin:nth-of-type(2)` devrait avoir une propriété `top` définie sur `93.5%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(2)')?.top, '93.5%');
```

Tu devrais avoir un sélecteur `.cabin:nth-of-type(3)`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(3)'));
```

Ton sélecteur `.cabin:nth-of-type(3)` devrait avoir une propriété `right` définie sur `67%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(3)')?.right,'67%');
```

Ton sélecteur `.cabin:nth-of-type(3)` devrait avoir une propriété `top` définie sur `93.5%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(3)')?.top, '93.5%');
```

Tu devrais avoir un sélecteur `.cabin:nth-of-type(4)`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(4)'));
```

Ton sélecteur `.cabin:nth-of-type(4)` devrait avoir une propriété `left` définie sur `-8.5%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(4)')?.left, '-8.5%');
```

Ton sélecteur `.cabin:nth-of-type(4)` devrait avoir une propriété `top` définie sur `50%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(4)')?.top, '50%');
```

Tu devrais avoir un sélecteur `.cabin:nth-of-type(5)`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(5)'));
```

Ton sélecteur `.cabin:nth-of-type(5)` devrait avoir une propriété `left` définie sur `17%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(5)')?.left, '17%');
```

Ton sélecteur `.cabin:nth-of-type(5)` devrait avoir une propriété `top` définie sur `7%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(5)')?.top, '7%');
```

Tu devrais avoir un sélecteur `.cabin:nth-of-type(6)`.

```js
const def = (s) => new __helpers.CSSHelp(document).getStyle(s)
assert.exists(def('.cabin:nth-of-type(6)') || def('.cabin:last-of-type'));
```

Ton sélecteur `.cabin:nth-of-type(6)` devrait avoir une propriété `right` définie sur `17%`.

```js
const right = (s) => new __helpers.CSSHelp(document).getStyle(s)?.right
assert.equal(right('.cabin:nth-of-type(6)') || right('.cabin:last-of-type'),'17%');
```

Ton sélecteur `.cabin:nth-of-type(6)` devrait avoir une propriété `top` définie sur `7%`.

```js
const top = (s) => new __helpers.CSSHelp(document).getStyle(s)?.top
assert.equal(top('.cabin:nth-of-type(6)') || top('.cabin:last-of-type'),'7%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}
.line:nth-of-type(3) {
  transform: rotate(120deg);
}
.line:nth-of-type(4) {
  transform: rotate(180deg);
}
.line:nth-of-type(5) {
  transform: rotate(240deg);
}
.line:nth-of-type(6) {
  transform: rotate(300deg);
}

.cabin {
  background-color: red;
  width: 20%;
  height: 20%;
  position: absolute;
  border: 2px solid;
  transform-origin: 50% 0%;
}

.cabin:nth-of-type(1) {
  right: -8.5%;
  top: 50%;
}
--fcc-editable-region--

--fcc-editable-region--
```
