---
id: 6140c7e645d8e905819f1dd4
title: Étape 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Pour commencer, ajoute un élément `link` pour le fichier `./styles.css`.

# --hints--

Ton code devrait contenir un élément `link`.

```js
assert.match(code, /<link/);
```

Tu devrais avoir un élément `link`.

```js
assert.lengthOf(document.querySelectorAll('link'), 1);
```

Ton élément `link` devrait être dans ton élément `head`.

```js
assert.exists(document.querySelector('head > link'));
```

Ton élément `link` devrait avoir un attribut `rel` avec la valeur `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Ton élément `link` devrait avoir un attribut `href` avec la valeur `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>

  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css

```
