---
id: 6140d10d50636e14695013b2
title: Étape 9
challengeType: 0
dashedName: step-9
---

# --description--

Crée un sélecteur pour cibler ton deuxième élément `.line`. Définis la propriété `transform` sur `rotate(60deg)`.

Rappelle-toi que la propriété `transform` te permet de manipuler la forme d'un élément. Ici, utiliser la valeur `rotate(60deg)` fera pivoter l'élément autour de son point `transform-origin` de 60 degrés dans le sens horaire.

# --hints--

Tu devrais créer un sélecteur `.line:nth-of-type(2)`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(2)'));
```

Ton sélecteur `.line:nth-of-type(2)` devrait avoir une propriété `transform` définie sur `rotate(60deg)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(2)')?.transform, 'rotate(60deg)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

--fcc-editable-region--

--fcc-editable-region--
```
