---
id: 6140cd32d018ed0f600eefce
title: Étape 4
challengeType: 0
dashedName: step-4
---

# --description--

Définis la propriété `position` du sélecteur `.wheel` sur `absolute`. Définis `height` et `width` sur `55vw`.

# --hints--

Ton sélecteur `.wheel` devrait avoir une propriété `position` définie sur `absolute`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.position, 'absolute');
```

Ton sélecteur `.wheel` devrait avoir une propriété `height` définie sur `55vw`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.height, '55vw');
```

Ton sélecteur `.wheel` devrait avoir une propriété `width` définie sur `55vw`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.width, '55vw');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
}
--fcc-editable-region--
```
