---
id: 6140d0069049f5139d78da40
title: Étape 8
challengeType: 0
dashedName: step-8
---

# --description--

La propriété `transform-origin` sert à définir le point autour duquel une transformation CSS est appliquée. Par exemple, quand tu appliques une transformation `rotate` (comme tu le feras plus tard dans ce projet), `transform-origin` détermine autour de quel point l'élément pivote.

Donne au sélecteur `.line` une propriété `transform-origin` de `0% 0%`. Cela décalera le point d'origine à `0%` depuis la gauche et `0%` depuis le haut, ce qui le place dans le coin supérieur gauche de l'élément.

# --hints--

Ton sélecteur `.line` devrait avoir une propriété `transform-origin` définie sur `0% 0%`.

```js
const transformOrigin = new __helpers.CSSHelp(document).getStyle('.line')?.transformOrigin;
assert.oneOf(transformOrigin, ['0% 0%', '0% 0% 0px']);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--
.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
}
--fcc-editable-region--
```
