---
id: 6140c9d35015ae0ba0c250e8
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

Ajoute un `div` dans ton élément `body` et donne-lui une `class` de `wheel`.

Dans ton nouveau `div`, ajoute six éléments `span` avec une `class` définie sur `line`, et six éléments `div` avec une `class` définie sur `cabin`.

# --hints--

Tu devrais créer un nouveau `div` dans ton élément `body`.

```js
const divs = [...document.querySelector('body')?.children].filter(child => child?.localName === 'div');
assert.lengthOf(divs, 1);
```

Ton `div` dans ton élément `body` devrait avoir une `class` de `wheel`.

```js
assert.isTrue(document.querySelector('body div')?.classList?.contains('wheel'));
```

Ton élément `.wheel` devrait contenir six éléments `span`.

```js
assert.lengthOf(document.querySelectorAll('.wheel span'), 6);
```

Tes six éléments `span` dans l'élément `.wheel` devraient avoir une `class` de `line`.

```js
const spans = [...document.querySelectorAll('.wheel span')];
assert.isTrue(spans?.every(span => span?.classList?.contains('line')));
assert.lengthOf(document.querySelectorAll('.line'), 6);
```

Ton élément `.wheel` devrait contenir six éléments `div`.

```js
assert.lengthOf(document.querySelectorAll('.wheel div'), 6);
```

Tes six éléments `div` dans l'élément `.wheel` devraient avoir une `class` de `cabin`.

```js
const divs = [...document.querySelectorAll('.wheel div')];
assert.isTrue(divs?.every(div => div?.classList?.contains('cabin')));
assert.lengthOf(document.querySelectorAll('.cabin'), 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css

```
