---
id: 6140d1a351e88f159ed54fca
title: Étape 10
challengeType: 0
dashedName: step-10
---

# --description--

En suivant le même modèle, crée un sélecteur séparé pour le troisième `.line`, le quatrième `.line`, le cinquième `.line` et le sixième `.line`.

Définis la propriété `transform` du troisième `.line` sur `rotate(120deg)`, celle du quatrième sur `rotate(180deg)`, celle du cinquième sur `rotate(240deg)` et celle du sixième sur `rotate(300deg)`.

# --hints--

Tu devrais créer un sélecteur `.line:nth-of-type(3)`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(3)'));
```

Ton sélecteur `.line:nth-of-type(3)` devrait avoir une propriété `transform` définie sur `rotate(120deg)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(3)')?.transform, 'rotate(120deg)');
```

Tu devrais créer un sélecteur `.line:nth-of-type(4)`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(4)'));
```

Ton sélecteur `.line:nth-of-type(4)` devrait avoir une propriété `transform` définie sur `rotate(180deg)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(4)')?.transform, 'rotate(180deg)');
```

Tu devrais créer un sélecteur `.line:nth-of-type(5)`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(5)'));
```

Ton sélecteur `.line:nth-of-type(5)` devrait avoir une propriété `transform` définie sur `rotate(240deg)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(5)')?.transform, 'rotate(240deg)');
```

Tu devrais créer un sélecteur `.line:nth-of-type(6)`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(6)'));
```

Ton sélecteur `.line:nth-of-type(6)` devrait avoir une propriété `transform` définie sur `rotate(300deg)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(6)')?.transform, 'rotate(300deg)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}

--fcc-editable-region--

--fcc-editable-region--
```
