---
id: 5f33071498eb2472b87ddee4
title: Étape 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Dans ce projet, tu vas apprendre les bases du **CSS** (Cascading Style Sheets — feuilles de style en cascade) en construisant un menu de café. Le CSS est le langage qui sert à **styler** un document HTML : il décrit comment les éléments doivent s'afficher à l'écran (couleurs, tailles, polices, espacements…).

Comme tu l'as vu dans les dernières étapes du Cat Photo App, toute page web a besoin d'une structure de base. Chaque document HTML doit avoir une déclaration `DOCTYPE` et un élément `html`. Le `DOCTYPE` indique au navigateur la version d'HTML utilisée. Et `html` est l'élément **racine** qui contient tous les autres.

```html
<!DOCTYPE html>
<html lang="en">
<!--all other elements go here-->
</html>
```

Ajoute la balise `<!DOCTYPE html>`, puis un élément `html` avec un attribut `lang="en"`.

# --hints--

Tu dois avoir la déclaration `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Tu dois avoir une balise ouvrante `<html lang="en">` avec la langue réglée sur l'anglais.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Tu dois avoir une balise fermante `</html>`.

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
