---
id: 5f3326b143638ee1a09ff1e3
title: Étape 4
challengeType: 0
dashedName: step-4
---

# --description--

Pour préparer du **vrai contenu**, ajoute un élément `body` en dessous du `head`.

# --hints--

Tu dois avoir une balise ouvrante `<body>`.

```js
assert(code.match(/<body>/i));
```

Tu dois avoir une balise fermante `</body>`.

```js
assert(code.match(/<\/body>/i));
```

Ne touche pas au `head` — vérifie que sa balise fermante est toujours là.

```js
assert(code.match(/<head>/i));
assert(code.match(/<\/head>/i));
```

Le `body` doit venir **après** le `head`.

```js
assert(code.match(/<\/head>[.\n\s]*<body>/im));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
