---
id: 5f331e55dfab7a896e53c3a1
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

Le `title` fait partie de **plusieurs éléments** qui apportent des informations supplémentaires non visibles sur la page, mais utiles pour les moteurs de recherche ou pour la façon dont la page est affichée.

À l'intérieur du `head`, imbrique un élément `meta` avec un attribut `charset` qui vaut `utf-8`, pour indiquer au navigateur comment encoder les caractères de la page (notamment les accents).

Rappel : `meta` est un élément **void** — pas besoin de balise fermante.

# --hints--

Tu dois avoir une balise `meta`.

```js
assert.match(code, /<meta.*?\/?>/is);
```

`meta` est un élément void — pas de balise fermante `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

Ton `meta` doit avoir un attribut `charset`.

```js
assert.match(code, /<meta\s+charset\s*/i);
```

L'attribut `charset` doit valoir `utf-8`.

```js
assert.match(code, /charset\s*=\s*('|")utf-8\1/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
