---
id: 5f3313e74582ad9d063e3a38
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

Ajoute un élément `head` à l'intérieur du `html`, pour pouvoir y mettre un élément `title`. Le texte du `title` doit être `Cafe Menu`.

# --hints--

Tu dois avoir une balise ouvrante `<head>`.

```js
assert(code.match(/<head>/i));
```

Tu dois avoir une balise fermante `</head>`.

```js
assert(code.match(/<head>/i));
```

Tu dois avoir une balise ouvrante `<title>`.

```js
assert(code.match(/<title>/i));
```

Tu dois avoir une balise fermante `</title>`.

```js
assert(code.match(/<\/title>/i));
```

Le `title` doit être imbriqué dans le `head`.

```js
assert(code.match(/<head>\s*<title>.*<\/title>\s*<\/head>/si));
```

Le texte du `title` doit être `Cafe Menu`. Vérifie ton orthographe.

```js
assert.match(document.querySelector('title')?.innerText, /Cafe Menu/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
