---
id: 5f33294a6af5e9188dbdb8f3
title: Étape 5
challengeType: 0
dashedName: step-5
---

# --description--

C'est le moment d'ajouter du contenu au menu. Ajoute un élément `main` à l'intérieur du `body`. Il contiendra plus tard les informations sur les cafés et desserts proposés par le café.

# --hints--

Ton code doit avoir une balise ouvrante `<main>`.

```js
assert(code.match(/<main>/i));
```

Ton code doit avoir une balise fermante `</main>`.

```js
assert(code.match(/<\/main>/i));
```

Ne touche pas au `body` — vérifie que sa balise fermante est toujours là.

```js
assert.lengthOf(document.querySelectorAll('body'), 1);
```

Le `main` doit être à l'intérieur du `body`.

```js
const main = document.querySelector('main');
assert(main.parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
