---
id: 5f332a88dc25a0fd25c7687a
title: Étape 6
challengeType: 0
dashedName: step-6
---

# --description--

Le café s'appelle `CAMPER CAFE`. Ajoute un élément `h1` à l'intérieur du `main`. Mets-y le nom du café **en majuscules** pour qu'il ressorte.

# --hints--

Tu dois avoir une balise ouvrante `<h1>`.

```js
assert(code.match(/<h1>/i));
```

Tu dois avoir une balise fermante `</h1>`.

```js
assert(code.match(/<\/h1>/i));
```

Ne touche pas au `main`.

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

Le `h1` doit être imbriqué dans le `main`.

```js
assert.equal(document.querySelector('h1')?.parentElement?.tagName, "MAIN");
```

Le texte du `h1` doit être `CAMPER CAFE`.

```js
assert(code.match(/<h1>CAMPER CAFE<\/h1>/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
    </main>
--fcc-editable-region--
  </body>
</html>
```
