---
id: 5f3477ae9675db8bb7655b30
title: Étape 12
challengeType: 0
dashedName: step-12
---

# --description--

Dans l'étape précédente, vous avez utilisé <dfn>sélecteur de type</dfn> pour modeler la `h1` élément. Centrer le contenu du `h2` et les `p` éléments en ajoutant un nouveau sélecteur de type pour chacun à l'existant `style` élément.

# --hints--

Vous ne devriez pas changer l'existant `h1` sélecteur.

```js
const hasH1 = new __helpers.CSSHelp(document).getStyle('h1');
assert(hasH1);
```

Vous ne devriez pas ajouter un nouveau `style` marque. Ajouter les nouvelles règles CSS à l'existant `style` marque.

```js
assert.isAtMost(document.querySelectorAll('style').length, 2);
```

Vous devriez ajouter un nouveau `h2` sélecteur.

```js
const hasH2 = new __helpers.CSSHelp(document).getStyle('h2');
assert(hasH2);
```

Vous devriez ajouter un nouveau `p` sélecteur.

```js
const hasP = new __helpers.CSSHelp(document).getStyle('p');
assert(hasP);
```

Votre `h1` élément devrait avoir un `text-align` des `center`.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert(h1TextAlign === 'center');
```

Votre `h2` élément devrait avoir un `text-align` des `center`.

```js
const h2TextAlign = new __helpers.CSSHelp(document).getStyle('h2')?.getPropertyValue('text-align');
assert(h2TextAlign === 'center');
```

Votre `p` élément devrait avoir un `text-align` des `center`.

```js
const pTextAlign = new __helpers.CSSHelp(document).getStyle('p')?.getPropertyValue('text-align');
assert(pTextAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
      h1 {
        text-align: center;
      }
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

