---
id: 5f3477aefa51bfc29327200b
title: Étape 14
challengeType: 0
dashedName: step-14
---

# --description--

Vous avez style trois éléments en écrivant CSS à l'intérieur du `style` des étiquettes. Cela fonctionne, mais comme il y aura beaucoup plus de styles, il est préférable de mettre tous les styles dans un fichier séparé et un lien vers lui.

Une `styles.css` fichier a été créé pour vous. Vous pouvez basculer entre les fichiers avec les onglets en haut de l'éditeur.

Commencez par réécrire les styles que vous avez créés dans le `styles.css` fichier. Assurez-vous d'exclure l'ouverture et la fermeture `style` des étiquettes.

# --hints--

Votre `styles.css` fichier devrait avoir le `h1, h2, p` sélecteur de type.

```js
assert(code.slice(code.indexOf('</html>')).replace(/[\s\n]*/g, "").match(__helpers.concatRegex(__helpers.permutateRegex(['h1', 'h2', 'p'], { elementsSeparator: ',' }), /{/)));
```
 
Votre sélecteur devrait définir `text-align` propriété `center`.

```js
assert(code.slice(code.indexOf('</html>')).match(/text-align:\s*center;?/));
```

Votre code ne doit pas contenir `<style>` des étiquettes.

```js
assert.isFalse(/<\/html>(\n|.)*<\/?\s*style\s*>/.test(code));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
