---
id: 5f344fc1520b6719f2e35605
title: Étape 8
challengeType: 0
dashedName: step-8
---

# --description--

Il y aura deux sections au menu, une pour les cafés et une pour les desserts. Ajouter un `section` élément à l'intérieur de `main` élément donc vous avez un endroit pour mettre tous les cafés disponibles.

# --hints--

Tu devrais avoir une ouverture. `<section>` marque.

```js
assert(code.match(/<section\s*>/i));
```

Tu devrais fermer. `</section>` marque.

```js
assert(code.match(/<\/section\s*>/i));
```

Vous ne devriez pas changer votre `main` élément. Assurez-vous que vous n'avez pas supprimé l'étiquette de fermeture.

```js
assert.lengthOf(document.querySelectorAll('main'),1);
```

Votre `section` élément devrait être dans votre `main` élément.

```js
const main = document.querySelector('main');
const section = document.querySelector('section');
assert(section.parentElement === main);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```
