---
id: 5f356ed60a5decd94ab66986
title: Étape 22
challengeType: 0
dashedName: step-22
---

# --description--

Les commentaires dans CSS ressemblent à ceci :

```css
/* comment here */
```

Dans votre feuille de style, commentez la ligne contenant la `background-color` propriété et la valeur, de sorte que vous pouvez voir l'effet de `#menu` élément. Cela rendra le fond à nouveau blanc.

# --hints--

Vous devriez commenter la `background-color: burlywood;` ligne dans votre CSS.

```js
assert(code.match(/\/\*\s*background-color:\s*burlywood;?\s*\*\//i));
```


Votre `body` devrait avoir un fond blanc.

```js
const body = document.querySelector("body");
const compStyles = window.getComputedStyle(body);
const backgroundColor = compStyles?.getPropertyValue('background-color'); 
assert.equal(backgroundColor,"rgba(0, 0, 0, 0)")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
--fcc-editable-region--
  background-color: burlywood;
--fcc-editable-region--
}

h1, h2, p {
  text-align: center;
}

#menu {
  width: 300px;
}
```

