---
id: 5f356ed656a336993abd9f7c
title: Étape 25
challengeType: 0
dashedName: step-25
---

# --description--

Ensuite, vous voulez centrer la `#menu` horizontalement. Vous pouvez le faire en définissant son `margin-left` et `margin-right` propriétés à `auto`. Pensez à la marge comme un espace invisible autour d'un élément. En utilisant ces deux propriétés de marge, centrer la `#menu` élément à l'intérieur de `body` élément.

# --hints--

Vous devriez mettre le `margin-left` propriété `auto`.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['margin-left'] === 'auto');
assert(hasMargin);
```

Vous devriez mettre le `margin-right` propriété `auto`.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['margin-right'] === 'auto');
assert(hasMargin);
```

Vous devriez mettre le `margin-left` et `margin-right` propriétés de votre `#menu` à `auto`.

```js
const divMarginRight = new __helpers.CSSHelp(document).getStyle('#menu')?.getPropertyValue('margin-right');
const divMarginLeft = new __helpers.CSSHelp(document).getStyle('#menu')?.getPropertyValue('margin-left');
assert(divMarginRight === 'auto');
assert(divMarginLeft === 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

--fcc-editable-region--
#menu {
  width: 80%;
  background-color: burlywood;
}
--fcc-editable-region--
```
