---
id: 5f356ed6199b0cdef1d2be8f
title: Étape 26
challengeType: 0
dashedName: step-26
---

# --description--

Jusqu'à présent, vous avez utilisé des sélecteurs de type et d'id pour styler des éléments. Cependant, il est plus courant d'utiliser un sélecteur différent pour styler vos éléments.

A <dfn>sélecteur de classe</dfn> est défini par un nom avec un point directement devant lui, comme ceci:

```css
.class-name {
  styles
}
```

Modifier l'actuel `#menu` sélecteur dans un sélecteur de classe en remplaçant `#menu` avec une classe nommée `.menu`.

# --hints--

Vous devriez avoir un `.menu` sélecteur de classe.

```js
const hasMenu = new __helpers.CSSHelp(document).getStyle('.menu');
assert(hasMenu);
```

Vous ne devriez pas avoir `#menu` sélecteur.

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle('#menu');
assert(!hasDiv);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

--fcc-editable-region--
#menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
--fcc-editable-region--
```

