---
id: 69708d59833ceded7d00541f
title: Implémenter les opérations d'une liste chaînée
challengeType: 26
dashedName: implement-linked-list-operations
---

# --description--

Dans ce lab, tu vas implémenter des opérations supplémentaires pour une structure de données liste chaînée, en t'appuyant sur les fonctionnalités de base d'ajout et de suppression de nœuds.

**Objectif** : Remplir les user stories ci-dessous et faire passer tous les tests pour terminer le lab.

**User Stories**

1. Tu dois avoir une fonction `contains` qui accepte une liste chaînée et un élément. Elle doit renvoyer `true` si l'élément spécifié existe dans la liste chaînée, et `false` sinon.
1. Tu dois avoir une fonction `getAt` qui accepte une liste chaînée et un index. Elle doit renvoyer l'élément à l'index donné dans la liste chaînée. Si l'index est hors limites, elle doit renvoyer `undefined`.
1. Tu dois avoir une fonction `insertAt` qui accepte une liste chaînée, un index et un élément. Elle doit insérer l'élément donné à la position spécifiée dans la liste chaînée. Si l'index est hors limites, elle ne doit pas modifier la liste.
1. Tu dois avoir une fonction `removeAt` qui accepte une liste chaînée et un index. Elle doit retirer le nœud à l'index donné dans la liste chaînée. Si l'index est hors limites, elle ne doit pas modifier la liste.
1. Tu dois avoir une fonction `clear` qui accepte une liste chaînée. Elle doit retirer tous les éléments de la liste chaînée, en la réinitialisant effectivement à un état vide.

**Note :** Certains tests ultérieurs s'appuient sur des méthodes précédentes. Par exemple, si `getAt` n'est pas correctement implémentée, les tests de fonctions comme `insertAt` et `removeAt` peuvent échouer même lorsque ces fonctions sont proches d'être correctes.

# --hints--

`initList` doit renvoyer un objet avec `head` défini à `null` et `length` défini à `0`.

```js
const list = initList();
assert.isNull(list.head);
assert.strictEqual(list.length, 0);
```

`add` doit augmenter la longueur de la liste de un.

```js
const list = initList();
add(list, 'apple');
assert.strictEqual(list.length, 1);
```

`remove` doit diminuer la longueur de la liste de un.

```js
const list = initList();
add(list, 'apple');
add(list, 'banana');
remove(list, 'apple');
assert.strictEqual(list.length, 1);
assert.strictEqual(list.head.element, 'banana');
assert.isNull(list.head.next);
```

Tu dois avoir une fonction `contains`.

```js
assert.isFunction(contains);
```

`contains` doit renvoyer `true` si l'élément existe dans la liste.

```js
const list = initList();
add(list, 'apple');
add(list, 'banana');
add(list, 'cherry');
assert.isTrue(contains(list, 'banana'));
```

`contains` doit renvoyer `false` si l'élément n'existe pas dans la liste.

```js
const list = initList();
add(list, 'apple');
add(list, 'banana');
assert.isFalse(contains(list, 'cherry'));
```

`contains` doit renvoyer `false` pour une liste vide.

```js
const list = initList();
assert.isFalse(contains(list, 'apple'));
```

Tu dois avoir une fonction `getAt`.

```js
assert.isFunction(getAt);
```

`getAt` doit renvoyer l'élément à un index donné.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
add(list, 'third');
assert.strictEqual(getAt(list, 0), 'first');
assert.strictEqual(getAt(list, 1), 'second');
assert.strictEqual(getAt(list, 2), 'third');
```

`getAt` doit renvoyer `undefined` pour des indices invalides.

```js
const list = initList();
add(list, 'apple');
assert.isUndefined(getAt(list, -1));
assert.isUndefined(getAt(list, 1));
assert.isUndefined(getAt(list, 5));
```

`getAt` doit renvoyer `undefined` pour une liste vide.

```js
const list = initList();
assert.isUndefined(getAt(list, 0));
```

Tu dois avoir une fonction `insertAt`.

```js
assert.isFunction(insertAt);
```

`insertAt` doit insérer un élément au début de la liste.

```js
const list = initList();
add(list, 'second');
add(list, 'third');
insertAt(list, 0, 'first');
assert.strictEqual(getAt(list, 0), 'first');
assert.strictEqual(getAt(list, 1), 'second');
assert.strictEqual(getAt(list, 2), 'third');
assert.strictEqual(list.length, 3);
```

`insertAt` doit insérer un élément au milieu de la liste.

```js
const list = initList();
add(list, 'first');
add(list, 'third');
insertAt(list, 1, 'second');
assert.strictEqual(getAt(list, 0), 'first');
assert.strictEqual(getAt(list, 1), 'second');
assert.strictEqual(getAt(list, 2), 'third');
assert.strictEqual(list.length, 3);
```

`insertAt` doit insérer un élément à la fin de la liste.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
insertAt(list, 2, 'third');
assert.strictEqual(getAt(list, 0), 'first');
assert.strictEqual(getAt(list, 1), 'second');
assert.strictEqual(getAt(list, 2), 'third');
assert.strictEqual(list.length, 3);
```

`insertAt` doit insérer dans une liste vide à l'index `0`.

```js
const list = initList();
insertAt(list, 0, 'first');
assert.strictEqual(list.length, 1);
assert.strictEqual(getAt(list, 0), 'first');
assert.isNull(list.head.next);
```

`insertAt` ne doit pas modifier la liste pour des indices invalides.

```js
const list = initList();
add(list, 'apple');
insertAt(list, -1, 'invalid');
insertAt(list, 10, 'invalid');
assert.strictEqual(list.length, 1);
assert.strictEqual(getAt(list, 0), 'apple');
assert.isUndefined(getAt(list, 1));
```

Tu dois avoir une fonction `removeAt`.

```js
assert.isFunction(removeAt);
```

`removeAt` doit retirer l'élément à un index spécifié.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
add(list, 'third');
removeAt(list, 1);
assert.strictEqual(getAt(list, 0), 'first');
assert.strictEqual(getAt(list, 1), 'third');
assert.isUndefined(getAt(list, 2));
assert.strictEqual(list.length, 2);
```

`removeAt` doit retirer correctement le premier élément.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
removeAt(list, 0);
assert.strictEqual(getAt(list, 0), 'second');
assert.strictEqual(list.length, 1);
assert.isNull(list.head.next);
```

`removeAt` doit retirer correctement le dernier élément.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
add(list, 'third');
removeAt(list, 2);
assert.strictEqual(getAt(list, 0), 'first');
assert.strictEqual(getAt(list, 1), 'second');
assert.isUndefined(getAt(list, 2));
assert.strictEqual(list.length, 2);
```

`removeAt` doit réinitialiser la liste lors du retrait du seul élément.

```js
const list = initList();
add(list, 'only');
removeAt(list, 0);
assert.strictEqual(list.length, 0);
assert.isNull(list.head);
assert.isTrue(isEmpty(list));
```

`removeAt` ne doit pas modifier la liste pour des indices invalides.

```js
const list = initList();
add(list, 'apple');
removeAt(list, -1);
removeAt(list, 10);
assert.strictEqual(list.length, 1);
assert.strictEqual(getAt(list, 0), 'apple');
```

Tu dois avoir une fonction `clear`.

```js
assert.isFunction(clear);
```

`clear` doit retirer tous les éléments de la liste.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
add(list, 'third');
clear(list);
assert.strictEqual(list.length, 0);
assert.isNull(list.head);
assert.isTrue(isEmpty(list));
```

`clear` doit permettre de réutiliser la liste.

```js
const list = initList();
add(list, 'first');
clear(list);
clear(list);
add(list, 'second');
assert.strictEqual(list.length, 1);
assert.strictEqual(getAt(list, 0), 'second');
assert.isNull(list.head.next);
```

# --seed--

## --seed-contents--

```js
function initList() {
  return {
    head: null,
    length: 0
  };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
  const node = { element, next: null };

  if (isEmpty(list)) {
    list.head = node;
  } else {
    let current = list.head;
    while (current.next !== null) {
      current = current.next;
    }
    current.next = node;
  }

  list.length++;
}

function remove(list, element) {
  let previous = null;
  let current = list.head;

  while (current !== null && current.element !== element) {
    previous = current;
    current = current.next;
  }

  if (current === null) return;

  if (previous !== null) {
    previous.next = current.next;
  } else {
    list.head = current.next;
  }

  list.length--;
}

function contains(list, element) {

}

function getAt(list, index) {

}

function insertAt(list, index, element) {

}

function removeAt(list, index) {

}

function clear(list) {

}
```

# --solutions--

```js
function initList() {
  return {
    head: null,
    length: 0
  };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
  const node = { element, next: null };

  if (isEmpty(list)) {
    list.head = node;
  } else {
    let current = list.head;
    while (current.next !== null) {
      current = current.next;
    }
    current.next = node;
  }

  list.length++;
}

function remove(list, element) {
  let previous = null;
  let current = list.head;

  while (current !== null && current.element !== element) {
    previous = current;
    current = current.next;
  }

  if (current === null) return;

  if (previous !== null) {
    previous.next = current.next;
  } else {
    list.head = current.next;
  }

  list.length--;
}

function contains(list, element) {
  let current = list.head;
  while (current !== null) {
    if (current.element === element) {
      return true;
    }
    current = current.next;
  }
  return false;
}

function getAt(list, index) {
  if (index < 0 || index >= list.length) return undefined;
  let current = list.head;
  let i = 0;
  while (i < index) {
    current = current.next;
    i++;
  }
  return current.element;
}

function insertAt(list, index, element) {
  if (index < 0 || index > list.length) return;
  const node = { element, next: null };
  if (index === 0) {
    node.next = list.head;
    list.head = node;
    list.length++;
    return;
  }
  let current = list.head;
  let previous = null;
  let i = 0;
  while (i < index) {
    previous = current;
    current = current.next;
    i++;
  }
  previous.next = node;
  node.next = current;
  list.length++;
}

function removeAt(list, index) {
  if (index < 0 || index >= list.length) return;
  if (index === 0) {
    list.head = list.head.next;
    list.length--;
    return;
  }
  let current = list.head;
  let previous = null;
  let i = 0;
  while (i < index) {
    previous = current;
    current = current.next;
    i++;
  }
  previous.next = current.next;
  list.length--;
}

function clear(list) {
  list.head = null;
  list.length = 0;
}
```
