---
id: 6842a6cd9836f0114a5b7a8a
title: Crée un générateur de motifs numériques
challengeType: 27
dashedName: build-a-number-pattern-generator
---

# --description--

Dans ce lab, tu vas t'entraîner aux bases de Python en construisant une petite application qui crée un motif numérique.

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois définir une fonction nommée `number_pattern` qui prend un seul paramètre `n` (représentant un entier positif).
1. `number_pattern` doit utiliser une boucle `for`.
1. `number_pattern(n)` doit renvoyer une chaîne contenant tous les entiers de 1 jusqu'à `n` (inclus), séparés par un espace. Par exemple, `number_pattern(4)` doit renvoyer la chaîne `1 2 3 4`.
1. Si l'argument passé à la fonction n'est pas une valeur entière, la fonction doit renvoyer `Argument must be an integer value.`
1. Si l'argument passé à la fonction est inférieur à 1, la fonction doit renvoyer `Argument must be an integer greater than 0.`

# --hints--

Tu dois avoir une fonction `number_pattern`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('number_pattern')
    `)
})
```

La fonction `number_pattern` doit avoir un paramètre nommé `n`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('number_pattern').has_args('n')
    `)
})
```

`number_pattern(4)` doit renvoyer `1 2 3 4`.

```js
({
    test: () => runPython(`
    assert number_pattern(4) == '1 2 3 4'
    `)
})
```

`number_pattern(12)` doit renvoyer `1 2 3 4 5 6 7 8 9 10 11 12`.

```js
({
    test: () => runPython(`
    assert number_pattern(12) == '1 2 3 4 5 6 7 8 9 10 11 12'
    `)
})
```

`number_pattern` doit renvoyer une liste de nombres séparés par des espaces pour tout entier positif.

```js
({
    test: () => runPython(`
    from random import randint
    n = randint(1,65)
    ls = ' '.join(map(lambda n: str(n), list(range(1, n+1))))
    assert number_pattern(n) == ls

    n2 = randint(34,122)
    ls2 = ' '.join(map(lambda n: str(n), list(range(1, n2+1))))
    assert number_pattern(n2) == ls2

    `)
})
```

`number_pattern` doit renvoyer `Argument must be an integer value.` lorsqu'on lui passe une valeur qui n'est pas un entier.

```js
({
    test: () => runPython(`
    string = 'Argument must be an integer value.'
    assert number_pattern('a') == string
    assert number_pattern(3.5) == string
    assert number_pattern(lambda n: n) == string
    `)
})
```

`number_pattern` doit renvoyer `Argument must be an integer greater than 0.` lorsqu'on lui passe un entier non positif.

```js
({
    test: () => runPython(`
    from random import randint
    string = 'Argument must be an integer greater than 0.'
    assert number_pattern(0) == string
    assert number_pattern(-1185) == string
    assert number_pattern(-45) == string
    assert number_pattern(randint(-9999999, 0)) == string
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def number_pattern(n):
    result = "1"
    if isinstance(n, int):
        if n > 0:
            for i in range(2, n + 1):
                result += f" {i}"
            return result
        return "Argument must be an integer greater than 0."
    return "Argument must be an integer value."

```
