---
id: 67116d7584d0b469b14579bf
title: Étape 2
challengeType: 1
dashedName: step-2
---

# --description--

Dans le tableau `library`, crée un objet avec les propriétés et valeurs suivantes :

| Propriété | Valeur   |
| ----------- | ------- |
| `title` | `"Your Next Five Moves: Master the Art of Business Strategy"`|
| `author` | `"Patrick Bet-David and Greg Dinkin"`|
| `about` | `"A book on how to plan ahead"`|
| `pages` | `320` |

# --hints--

Tu dois avoir un objet dans le tableau `library`.

```js
assert.isObject(library[0])
```

Ton objet doit avoir une propriété `title`.

```js
assert.property(library[0], "title")
```

La valeur de ta propriété `title` doit être `"Your Next Five Moves: Master the Art of Business Strategy"`.

```js
assert.propertyVal(library[0], "title", "Your Next Five Moves: Master the Art of Business Strategy");
```

Ton objet doit avoir une propriété `author`.

```js
assert.property(library[0], "author")
```

Ta propriété `author` doit avoir la valeur `"Patrick Bet-David and Greg Dinkin"`.

```js
assert.propertyVal(library[0], "author", "Patrick Bet-David and Greg Dinkin");
```

Ton objet doit avoir une propriété `about`.

```js
assert.property(library[0], "about")
```

Ta propriété `about` doit avoir la valeur `"A book on how to plan ahead"`.

```js
assert.propertyVal(library[0], "about", "A book on how to plan ahead");
```

Ton objet doit avoir une propriété `pages`.

```js
assert.property(library[0], "pages")
```

La valeur de ta propriété `pages` doit être un nombre.

```js
assert.isNumber(library[0]?.pages)
```

Ta propriété `pages` doit être définie sur `320`.

```js
assert.propertyVal(library[0], "pages", 320);
```

# --seed--

## --seed-contents--

```js
const library = [
--fcc-editable-region--

--fcc-editable-region--
];
```
