---
id: 67116d7584d0b469b14579c0
title: Étape 3
challengeType: 1
dashedName: step-3
---

# --description--

Crée un autre objet dans le tableau `library` avec les propriétés et valeurs suivantes :

| Propriété | Valeur   |
| ----------- | ------- |
| `title` | `"Atomic Habits"`|
| `author` | `"James Clear"`|
| `about` | `"A practical book about discarding bad habits and building good ones"`|
| `pages` | `320` |

# --hints--


Tu dois avoir un objet comme deuxième élément de ton tableau `library`.

```js
assert.isObject(library[1]);
```

Ton objet doit avoir une propriété `title`.

```js
assert.property(library[1], "title")
```

Ta propriété `title` doit avoir `"Atomic Habits"` comme valeur.

```js
assert.propertyVal(library[1], "title", "Atomic Habits");
```

Ton objet doit avoir une propriété `author`.

```js
assert.property(library[1], "author")
```

Ta propriété `author` doit avoir `"James Clear"` comme valeur.

```js
assert.propertyVal(library[1], "author", "James Clear");
```

Ton objet doit avoir une propriété `about`.

```js
assert.property(library[1], "about")
```

Ta propriété `about` doit avoir la valeur `"A practical book about discarding bad habits and building good ones"`.

```js
assert.propertyVal(library[1], "about", "A practical book about discarding bad habits and building good ones");
```

Ton objet doit avoir une propriété `pages`.

```js
assert.property(library[1], "pages")
```

La valeur de ta propriété `pages` doit être un nombre.

```js
assert.isNumber(library[1]?.pages)
```

Ta propriété `pages` doit être définie sur `320`.

```js
assert.propertyVal(library[1], "pages", 320);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: 'A book on how to plan ahead',
    pages: 320,
  },
--fcc-editable-region--

--fcc-editable-region--
];
```
