---
id: 67116c175a77cb64c5c10d49
title: Étape 1
challengeType: 1
dashedName: step-1
---

# --description--

Dans les leçons précédentes, tu as appris à travailler avec des fonctions d'ordre supérieur comme `map`, `filter` et `reduce`.

Dans ce workshop, tu vas créer une application de gestion de bibliothèque qui te donnera l'occasion de t'entraîner avec différentes fonctions d'ordre supérieur.

Pour commencer, crée une variable appelée `library` et assigne-lui un tableau vide.

# --hints--

Tu dois créer un tableau `library`.

```js
assert.isArray(library)
```

Ton tableau `library` doit être vide.

```js
assert.isEmpty(library)
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
