---
id: 694648abde178bb8202d9511
title: Étape 3
challengeType: 20
dashedName: step-3
---

# --description--

Tu devrais maintenant voir le nom de l'élève affiché dans le terminal.

Python fournit une fonction nommée `type()` que tu peux utiliser pour vérifier le type d'une valeur.

```py
platform = 'freeCodeCamp'
print(type(platform)) # Output: <class 'str'>
```

Dans l'exemple ci-dessus, la sortie `<class 'str'>` signifie que la variable passée à la fonction `type()` est une chaîne.

Utilise la fonction `type()` avec `name` comme argument et affiche le résultat comme dans l'exemple. Vérifie dans le terminal que la sortie montre que `name` est du type `str` (chaîne).

# --hints--

Tu devrais afficher le résultat de l'appel `type(name)`.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(type(name))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
--fcc-editable-region--

--fcc-editable-region--
```
