---
id: 694648acde178bb8202d9514
title: Étape 7
challengeType: 20
dashedName: step-7
---

# --description--

Maintenant, tu dois ajouter l'âge de l'élève au bulletin. Pour cela, tu vas utiliser un entier, l'un des types numériques de Python.

Déclare une variable nommée `age` et assigne-lui la valeur entière `20`.

Puis, affiche la valeur et le type de données de `age` ensemble, séparés par une virgule. Vérifie dans le terminal que la sortie montre la valeur de `age` et son type `int` (entier).

# --hints--

Tu devrais avoir une variable `age`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("age")`) })
```

La variable `age` devrait stocker la valeur `20`. N'entoure pas la valeur de quotes.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("age").is_equivalent("age = 20")`) })
```

Tu devrais afficher `age` et `type(age)` en les séparant par une virgule.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(age, type(age))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

--fcc-editable-region--

--fcc-editable-region--

```
