---
id: 694648acde178bb8202d9512
title: Étape 4
challengeType: 20
dashedName: step-4
---

# --description--

Le bulletin devrait aussi indiquer si l'élève est actuellement inscrit. Cela peut se représenter avec une valeur booléenne.

Les valeurs booléennes représentent une condition oui-ou-non, et elles servent souvent à prendre des décisions dans le code. Il n'existe que deux valeurs booléennes : `True` et `False`.

Déclare une variable nommée `is_student` et assigne-lui la valeur `True`.

# --hints--

Tu devrais avoir une variable `is_student`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("is_student")`) })
```

La variable `is_student` devrait stocker la valeur `True`. N'entoure pas la valeur de quotes.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("is_student").is_equivalent("is_student = True")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
print(type(name))

--fcc-editable-region--

--fcc-editable-region--
```
