---
id: 694648acde178bb8202d9515
title: Étape 8
challengeType: 20
dashedName: step-8
---

# --description--

Maintenant, ajoute la note de l'élève.

Déclare une variable nommée `score` et assigne-lui la valeur `80.5`.

Même si `age` et `score` sont tous les deux des nombres, ils peuvent ne pas être du même type. Python fournit une fonction appelée `isinstance()` pour le vérifier.

```py
x = 10
print(isinstance(x, int)) # Output: True
```

Utilise `isinstance()` pour vérifier si `score` est un `int`, et affiche le résultat dans le terminal comme dans l'exemple ci-dessus.

# --hints--

Tu devrais avoir une variable `score`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("score")`) })
```

La variable `score` devrait stocker la valeur `80.5`. N'entoure pas la valeur de quotes.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("score").is_equivalent("score = 80.5")`) })
```

Tu devrais afficher le résultat de `isinstance(score, int)`.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(isinstance(score, int))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

--fcc-editable-region--

--fcc-editable-region--

```
