---
id: 694648acde178bb8202d9513
title: Étape 6
challengeType: 20
dashedName: step-6
---

# --description--

Le nom de l'élève devrait suivre le même format que les autres informations.

Supprime les affichages précédents de la variable `name`. Puis, affiche `name` et `type(name)` ensemble sur une seule ligne, séparés par une virgule, comme à l'étape précédente.

# --hints--

Tu ne devrais pas avoir `print(name)` dans ton code.

```js
({ test: () => runPython(`assert not _Node(_code).has_call("print(name)")`)})
```

Tu ne devrais pas avoir `print(type(name))` dans ton code.

```js
({ test: () => runPython(`assert not _Node(_code).has_call("print(type(name))")`)})
```

Tu devrais afficher `name` et `type(name)` dans le même appel `print()` en les séparant par une virgule.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(name, type(name))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
--fcc-editable-region--
print(name)
print(type(name))
--fcc-editable-region--

is_student = True
print(is_student, type(is_student))
```
