---
id: 694648acde178bb8202d9516
title: Étape 9
challengeType: 20
dashedName: step-9
---

# --description--

La sortie est `False`, ce qui montre que `score` n'est pas un `int`.

Un autre type de nombre courant en Python est `float`, qui représente un nombre avec des décimales. Remplace `int` par `float` dans l'appel `isinstance()` existant pour le confirmer.

# --hints--

Tu devrais remplacer `int` par `float` dans l'appel `isinstance(score, int)` existant.

```js
({ test: () => runPython(`
  assert not _Node(_code).has_call("print(isinstance(score, int))")
  assert _Node(_code).has_call("print(isinstance(score, float))")
`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
--fcc-editable-region--
print(isinstance(score, int))
--fcc-editable-region--

```
