---
id: 69461e4f98594012939bc1c6
title: Étape 1
challengeType: 20
dashedName: step-1
---

# --description--

Dans ce workshop, tu vas t'entraîner aux types de données Python en construisant une simple imprimante de bulletins.

Comme tu l'as appris dans les leçons précédentes, les variables s'assignent de cette façon :

```py
greeting = 'Hello World'
```

Dans l'exemple ci-dessus, la variable `greeting` a pour valeur une chaîne, c'est-à-dire une séquence de caractères entourée de quotes.

Crée une variable appelée `name` et assigne-lui la chaîne `Alice`. N'oublie pas d'entourer la valeur de quotes simples ou doubles, comme dans l'exemple.

**Important :** Python utilise l'indentation (les espaces au début d'une ligne) pour organiser le code. Pour ce workshop, assure-toi qu'il n'y a pas d'espaces en trop au début de chaque ligne de code. Ajouter des espaces supplémentaires provoquera une `IndentationError` et empêchera ton code de s'exécuter.

# --hints--

Tu devrais avoir une variable `name`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable('name')`) })
```

La variable `name` devrait stocker la valeur `'Alice'`.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("name").is_equivalent("name = 'Alice'")`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
