---
id: 694648acde178bb8202d9517
title: Étape 10
challengeType: 20
dashedName: step-10
---

# --description--

La sortie est `True`, ce qui confirme que `score` est un `float`.

Termine le bulletin en affichant la valeur de `score` ainsi que son type de données avec une seule instruction `print()`.

# --hints--

Tu devrais afficher `score` et `type(score)` en les séparant par une virgule.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(score, type(score))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
print(isinstance(score, float))
--fcc-editable-region--

--fcc-editable-region--

```

# --solutions--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
print(isinstance(score, float))
print(score, type(score))
```
