---
id: 6947a3e516bb8a4a681274a7
title: Étape 5
challengeType: 20
dashedName: step-5
---

# --description--

La fonction `print()` peut afficher plus d'une valeur à la fois. Sépare les valeurs par une virgule (`,`) pour les afficher sur la même ligne.

```py
subject = 'Python'
print(subject, type(subject)) # Output: Python <class 'str'>
```

Affiche `is_student` et `type(is_student)` sur la même ligne en utilisant une virgule `,` comme dans l'exemple. Puis, vérifie dans le terminal que la sortie montre la valeur de `is_student` et son type `bool` (booléen).

# --hints--

Tu devrais afficher `is_student` et `type(is_student)` en les séparant par une virgule.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(is_student, type(is_student))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
print(type(name))

is_student = True
--fcc-editable-region--

--fcc-editable-region--
```
