---
id: 6947a0c2b6de035dd8d3eea1
title: Étape 2
challengeType: 20
dashedName: step-2
---

# --description--

Tu peux afficher la valeur d'une variable avec la fonction `print()`.

```py
greeting = 'Hello World'
print(greeting) # Output: Hello World
```

Tu en apprendras plus sur les fonctions dans les leçons suivantes. Pour l'instant, sache qu'une fonction est un bloc de code réutilisable que tu peux appeler, ou invoquer, pour exécuter son code, et auquel tu peux passer des arguments.

Dans l'exemple ci-dessus, `print(greeting)` est un appel de fonction, et `greeting` est l'argument de la fonction.

En t'aidant de l'exemple, affiche la variable `name`. Vérifie la sortie dans le terminal.

# --hints--

Tu devrais appeler `print()` avec `name` comme argument.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(name)")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
--fcc-editable-region--

--fcc-editable-region--
```
