---
id: 699d20623e57e8ffee885457
title: Étape 8
challengeType: 1
dashedName: step-8
---

# --description--

Certaines fiches de catalogue ont des champs manquants. Par exemple, l'entrée 16 a un auteur vide, donc `parseCard` retourne actuellement une chaîne vide pour ce champ.

L'opérateur `||` fournit une valeur de repli : `a || b` vaut `a` lorsque `a` est truthy, et `b` sinon. Comme une chaîne vide est falsy, `"" || "Unknown"` vaut `"Unknown"`.

Mets à jour `title`, `author` et `location` dans l'instruction return pour utiliser `|| "Unknown"` comme valeur de repli. Laisse `year` pour l'étape suivante.

# --hints--

Lorsqu'elle est appelée avec une chaîne ayant un champ author vide, la propriété `author` de l'objet retourné doit être `"Unknown"`.

```js
assert.strictEqual(parseCard("X |  | 2000 | Y").author, "Unknown");
```

Lorsqu'elle est appelée avec une chaîne ayant un champ title vide, la propriété `title` de l'objet retourné doit être `"Unknown"`.

```js
assert.strictEqual(parseCard(" | A | 2000 | Y").title, "Unknown");
```

Lorsqu'elle est appelée avec une chaîne ayant un champ location vide, la propriété `location` de l'objet retourné doit être `"Unknown"`.

```js
assert.strictEqual(parseCard("X | A | 2000 | ").location, "Unknown");
```

Tu ne dois pas encore ajouter de valeur de repli pour `year`.

```js
assert.strictEqual(parseCard("X | A |  | Y").year, "");
```

Lorsque tous les champs sont présents, `parseCard` doit retourner leurs valeurs réelles et non `"Unknown"`.

```js
assert.strictEqual(parseCard("The Stand | King, Stephen | 1978 | Shelf K2").title, "The Stand");
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
--fcc-editable-region--
  return {
    title: title,
    author: author,
    year: year,
    location: location
  };
--fcc-editable-region--
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
