---
id: 699d1f8e51fcdc8230f779fb
title: Étape 1
challengeType: 1
dashedName: step-1
---

# --description--

Dans ce workshop, tu vas construire un système de catalogue pour une bibliothèque patrimoniale. Tu analyseras des données brutes de livres, tu rechercheras et regrouperas des entrées, tu valideras la qualité des données, et tu exporteras les résultats vers différents formats.

Chaque fiche de catalogue est stockée sous forme d'une seule chaîne au format `"Title | Author | Year | Location"`, avec des champs séparés par le caractère pipe (`|`).

Tu utiliseras un tableau de ces chaînes tout au long du workshop pour construire et tester tes fonctions de catalogue. Crée un tableau nommé `rawCatalogCards` et ajoute-lui les deux chaînes suivantes :

```js
"From a Buick 8 | King, Stephen | 2002 | Shelf K7",
"The Shining | King, Stephen | 1977 | Shelf K1"
```

# --hints--

Tu dois déclarer une variable nommée `rawCatalogCards`.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+rawCatalogCards\s*=/);
```

`rawCatalogCards` doit être un tableau.

```js
assert.isArray(rawCatalogCards);
```

Le premier élément de `rawCatalogCards` doit être `"From a Buick 8 | King, Stephen | 2002 | Shelf K7"`.

```js
assert.strictEqual(rawCatalogCards[0], "From a Buick 8 | King, Stephen | 2002 | Shelf K7");
```

Le deuxième élément de `rawCatalogCards` doit être `"The Shining | King, Stephen | 1977 | Shelf K1"`.

```js
assert.strictEqual(rawCatalogCards[1], "The Shining | King, Stephen | 1977 | Shelf K1");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
