---
id: 699d20623e57e8ffee885455
title: Étape 5
challengeType: 1
dashedName: step-5
---

# --description--

Comme tu peux le voir dans la sortie de la console, après le découpage, chaque partie a encore des espaces superflus au début et/ou à la fin. Pour nettoyer chaque élément, tu dois traiter chacun d'eux individuellement.

C'est un cas d'usage parfait pour une boucle `for`, qui répète un bloc de code une fois pour chaque élément d'un itérable, comme un tableau :

```js
for (let i = 0; i < array.length; i++) {
  // code to run on each iteration
}
```

`i` commence à `0` (le premier index), la boucle continue tant que `i` est inférieur à la longueur du tableau, et `i++` augmente `i` de un après chaque itération.

À l'intérieur de `parseCard`, déclare un tableau vide nommé `trimmedParts`. Puis ajoute une boucle `for` qui itère sur `parts`. Laisse le corps de la boucle vide pour l'instant, car tu le rempliras à l'étape suivante.

# --hints--

Tu dois déclarer une variable nommée `trimmedParts` définie sur un tableau vide.

```js
assert.match(parseCard.toString(), /(const|let|var)\s+trimmedParts\s*=\s*\[\s*\]/);
```

Tu dois utiliser une boucle `for` pour itérer sur `parts`.

```js
assert.match(parseCard.toString(), /for\s*\(\s*(var|let)\s+(\w+)\s*=\s*0\s*;\s*\2\s*<\s*parts\s*\.\s*length\s*;\s*\2\+\+\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
--fcc-editable-region--

--fcc-editable-region--
  return parts;
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
