---
id: 69aa11c2584523e6e991330c
title: Étape 6
challengeType: 1
dashedName: step-6
---

# --description--

La méthode `.trim()` supprime les espaces aux deux extrémités d'une chaîne. Par exemple, `" hello ".trim()` retourne `"hello"`.

À l'intérieur du corps de la boucle, ajoute `parts[i].trim()` dans `trimmedParts`. Puis mets à jour l'instruction `return` existante pour retourner `trimmedParts` au lieu de `parts`.

Après cette étape, la sortie affichée doit montrer des valeurs propres sans espaces superflus.

# --hints--

Tu dois ajouter `parts[i].trim()` dans `trimmedParts` à l'intérieur de la boucle.

```js
assert.match(parseCard.toString(), /trimmedParts\s*\.\s*push\s*\(\s*parts\s*\[\s*\w+\s*\]\s*\.\s*trim\s*\(\s*\)\s*\)/);
```

`parseCard` doit retourner `trimmedParts`.

```js
const explorer = await __helpers.Explorer(code);
const { parseCard } = explorer.allFunctions;
assert.isTrue(parseCard.hasReturn("trimmedParts"));
```

`parseCard` doit retourner un tableau où chaque élément n'a pas d'espaces au début ni à la fin.

```js
assert.isArray(parseCard("  A  |  B  |  C  |  D  "));
```

Lorsqu'elle est appelée avec `"  A  |  B  |  C  |  D  "`, le premier élément du tableau retourné doit être `"A"`.

```js
assert.strictEqual(parseCard("  A  |  B  |  C  |  D  ")[0], "A");
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
--fcc-editable-region--

  }
  return parts;
--fcc-editable-region--
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
