---
id: 5d822fd413a79914d39e98cd
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

En CSS, tu peux cibler tous les éléments avec un astérisque. Ajoute une bordure à tous les éléments en utilisant le sélecteur `*` et en lui donnant un `border` de `1px solid black`. C'est une astuce qui aide à visualiser où se trouvent les éléments et leur taille. Tu la retireras plus tard.

# --hints--

Tu devrais utiliser le sélecteur `*`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

Tu devrais utiliser la propriété `border` pour styliser tous les éléments.

```js
assert.exists(new __helpers.CSSHelp(document).isPropertyUsed('border'));
```

Tous les éléments devraient avoir une bordure `1px solid black`.

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles?.border, '1px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
