---
id: 5d822fd413a79914d39e98cc
title: Étape 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Bienvenue dans le projet Skyline avec variables CSS ! Commence par ajouter dans l'élément `head` un élément `link` qui relie ton fichier `styles.css`.

# --hints--

Ton code devrait avoir un élément `link`.

```js
assert.isNotNull(document.querySelector('link'));
```

Tu ne devrais pas modifier tes balises `head` existantes. Assure-toi de ne pas avoir supprimé la balise fermante.

```js
const headElements = document.querySelectorAll('head');
assert.strictEqual(headElements?.length, 1);
```

Tu devrais avoir un élément `link` vide.

```js
assert.lengthOf(document.querySelectorAll('link'), 1);
```

Ton élément `link` devrait être dans ton élément `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

Ton élément `link` devrait avoir un attribut `rel` avec la valeur `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Ton élément `link` devrait avoir un attribut `href` avec la valeur `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
