---
id: 5d822fd413a79914d39e98ce
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

Ajoute aussi un `box-sizing` de `border-box` à tous les éléments. Ainsi, la bordure que tu as ajoutée n'ajoutera aucune taille supplémentaire à tes éléments.

# --hints--

Tu devrais utiliser la propriété `box-sizing`.

```js
assert.exists(new __helpers.CSSHelp(document).isPropertyUsed('box-sizing'));
```

Tu devrais utiliser le sélecteur `*` existant.

```js
// Two selectors create two CSSStyleRule objects
assert.lengthOf(new __helpers.CSSHelp(document).getStyleDeclarations('*'), 1);
```

Tous les éléments devraient avoir un `box-sizing` de `border-box`.

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles.boxSizing, 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--
* {
  border: 1px solid black;
}

--fcc-editable-region--

```
