---
id: 5d822fd413a79914d39e98cf
title: Étape 4
challengeType: 0
dashedName: step-4
---

# --description--

Tu peux voir le `body` (c'est la boîte la plus interne de ta page) ; la boîte qui l'entoure est l'élément `html`. Fais remplir tout le viewport à ton `body` en lui donnant un `height` de `100vh`. Supprime la `margin` par défaut du `body` en definissant la `margin` sur `0`. Enfin, définis la propriété `overflow` sur `hidden` pour masquer les barres de défilement qui apparaissent quand quelque chose depasse le viewport.

# --hints--

Tu devrais utiliser le sélecteur `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

Ton `body` devrait avoir un `height` de `100vh`.

```js
const bodyStyles = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(bodyStyles?.height, '100vh');
```

Ton `body` devrait avoir un `margin` de `0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

Ton `body` devrait avoir la propriété `overflow` définie sur `hidden`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.overflow, 'hidden');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

--fcc-editable-region--


--fcc-editable-region--

```
