---
id: 5d822fd413a79914d39e98d4
title: Ătape 9
challengeType: 0
dashedName: step-9
---

# --description--

Donne aux parties de ton immeuble les propriĂŠtĂŠs `width` et `height` avec ces valeurs : `70%` et `10%` pour `.bb1a`, `80%` et `10%` pour `.bb1b`, `90%` et `10%` pour `.bb1c`, puis `100%` et `70%` pour `.bb1d`. Souviens-toi que ces pourcentages sont relatifs au parent, et remarque que les hauteurs s'additionnent pour faire 100% : elles remplissent donc verticalement le conteneur.

# --hints--

Tu devrais utiliser un sĂŠlecteur de classe pour styliser `.bb1a`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1a'));
```

Tu devrais donner Ă  `.bb1a` un `width` de `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.width, '70%');
```

Tu devrais donner Ă  `.bb1a` un `height` de `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.height, '10%');
```

Tu devrais utiliser un sĂŠlecteur de classe pour styliser `.bb1b`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1b'));
```

Tu devrais donner Ă  `.bb1b` un `width` de `80%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.width, '80%');
```

Tu devrais donner Ă  `.bb1b` un `height` de `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.height, '10%');
```

Tu devrais utiliser un sĂŠlecteur de classe pour styliser `.bb1c`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1c'));
```

Tu devrais donner Ă  `.bb1c` un `width` de `90%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.width, '90%');
```

Tu devrais donner Ă  `.bb1c` un `height` de `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.height, '10%');
```

Tu devrais utiliser un sĂŠlecteur de classe pour styliser `.bb1d`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1d'));
```

Tu devrais donner Ă  `.bb1d` un `width` de `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.width, '100%');
```

Tu devrais donner Ă  `.bb1d` un `height` de `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}
--fcc-editable-region--

--fcc-editable-region--

```
