---
id: 5d822fd413a79914d39e98d0
title: Étape 5
challengeType: 0
dashedName: step-5
---

# --description--

Crée dans le `body` un élément `div` avec une classe `background-buildings`. Ce sera un conteneur pour un groupe d'immeubles.

# --hints--

Tu devrais créer un élément `div`.

```js
assert.exists(document.querySelector('div'));
```

Ton élément `div` devrait être dans le `body`.

```js
assert.equal(document.querySelector('div')?.parentElement?.localName, 'body');
```

Ton élément `div` devrait avoir une classe `background-buildings`.

```js
assert.isTrue([...document.querySelector('div')?.classList]?.includes('background-buildings'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

```
