---
id: 5d822fd413a79914d39e98db
title: Étape 16
challengeType: 0
dashedName: step-16
---

# --description--

Donne aux nouveaux immeubles les propriétés `width` et `height` suivantes : `10%` et `50%` pour `.bb2`, `10%` et `55%` pour `.bb3`, puis `11%` et `58%` pour `.bb4`. Tu utiliseras presque uniquement des unites en pourcentage et un peu de flexbox pour ce projet, donc tout sera completement responsive.

# --hints--

Tu devrais donner à `.bb2` un `width` de `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb2')?.width, '10%');
```

Tu devrais donner à `.bb2` un `height` de `50%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb2')?.height, '50%');
```

Tu devrais donner à `.bb3` un `width` de `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb3')?.width, '10%');
```

Tu devrais donner à `.bb3` un `height` de `55%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb3')?.height, '55%');
```

Tu devrais donner à `.bb4` un `width` de `11%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb4')?.width, '11%');
```

Tu devrais donner à `.bb4` un `height` de `58%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb4')?.height, '58%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
      <div class="bb2"></div>
      <div class="bb3"></div>
      <div class="bb4"></div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
  --building-color1: #aa80ff;
}

.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}

--fcc-editable-region--

--fcc-editable-region--

```
