---
id: 5d822fd413a79914d39e98d7
title: Étape 12
challengeType: 0
dashedName: step-12
---

# --description--

Pour utiliser une variable, mets le nom de la variable entre parentheses avec `var` devant, comme ceci : `var(--variable-name)`. La valeur que tu as donnee à la variable sera appliquée à la propriété où tu l'utilises.

Ajoute la variable `--building-color1` que tu as creee à l'étape précédente comme valeur de la propriété `background-color` de la classe `.bb1a`.

# --hints--

Le `background-color` de l élément `.bb1a` devrait être défini.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1a')?.backgroundColor)
```

Tu devrais utiliser `var(--building-color1)` pour définir le `background-color` de l élément `.bb1a`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.getPropVal('background-color', true), 'var(--building-color1)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
  --building-color1: #999;
}
--fcc-editable-region--
.bb1a {
  width: 70%;
  height: 10%;
}
--fcc-editable-region--
.bb1b {
  width: 80%;
  height: 10%;
}

.bb1c {
  width: 90%;
  height: 10%;
}

.bb1d {
  width: 100%;
  height: 70%;
}

```
