---
id: 5d822fd413a79914d39e98d3
title: Étape 8
challengeType: 0
dashedName: step-8
---

# --description--

Imbrique quatre éléments `div` dans le conteneur `.bb1`. Donne-leur les classes `bb1a`, `bb1b`, `bb1c` et `bb1d`, dans cet ordre. Cet immeuble aura quatre sections.

# --hints--

Tu devrais créer quatre nouveaux éléments `div`.

```js
assert.equal(document.querySelectorAll('div')?.length, 6);
```

Tu devrais donner à l'un des nouveaux éléments `div` une classe `bb1a`.

```js
assert.exists(document.querySelector('div.bb1a'));
```

Tu devrais donner à l'un des nouveaux éléments `div` une classe `bb1b`.

```js
assert.exists(document.querySelector('div.bb1b'));
```

Tu devrais donner à l'un des nouveaux éléments `div` une classe `bb1c`.

```js
assert.exists(document.querySelector('div.bb1c'));
```

Tu devrais donner à l'un des nouveaux éléments `div` une classe `bb1d`.

```js
assert.exists(document.querySelector('div.bb1d'));
```

Tu devrais placer les nouveaux éléments `div` dans le bon ordre.

```js
function __t(a, b) {
  return [...document.querySelector(a)?.nextElementSibling?.classList]?.includes(b);
}
assert.isTrue(__t('div.bb1a','bb1b') && __t('div.bb1b','bb1c') && __t('div.bb1c','bb1d'));
```

Tu devrais placer les nouveaux éléments `div` dans l élément `.bb1`.

```js
assert.equal(document.querySelectorAll('div.bb1 > div')?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings">
--fcc-editable-region--
      <div class="bb1"></div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}

```
