---
id: 5d822fd413a79914d39e98d1
title: Étape 6
challengeType: 0
dashedName: step-6
---

# --description--

Donne à ton élément `.background-buildings` un `width` et un `height` de `100%` pour qu'il prenne toute la largeur et toute la hauteur de son parent, le `body`.

# --hints--

Tu devrais utiliser la classe `background-buildings` pour selectionner le bon élément.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.background-buildings'));
```

Ton élément `.background-buildings` devrait avoir un `width` de `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.width, '100%');
```

Ton élément `.background-buildings` devrait avoir un `height` de `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.height, '100%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings"></div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}
--fcc-editable-region--


--fcc-editable-region--

```
