---
id: 5d822fd413a79914d39e98dd
title: Étape 18
challengeType: 0
dashedName: step-18
---

# --description--

Les immeubles sont trop espaces. Resserre-les en ajoutant deux éléments `div` vides au début de l'élément `.background-buildings`, deux autres à la fin, et un autre entre `.bb3` et `.bb4`. Ils seront ajoutes comme éléments régulièrement espaces dans le conteneur, ce qui rapprochera effectivement les immeubles du centre.

# --hints--

Tu devrais ajouter deux nouveaux éléments `div` avant l élément `.bb1`.

```js
const bBuildings = document.querySelector('.background-buildings')?.children;
assert.isTrue(![...bBuildings?.[0]?.classList]?.includes('bb1'));
assert.isTrue(![...bBuildings?.[1]?.classList]?.includes('bb1'));
```

Tu devrais ajouter un nouvel élément `div` entre les éléments `.bb3` et `.bb4`.

```js
assert.equal(document.querySelector('.bb3')?.nextElementSibling, document.querySelector('.bb4')?.previousElementSibling);
```

Tu devrais ajouter deux nouveaux éléments `div` après l élément `.bb4`.

```js
const bb4 = document.querySelector('.bb4');
assert.exists(bb4?.nextElementSibling);
assert.exists(bb4?.nextElementSibling?.nextElementSibling);
```

Tu devrais ajouter 5 nouveaux éléments `div`.

```js
assert.lengthOf(document.querySelectorAll('div'), 14);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
--fcc-editable-region--
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
      <div class="bb2"></div>
      <div class="bb3"></div>
      <div class="bb4"></div>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
  display: flex;
  align-items: flex-end;
  justify-content: space-evenly;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
  --building-color1: #aa80ff;
}

.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}

.bb2 {
  width: 10%;
  height: 50%;
}

.bb3 {
  width: 10%;
  height: 55%;
}

.bb4 {
  width: 11%;
  height: 58%;
}

```
