---
id: 5d822fd413a79914d39e98d6
title: Étape 11
challengeType: 0
dashedName: step-11
---

# --description--

Tu as maintenant quelque chose qui ressemble à un immeuble. Tu es prêt à créer ta première variable. Dans les leçons précédentes, tu as appris que les déclarations de variables commencent par deux tirets (`-`) et reçoivent un nom et une valeur comme ceci : `--variable-name: value;`. Dans la règle de la classe `bb1`, crée une variable nommée `--building-color1` et donne-lui la valeur `#999`.

# --hints--

Tu devrais créer une nouvelle variable nommée `--building-color1`.

```js
assert.exists(new __helpers.CSSHelp(document).isPropertyUsed('--building-color1'));
```

Tu devrais définir la variable `--building-color1` dans `.bb1`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color1'));
```

Tu devrais donner à `--building-color1` une valeur de `#999`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color1').trim(), '#999');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}
--fcc-editable-region--
.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
}
--fcc-editable-region--
.bb1a {
  width: 70%;
  height: 10%;
}

.bb1b {
  width: 80%;
  height: 10%;
}

.bb1c {
  width: 90%;
  height: 10%;
}

.bb1d {
  width: 100%;
  height: 70%;
}

```
