---
id: 5d822fd413a79914d39e98d2
title: Étape 7
challengeType: 0
dashedName: step-7
---

# --description--

Imbrique un `div` avec une classe `bb1` dans le conteneur des immeubles d'arrière-plan. Ouvre ton fichier `styles.css`, puis donne à `.bb1` un `width` de `10%` et un `height` de `70%`. "bb" signifie "background building" : ce sera ton premier immeuble.

# --hints--

Tu devrais créer un nouvel élément `div`.

```js
assert.equal(document.querySelectorAll('div').length, 2);
```

Tu devrais donner au nouveau `div` une classe `bb1`.

```js
assert.exists(document.querySelector('div.bb1'));
```

Tu devrais utiliser un sélecteur de classe `.bb1` pour styliser l élément.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1'));
```

Tu devrais donner à l élément `.bb1` un `width` de `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.width, '10%');
```

Tu devrais donner à l élément `.bb1` un `height` de `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
--fcc-editable-region--
    <div class="background-buildings"></div>
--fcc-editable-region--
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

```
