---
id: 6718d2d59337c822ecb697ff
title: Construire un programme de gestion de compte bancaire
challengeType: 26
dashedName: build-a-bank-account-management-program
saveSubmissionToDB: true
---

# --description--

**Objectif :** Remplir les user stories ci-dessous et faire passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu dois définir une classe nommée `BankAccount` avec un constructeur. Le constructeur doit définir le `balance` par défaut à `0` et initialiser un tableau vide nommé `transactions` pour stocker les enregistrements de transactions sous forme d'objets.

2. Chaque transaction stockée dans le tableau `transactions` doit être un objet avec deux propriétés : `type` et `amount`. La propriété `type` doit être soit `deposit` soit `withdraw`, et la propriété `amount` doit être le montant déposé ou retiré.

3. Tu dois définir une méthode nommée `deposit` qui prend le montant du dépôt comme paramètre. Quand le montant du dépôt est supérieur à `0`, elle doit :

   - Ajouter un nouvel objet au tableau des transactions avec un type `deposit` et le montant déposé.
   - Mettre à jour le solde.
   - Renvoyer `"Successfully deposited $[amount]. New balance: $[balance]"`.

4. Si le montant est inférieur ou égal à `0`, la méthode `deposit` doit renvoyer `"Deposit amount must be greater than zero."`.

5. Tu dois définir une méthode nommée `withdraw` qui prend un montant comme paramètre. Cette méthode doit mettre à jour le solde courant selon les retraits. Quand le montant à retirer est supérieur à `0` et inférieur ou égal au solde courant, elle doit :

   - Ajouter un nouvel objet au tableau des transactions avec un type `withdraw` et le montant retiré.
   - Mettre à jour le solde.
   - Renvoyer `"Successfully withdrew $[amount]. New balance: $[balance]"`.

6. Si le montant à retirer est inférieur ou égal à `0` ou supérieur au solde courant, la méthode `withdraw` doit renvoyer `"Insufficient balance or invalid amount."`.


7. Tu dois définir une méthode nommée `checkBalance` qui renvoie le solde courant au format `"Current balance: $[balance]"`.

8. Tu dois définir une méthode nommée `listAllDeposits` qui parcourt le tableau des transactions et renvoie tous les dépôts au format `"Deposits: amount,amount,..."`.

9. Tu dois définir une méthode nommée `listAllWithdrawals` qui parcourt le tableau des transactions et renvoie tous les retraits au format `"Withdrawals: amount,amount,..."`.

10. Tu dois créer une nouvelle instance de `BankAccount` nommée `myAccount`.

11. Ton compte bancaire `myAccount` doit avoir au moins cinq transactions.

12. Ton compte bancaire `myAccount` doit avoir au moins deux dépôts.

13. Ton compte bancaire `myAccount` doit avoir au moins deux retraits.

14. Ton compte bancaire `myAccount` doit avoir un solde supérieur à `$100`.

# --hints--

Tu dois définir une classe nommée `BankAccount`.

```js
assert.isFunction(BankAccount);
assert.isObject(new BankAccount());
```

L'objet `BankAccount` doit initialement avoir un `balance` de `0` et un tableau vide `transactions` pour stocker les enregistrements de transactions.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.balance, 0);
assert.equal(myAccount.transactions.length, 0);
```

Tu dois avoir une méthode `deposit` qui prend le montant du dépôt comme paramètre.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.deposit);
assert.lengthOf(myAccount.deposit, 1);
```

Tu dois avoir une méthode `withdraw` qui prend le montant du retrait comme paramètre.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.withdraw);
assert.lengthOf(myAccount.withdraw, 1);
```

Tu dois avoir une méthode `checkBalance` qui vérifie le solde courant.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.checkBalance);
```

Tu dois avoir une méthode `listAllDeposits` qui liste tous les dépôts.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllDeposits);
```

Tu dois avoir une méthode `listAllWithdrawals` qui liste tous les retraits.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllWithdrawals);
```

`BankAccount.deposit(100)` devrait renvoyer `"Successfully deposited $100. New balance: $100"`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(100), "Successfully deposited $100. New balance: $100")

// prevent hardcoding
assert.equal(myAccount.deposit(200), "Successfully deposited $200. New balance: $300")

```

`BankAccount.deposit(-50)` devrait renvoyer `"Deposit amount must be greater than zero."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(-50), "Deposit amount must be greater than zero.")

// prevent hardcoding
assert.equal(myAccount.deposit(-30), "Deposit amount must be greater than zero.")

```

`BankAccount.deposit(0)` devrait renvoyer `"Deposit amount must be greater than zero."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(0), "Deposit amount must be greater than zero.")
```

Quand le solde du compte est `100`, `BankAccount.withdraw(150)` devrait renvoyer `"Insufficient balance or invalid amount."`.

```js
const myAccount = new BankAccount();
myAccount.deposit(100);
assert.equal(myAccount.withdraw(150), "Insufficient balance or invalid amount.")

// prevent hardcoding
assert.equal(myAccount.withdraw(250), "Insufficient balance or invalid amount.")

```

`BankAccount.withdraw(-50)` devrait renvoyer `"Insufficient balance or invalid amount."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.withdraw(-50), "Insufficient balance or invalid amount.")

// prevent hardcoding
assert.equal(myAccount.withdraw(-30), "Insufficient balance or invalid amount.")
```

`BankAccount.withdraw(0)` devrait renvoyer `"Insufficient balance or invalid amount."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.withdraw(0), "Insufficient balance or invalid amount.")
```

Quand le solde du compte est `200`, `BankAccount.withdraw(150)` devrait renvoyer `"Successfully withdrew $150. New balance: $50"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(200);
assert.equal(myAccount.withdraw(150), "Successfully withdrew $150. New balance: $50")
```

Quand le solde du compte est `200`, `BankAccount.checkBalance()` devrait renvoyer `"Current balance: $200"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(200);
assert.equal(myAccount.checkBalance(), "Current balance: $200")

// prevent hardcoding
myAccount.deposit(300);
assert.equal(myAccount.checkBalance(), "Current balance: $500")
```

Quand tu déposes `10`, `35`, `90`, la méthode `listAllDeposits` devrait renvoyer `"Deposits: 10,35,90"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(10);
myAccount.deposit(35);
myAccount.deposit(90);
assert.equal(myAccount.listAllDeposits(), "Deposits: 10,35,90")
```

Quand tu retires `20`, `50`, `100`, la méthode `listAllWithdrawals` devrait renvoyer `"Withdrawals: 20,50,100"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(500);

myAccount.withdraw(20);
myAccount.withdraw(50);
myAccount.withdraw(100);
assert.equal(myAccount.listAllWithdrawals(), "Withdrawals: 20,50,100")
```

Tu dois avoir une instance de `BankAccount` nommée `myAccount`.

```js
assert.exists(myAccount);
assert.instanceOf(myAccount, BankAccount);
```

Ton compte bancaire `myAccount` doit avoir au moins cinq transactions.

```js
assert.isAtLeast(myAccount.transactions.length, 5);
```

Ton compte bancaire `myAccount` doit avoir au moins deux dépôts.

```js
const depositCount = myAccount.transactions.filter(transaction => transaction.type === "deposit").length;
assert.isAtLeast(depositCount, 2);
```

Ton compte bancaire `myAccount` doit avoir au moins deux retraits.

```js
const withdrawalCount = myAccount.transactions.filter(transaction => transaction.type === "withdraw").length;
assert.isAtLeast(withdrawalCount, 2);
```


Ton compte bancaire `myAccount` doit avoir un solde supérieur à `$100`.

```js
assert.isAbove(myAccount.balance, 100);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
class BankAccount {
    constructor() {
        this.balance = 0;
        this.transactions = [];
    }

    deposit(amount) {
        if (amount > 0) {
            this.transactions.push({ type: 'deposit', amount });
            this.balance += amount;
            return `Successfully deposited $${amount}. New balance: $${this.balance}`;
        } else {
            return 'Deposit amount must be greater than zero.';
        }
    }

    withdraw(amount) {
        if (amount > 0 && amount <= this.balance) {
            this.transactions.push({ type: 'withdraw', amount });
            this.balance -= amount;
            return `Successfully withdrew $${amount}. New balance: $${this.balance}`;
        } else {
            return 'Insufficient balance or invalid amount.';
        }
    }

    checkBalance() {
        return `Current balance: $${this.balance}`;
    }

    listAllDeposits() {
        const deposits = this.transactions
            .filter(transaction => transaction.type === 'deposit')
            .map(transaction => transaction.amount);
        return `Deposits: ${deposits}`;
    }

    listAllWithdrawals() {
        const withdrawals = this.transactions
            .filter(transaction => transaction.type === 'withdraw')
            .map(transaction => transaction.amount);
        return `Withdrawals: ${withdrawals}`;
    }
}

const myAccount = new BankAccount();

console.log(myAccount.deposit(0));
console.log(myAccount.deposit(20));
console.log(myAccount.deposit(500));
console.log(myAccount.withdraw(30));
console.log(myAccount.withdraw(55));
console.log(myAccount.withdraw(100));
console.log(myAccount.checkBalance());
console.log(myAccount.listAllDeposits());
console.log(myAccount.listAllWithdrawals());
```
