---
id: 695774002591bbc5f8cf3e53
title: Crée une fonction d'application de remise
challengeType: 27
dashedName: build-a-discount-calculator
---

# --description--

Dans ce laboratoire, tu vas écrire une fonction qui calcule le prix final d'un article après application d'une remise en pourcentage.

Par exemple, si le prix d'un article est `50` et qu'une remise de `20` est appliquée, le montant de la remise est `10`, et le prix final est `40`.

**Objectif :** remplis les user stories ci-dessous pour terminer le lab.

**User Stories :**

1. Tu devrais définir une fonction nommée `apply_discount`.

2. La fonction `apply_discount` devrait prendre exactement deux paramètres : `price` et `discount`.

3. Si `price` n'est pas un nombre (`int` ou `float`), la fonction devrait renvoyer la chaîne `The price should be a number`.

4. Si `discount` n'est pas un nombre (`int` ou `float`), la fonction devrait renvoyer la chaîne `The discount should be a number`.

5. Si `price` est inférieur ou égal à `0`, la fonction devrait renvoyer la chaîne `The price should be greater than 0`.

6. Si `discount` est inférieur à `0` ou supérieur à `100`, la fonction devrait renvoyer la chaîne `The discount should be between 0 and 100`.

7. Si les deux entrées sont valides, la fonction devrait calculer la remise comme un pourcentage du prix.

8. La fonction devrait renvoyer le prix final après application de la remise.

# --hints--

Tu devrais avoir une fonction nommée `apply_discount`.

```js
({
  test: () => runPython(
    `assert _Node(_code).has_function("apply_discount")`
  )
})
```

Ta fonction `apply_discount` devrait prendre deux paramètres : `price` et `discount`.

```js
({
  test: () => runPython(
`
func = _Node(_code).find_function("apply_discount")

assert func.has_args("price, discount")
`
  )
})
```

Quand `apply_discount` est appelée avec un `price` (premier argument) qui n'est pas un nombre (`int` ou `float`), elle devrait renvoyer `The price should be a number`.

```js
({
  test: () => runPython(
`
wrong_values = ["HelloWorld", "100", [1,2], {"a":1}, (1,2), None]
for v in wrong_values:
    assert apply_discount(v, 10) == "The price should be a number"
`
  )
})
```

Quand `apply_discount` est appelée avec un `discount` (deuxième argument) qui n'est pas un nombre (`int` ou `float`), elle devrait renvoyer `The discount should be a number`.

```js
({
  test: () => runPython(
`
wrong_values = ["HelloWorld", "10", [1], {"a":1}, (1,), None]
for v in wrong_values:
    assert apply_discount(100, v) == "The discount should be a number"
`
  )
})
```

Quand `apply_discount` est appelée avec un `price` inférieur ou égal à `0`, elle devrait renvoyer `The price should be greater than 0`.

```js
({
  test: () => runPython(
`
assert apply_discount(0, 10) == "The price should be greater than 0"
assert apply_discount(-5, 10) == "The price should be greater than 0"
`
  )
})
```

Quand `apply_discount` est appelée avec un `discount` inférieur à `0` ou supérieur à `100`, elle devrait renvoyer `The discount should be between 0 and 100`.

```js
({
  test: () => runPython(
`
assert apply_discount(100, -1) == "The discount should be between 0 and 100"
assert apply_discount(100, 101) == "The discount should be between 0 and 100"
`
  )
})
```

`apply_discount(100, 20)` devrait renvoyer `80`.

```js
({
  test: () => runPython(
`
assert apply_discount(100, 20) == 80
`
  )
})
```

`apply_discount(200, 50)` devrait renvoyer `100`.

```js
({
  test: () => runPython(
`
assert apply_discount(200, 50) == 100
`
  )
})
```

`apply_discount(50, 0)` devrait renvoyer `50`.

```js
({
  test: () => runPython(
`
assert apply_discount(50, 0) == 50
`
  )
})
```

Quand `apply_discount` est appelée avec une remise de `100`, elle devrait renvoyer `0`.

```js
({
  test: () => runPython(
`
assert apply_discount(120, 100) == 0
`
  )
})
```

`apply_discount(74.5, 20.0)` devrait renvoyer `59.6`.

```js
({
  test: () => runPython(
`
result = apply_discount(74.5, 20.0)
assert result == 59.6
`
  )
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def apply_discount(price, discount):
    # Validation steps
    if not isinstance(price, (int, float)):
        return "The price should be a number"
    if not isinstance(discount, (int, float)):
        return "The discount should be a number"
    if price <= 0:
        return "The price should be greater than 0"
    if discount < 0 or discount > 100:
        return "The discount should be between 0 and 100"

    discount_amount = price * discount / 100
    final_price = price - discount_amount
    return final_price
```
