---
id: 68f890b73bbfc9efa7379e7b
title: Déboguer des erreurs de coercition de type dans une application boguée
challengeType: 26
dashedName: debug-type-coercion-errors-in-a-buggy-app
---

# --description--

Tu viens de rejoindre une agence locale de développement web, et ta première mission est de nettoyer du code bogué laissé par les développeurs précédents.

Ils tentaient d'effectuer de simples opérations arithmétiques, mais quelque chose a mal tourné : les résultats n'ont pas de sens. Ta tâche est de relire, déboguer et corriger le code pour qu'il fonctionne comme prévu et soit plus facile à lire.

**Objectif :** Réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu devrais avoir une variable nommée `firstResult` qui additionne correctement les nombres `5` et `10` pour produire la valeur `15`.
2. Tu devrais avoir une variable nommée `secondResult` qui soustrait correctement `5` de `8` pour produire la valeur `3`.
3. Tu devrais avoir une variable nommée `thirdResult` qui produit la valeur `6` en additionnant deux nombres. Remplace le booléen actuellement utilisé dans l'expression par un nombre.
4. Tu devrais avoir une variable nommée `fourthResult` qui produit la valeur `8` en additionnant deux nombres. Remplace le booléen actuellement utilisé dans l'expression par un nombre.
5. Tu devrais avoir une variable nommée `fifthResult` qui multiplie correctement deux nombres pour obtenir le produit `20`. Remplace la chaîne actuellement utilisée dans l'expression par un nombre.
6. Tu devrais avoir une variable nommée `sixthResult` qui additionne correctement deux nombres pour produire la valeur `22`. Remplace le `null` actuellement utilisé dans l'expression par un nombre.

# --hints--

Tu devrais avoir une variable appelée `firstResult`.

```js
assert.match(code, /const\s+firstResult\s*=/);
```

`firstResult` devrait avoir la valeur numérique `15`.

```js
assert.strictEqual(firstResult, 15);
```

Tu devrais avoir une variable appelée `secondResult`.

```js
assert.match(code, /const\s+secondResult\s*=/);
```

`secondResult` devrait avoir la valeur numérique `3`.

```js
assert.strictEqual(secondResult, 3);
```

Tu devrais avoir une variable appelée `thirdResult`.

```js
assert.match(code, /const\s+thirdResult\s*=/);
```

`thirdResult` devrait avoir la valeur numérique `6`.

```js
assert.strictEqual(thirdResult, 6);
```

`thirdResult` ne devrait pas inclure de valeur booléenne (`true` ou `false`).

```js
const match = code.match(/const\s+thirdResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\btrue\b|\bfalse\b/);
```

Tu devrais avoir une variable appelée `fourthResult`.

```js
assert.match(code, /const\s+fourthResult\s*=/);
```

`fourthResult` devrait avoir la valeur numérique `8`.

```js
assert.strictEqual(fourthResult, 8);
```

`fourthResult` ne devrait pas inclure de valeur booléenne (`true` ou `false`).

```js
const match = code.match(/const\s+fourthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\btrue\b|\bfalse\b/);
```

Tu devrais avoir une variable appelée `fifthResult`.

```js
assert.match(code, /const\s+fifthResult\s*=/);
```

`fifthResult` devrait avoir la valeur numérique `20`.

```js
assert.strictEqual(fifthResult, 20);
```

`fifthResult` ne devrait pas inclure de multiplication de chaîne (par ex. "10" * 2).

```js
const match = code.match(/const\s+fifthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /["'`].*["'`]\s*\*/);
```

Tu devrais avoir une variable appelée `sixthResult`.

```js
assert.match(code, /const\s+sixthResult\s*=/);
```

`sixthResult` devrait avoir la valeur numérique `22`.

```js
assert.strictEqual(sixthResult, 22);
```

`sixthResult` ne devrait pas inclure `null` dans son expression.

```js
const match = code.match(/const\s+sixthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\bnull\b/);
```

# --seed--

## --seed-contents--

```js
const firstResult = 5 + "10";
console.log(`5 + 10 = ${firstResult}`);

const secondResult = "Eight" - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = true + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = false + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = "10" * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = null + 22;
console.log(`0 + 22 = ${sixthResult}`);
```

# --solutions--

```js
const firstResult = 5 + 10;
console.log(`5 + 10 = ${firstResult}`);

const secondResult = 8 - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = 1 + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = 0 + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = 10 * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = 0 + 22;
console.log(`0 + 22 = ${sixthResult}`);
```
