---
id: 66a207974c806a19d6607073
title: Construire une application d'inventaire de livres
challengeType: 25
dashedName: build-a-book-inventory-app
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**Objectif :** Réponds aux user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu dois avoir un élément `h1` avec le texte `Book Inventory`.
1. Tu dois avoir un élément `table` avec des colonnes nommées `Title`, `Author`, `Category`, `Status` et `Rate`.
1. Ton tableau doit avoir au moins quatre lignes, la première pour les en-têtes de colonnes et les autres remplies avec des informations sur des livres.
1. Chaque ligne de tableau à l'intérieur du corps du tableau doit avoir soit la classe `read`, `to-read` ou `in-progress`.
1. Les éléments `td` de la colonne `Status` doivent contenir un élément `span` avec la `class` `status` autour du texte `Read`, `To Read` ou `In Progress`, selon la classe de cette ligne.
1. Les éléments `td` de la colonne `Rate` doivent contenir un élément `span` avec la `class` `rate` qui enveloppe trois éléments `span` vides.
1. Les éléments `.rate` placés dans les lignes `.read` doivent avoir une classe supplémentaire avec la valeur `one`, `two` ou `three`, selon la note personnelle. Cette valeur doit venir après `rate`.
1. Tu dois créer trois sélecteurs d'attribut pour cibler les éléments avec la classe `read`, `to-read` et `in-progress`, puis définir leur propriété `background-image` pour utiliser un `linear-gradient` de ton choix.
1. Tu dois définir la propriété `display` de chaque élément `span` sur `inline-block`.
1. Tu dois utiliser un sélecteur d'attribut pour cibler les éléments `span` avec la classe `status` qui sont des descendants d'éléments `tr` avec la classe `to-read`, puis définir leurs propriétés `border` et `background-image`.
1. Tu dois utiliser un sélecteur d'attribut pour cibler les éléments `span` avec la classe `status` qui sont des descendants d'éléments `tr` avec la classe `read`, puis définir leurs propriétés `border` et `background-image`.
1. Tu dois utiliser un sélecteur d'attribut pour cibler les éléments `span` avec la classe `status` qui sont des descendants d'éléments `tr` avec la classe `in-progress`, puis définir leurs propriétés `border` et `background-image`.
1. Tu dois utiliser un sélecteur d'attribut pour cibler les éléments `span` avec la classe `status` et les éléments `span` dont la valeur de classe commence par `rate`, puis définir leurs propriétés `height`, `width` et `padding`.
1. Tu dois utiliser un sélecteur d'attribut pour cibler les éléments `span` qui sont des enfants directs d'éléments `span` dont la valeur de `class` commence par `rate`, puis définir leurs propriétés `border`, `border-radius`, `margin`, `height`, `width` et `background-color`.
1. Tu dois utiliser un sélecteur d'attribut pour cibler le premier descendant des éléments `span` qui ont le mot `one` comme partie de leur valeur de `class`, puis définir sa propriété `background-image` pour utiliser un `linear-gradient`.
1. Tu dois utiliser un sélecteur d'attribut pour cibler les deux premiers descendants des éléments `span` qui ont le mot `two` comme partie de leur valeur de `class`, puis définir leur propriété `background-image` pour utiliser un `linear-gradient`.
1. Tu dois utiliser un sélecteur d'attribut pour cibler les trois éléments `span` qui sont des descendants d'éléments `span` qui ont le mot `three` comme partie de leur valeur de `class`, puis définir leur propriété `background-image` pour utiliser un `linear-gradient`.

# --hints--

Tu dois avoir un élément `h1` avec le texte `Book Inventory`.

```js
assert.equal(document.querySelector('h1')?.innerText.trim(), 'Book Inventory');
```

Tu dois avoir un seul élément `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Tu dois avoir un élément `table`.

```js
assert.lengthOf(document.querySelectorAll('table'), 1);
```

Tu dois avoir un élément `thead` et un élément `tbody` dans `table`.

```js
assert.lengthOf(document.querySelectorAll("table > thead"), 1);
assert.lengthOf(document.querySelectorAll("table > tbody"), 1);
```

Dans le `thead`, il doit y avoir un `tr` avec 5 éléments `th`.

```js
assert.lengthOf(document.querySelectorAll("table > thead > tr"), 1);
assert.lengthOf(document.querySelectorAll("table > thead > tr > th"), 5);
```

Ta première colonne doit avoir le texte `Title` comme en-tête.

```js
assert.equal(document.querySelectorAll('th')[0]?.textContent.trim(), 'Title');
```

Ta deuxième colonne doit avoir le texte `Author` comme en-tête.

```js
assert.equal(document.querySelectorAll('th')[1]?.textContent.trim(), 'Author');
```

Ta troisième colonne doit avoir le texte `Category` comme en-tête.

```js
assert.equal(document.querySelectorAll('th')[2]?.textContent.trim(), 'Category');
```

Ta quatrième colonne doit avoir le texte `Status` comme en-tête.

```js
assert.equal(document.querySelectorAll('th')[3]?.textContent.trim(), 'Status');
```

Ta cinquième colonne doit avoir le texte `Rate` comme en-tête.

```js
assert.equal(document.querySelectorAll('th')[4]?.textContent.trim(), 'Rate');
```

Ton tableau doit avoir au moins quatre lignes.

```js
const rows = document.querySelectorAll('tr');
assert.isAtLeast(rows.length, 4);
```

Chaque ligne doit toujours avoir 5 colonnes.

```js
const rows = document.querySelectorAll('tbody tr');
assert.isNotEmpty(rows);
for (const row of rows) {
  assert.lengthOf(row.children, 5);
}
```

Chaque ligne de tableau, sauf la ligne d'en-tête, doit avoir soit la classe `read`, `to-read` ou `in-progress`.

```js
const rows = document.querySelectorAll('tr');
assert.isNotEmpty(rows);
for (let i = 1; i < rows.length; i++) {
  const classList = [...rows[i].classList];
  const currentProgress = classList[0]; 
  assert.oneOf(currentProgress, ['read', 'to-read', "in-progress"]); 
}
```

Les éléments `td` de la colonne `Status` doivent contenir un élément `span`.

```js
const statusColumnData = document.querySelectorAll('td:nth-child(4)');
assert.isNotEmpty(statusColumnData);
for (let e of statusColumnData) {
  assert.equal(e?.children[0]?.tagName, 'SPAN');
}
```

Chaque élément `span` de la colonne `Status` doit avoir la classe `status`.

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
assert.isNotEmpty(statusSpans);
for (let e of statusSpans) {
    assert.isTrue(e.classList.contains('status'));
}
```

Chaque élément `.status` doit avoir le texte `Read`, `To Read` ou `In Progress`, selon la classe de sa ligne.

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
const rows = Array.from(document.querySelectorAll('tr')).slice(1);
assert.isNotEmpty(statusSpans);
for (let i = 0; i < rows.length; i++) {
  switch (statusSpans[i]?.innerText.trim()) {
        case 'Read':
            assert.isTrue(rows[i].classList.contains('read'));
            break;
        case 'To Read':
            assert.isTrue(rows[i].classList.contains('to-read'));
            break;
        case 'In Progress':
            assert.isTrue(rows[i].classList.contains('in-progress'));
            break;
        default:
            assert.fail(); 
    }
}
```

Les éléments `td` de la colonne `Rate` doivent contenir un élément `span`.

```js
const rateColumnData = document.querySelectorAll('tr td:last-child');
assert.isNotEmpty(rateColumnData);
for (let e of rateColumnData) {
    assert.equal(e.children[0]?.tagName, 'SPAN')
}
```

Chaque élément `span` qui est un enfant direct d'un élément `td` de la colonne `Rate` doit avoir la classe `rate` comme première classe.

```js
const rateSpans = document.querySelectorAll('tr td:last-child > span:first-child');
assert.isNotEmpty(rateSpans);
for (let e of rateSpans) {
    assert.equal(e.classList[0], 'rate');
}
```

Chaque élément `.rate` doit contenir trois éléments `span` vides.

```js
const rateSpans = document.getElementsByClassName('rate');
assert.isNotEmpty(rateSpans);
for (let e of rateSpans) {
    assert.equal(e.children.length, 3);
    for (let child of e.children) {
        assert.equal(child.tagName, 'SPAN');
        assert.equal(child.innerText.length, 0);
    }
}
```

Les éléments `.rate` placés dans les lignes `.read` doivent avoir une classe supplémentaire après la classe `rate`, avec la valeur `one`, `two` ou `three`, selon la note personnelle.

```js
const readBooksRates = document.querySelectorAll('.read .rate');
assert.isNotEmpty(readBooksRates);
for (let e of readBooksRates) {
    assert.oneOf(e.classList[1], ['one', 'two', 'three']);
}
```

Tu dois avoir un sélecteur d'attribut pour cibler les lignes qui ont la classe `read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"]'));
```

Tu dois utiliser un sélecteur d'attribut pour cibler les lignes qui ont la classe `read` et définir leur propriété `background-image` sur un dégradé linéaire de ton choix.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="read"]')?.backgroundImage, 'linear-gradient(');
```

Tu dois avoir un sélecteur d'attribut pour cibler les lignes qui ont la classe `to-read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]'));
```

Tu dois utiliser un sélecteur d'attribut pour cibler les lignes qui ont la classe `to-read` et définir leur propriété `background-image` sur un dégradé linéaire de ton choix.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]')?.backgroundImage, 'linear-gradient(');
```

Tu dois avoir un sélecteur d'attribut pour cibler les lignes qui ont la classe `in-progress`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]'));
```

Tu dois utiliser un sélecteur d'attribut pour cibler les lignes qui ont la classe `in-progress` et définir leur propriété `background-image` sur un dégradé linéaire de ton choix.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]')?.backgroundImage, 'linear-gradient(');
```

Tu dois définir la propriété `display` de chaque élément `span` sur `inline-block`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('span')?.getPropVal('display'), 'inline-block');
```

Tu dois avoir un sélecteur d'attribut pour cibler les éléments `span` avec la classe `status` qui sont des descendants d'éléments `tr` avec la classe `to-read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]'));
```

Tu dois utiliser un sélecteur d'attribut pour cibler les éléments `span` avec la classe `status` qui sont des descendants d'éléments `tr` avec la classe `to-read` et définir leur propriété `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.border);
```

Tu dois utiliser un sélecteur d'attribut pour cibler les éléments `span` avec la classe `status` qui sont des descendants d'éléments `tr` avec la classe `to-read` et définir leur propriété `background-image`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.backgroundImage);
```

Tu dois avoir un sélecteur d'attribut pour cibler les éléments `span` avec la classe `status` qui sont des descendants d'éléments `tr` avec la classe `read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]'));
```

Tu dois utiliser un sélecteur d'attribut pour cibler les éléments `span` avec la classe `status` qui sont des descendants d'éléments `tr` avec la classe `read` et définir leur propriété `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.border);
```

Tu dois utiliser un sélecteur d'attribut pour cibler les éléments `span` avec la classe `status` qui sont des descendants d'éléments `tr` avec la classe `read` et définir leur propriété `background-image`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.backgroundImage);
```

Tu dois avoir un sélecteur d'attribut pour cibler les éléments `span` avec la classe `status` qui sont des descendants d'éléments `tr` avec la classe `in-progress`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]'));
```

Tu dois utiliser un sélecteur d'attribut pour cibler les éléments `span` avec la classe `status` qui sont des descendants d'éléments `tr` avec la classe `in-progress` et définir leur propriété `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.border);
```

Tu dois utiliser un sélecteur d'attribut pour cibler les éléments `span` avec la classe `status` qui sont des descendants d'éléments `tr` avec la classe `in-progress` et définir leur propriété `background-image`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.backgroundImage);
```

Tu dois avoir un sélecteur d'attribut pour cibler les éléments `span` avec la classe `status` et les éléments `span` dont la valeur de classe commence par `rate`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.exists(selector1 || selector2);
```

Tu dois utiliser un sélecteur d'attribut pour cibler les éléments `span` avec la classe `status` et les éléments `span` dont la valeur de classe commence par `rate`, puis définir leur propriété `height`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.height || selector2?.height);
```

Tu dois utiliser un sélecteur d'attribut pour cibler les éléments `span` avec la classe `status` et les éléments `span` dont la valeur de classe commence par `rate`, puis définir leur propriété `width`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.width || selector2?.width);
```

Tu dois utiliser un sélecteur d'attribut pour cibler les éléments `span` avec la classe `status` et les éléments `span` dont la valeur de classe commence par `rate`, puis définir leur propriété `padding`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.padding || selector2?.padding);
```

Tu dois avoir un sélecteur d'attribut pour cibler les éléments `span` qui sont des enfants directs d'éléments `span` dont la valeur de `class` commence par `rate`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span'));
```

Tu dois utiliser un sélecteur d'attribut pour cibler les éléments `span` qui sont des enfants directs d'éléments `span` dont la valeur de `class` commence par `rate` et définir leur propriété `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border'));
```

Tu dois utiliser un sélecteur d'attribut pour cibler les éléments `span` qui sont des enfants directs d'éléments `span` dont la valeur de `class` commence par `rate` et définir leur propriété `border-radius`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border-radius'));
```

Tu dois utiliser un sélecteur d'attribut pour cibler les éléments `span` qui sont des enfants directs d'éléments `span` dont la valeur de `class` commence par `rate` et définir leur propriété `margin`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('margin'));
```

Tu dois utiliser un sélecteur d'attribut pour cibler les éléments `span` qui sont des enfants directs d'éléments `span` dont la valeur de `class` commence par `rate` et définir leur propriété `height`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('height'));
```

Tu dois utiliser un sélecteur d'attribut pour cibler les éléments `span` qui sont des enfants directs d'éléments `span` dont la valeur de `class` commence par `rate` et définir leur propriété `width`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('width'));
```

Tu dois utiliser un sélecteur d'attribut pour cibler les éléments `span` qui sont des enfants directs d'éléments `span` dont la valeur de `class` commence par `rate` et définir leur propriété `background-color`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('background-color'));
```

Tu dois avoir un sélecteur d'attribut pour cibler le premier descendant des éléments `span` qui ont le mot `one` comme partie de leur valeur de `class`.

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

Tu dois utiliser un sélecteur d'attribut pour cibler le premier descendant des éléments `span` qui ont le mot `one` comme partie de leur valeur de `class`, puis définir sa propriété `background-image` pour utiliser un `linear-gradient`.

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image').includes('linear-gradient('));
```

Tu dois avoir un sélecteur d'attribut pour cibler les deux premiers descendants des éléments `span` qui ont le mot `two` comme partie de leur valeur de `class`.

```js

const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-of-type(2)',
  'span[class~="two"] > :nth-of-type(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

Tu dois utiliser un sélecteur d'attribut pour cibler les deux premiers descendants des éléments `span` qui ont le mot `two` comme partie de leur valeur de `class`, puis définir leur propriété `background-image` pour utiliser un `linear-gradient`.

```js
const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-of-type(2)',
  'span[class~="two"] > :nth-of-type(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]

assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.backgroundImage.includes('linear-gradient(')) 
```

Tu dois avoir un sélecteur d'attribut pour cibler les éléments `span` qui sont des descendants d'éléments `span` qui ont le mot `three` comme partie de leur valeur de `class`.

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.exists(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

Tu dois utiliser un sélecteur d'attribut pour cibler les éléments `span` qui sont des descendants d'éléments `span` qui ont le mot `three` comme partie de leur valeur de `class` et définir leur propriété `background-image` pour utiliser un `linear-gradient`.

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.include(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image'), 'linear-gradient(');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main>
    <h1>Book Inventory</h1>
    <div class="table-container">
      <table>
        <thead>
          <tr>
            <th>Title</th>
            <th>Author</th>
            <th>Category</th>
            <th>Status</th>
            <th>Rate</th>
          </tr>
        </thead>
        <tbody>
          <tr class="read">
            <td>The Three Musketeers</td>
            <td>Alexandre Dumas</td>
            <td>Historical Novel</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate three">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="to-read">
            <td>The Plague</td>
            <td>Albert Camus</td>
            <td>Philosophical Novel</td>
            <td><span class="status">To Read</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>The Metamorphosis</td>
            <td>Franz Kafka</td>
            <td>Novella</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate one">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Dead Souls</td>
            <td>Nikolai Gogol</td>
            <td>Picaresque</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="in-progress">
            <td>Lord of the Flies</td>
            <td>William Golding</td>
            <td>Allegorical Novel</td>
            <td><span class="status">In Progress</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Do Androids Dream of Electric Sheep?</td>
            <td>Philip K. Dick</td>
            <td>Science Fiction</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
        </tbody>
      </table>
    </div>
  </main>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
    font-family: Arial, sans-serif;
}

.table-container {
  overflow-x: auto;
}

table {
    border-collapse: collapse;
    width: 100vw;
}

thead {
    background-image: linear-gradient(hsl(197, 92%, 77%, 0.60), hsl(197, 92%, 50%, 0.60));
}

th {
    padding: 0.4em;
}

td {
    text-align: center;
    padding: 0.3em;
}

tr[class="read"] {
    background-image: linear-gradient(hsl(120, 100%, 85%, 0.6), hsl(120, 100%, 65%, 0.6));
}

tr[class="to-read"] {
    background-image: linear-gradient(hsl(120, 20%, 95%, 0.6), hsl(120, 20%, 75%, 0.6));
}

tr[class="in-progress"] {
    background-image: linear-gradient(hsl(40, 100%, 85%, 0.60), hsl(40, 100%, 65%, 0.6));
}

span {
    display: inline-block;
}

span[class="status"] {
    border-radius: 20%/60%;
}

span[class="status"],
span[class^="rate"] {
    height: 1.85em;
    width: 7em;
    padding: 0.3em;
}

tr[class="to-read"] span[class="status"] {
    border: 0.1em solid hsl(5, 100%, 45%);
    background-image: linear-gradient(hsl(5, 100%, 75%), hsl(5, 100%, 50%));
}

tr[class="read"] span[class="status"] {
    border: 0.1em solid hsl(120, 75%, 45%);
    background-image: linear-gradient(hsl(120, 75%, 75%), hsl(120, 75%, 50%));
}

tr[class="in-progress"] span[class="status"] {
    border: 0.1em solid hsl(40, 90%, 40%);
    background-image: linear-gradient(hsl(40, 90%, 75%), hsl(40, 90%, 50%));
}

span[class^="rate"] > span {
    border: 0.1em solid hsl(0, 0%, 50%);
    border-radius: 50%;
    background-color: hsl(0, 15%, 95%);
    height: 100%;
    width: 20%;
    margin: 0.1em;
}

span[class~="three"] span {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="two"] :nth-child(1),
span[class~="two"] :nth-child(2) {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="one"] :first-child {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}
```
