---
id: 66b102ee0bdbad34a9f85ba0
title: Étape 1
challengeType: 1
dashedName: step-1
---

# --description--

Dans ce workshop, tu vas continuer à apprendre les chaînes en construisant un chatbot enseignant.

Pour commencer, ajoute une instruction `console`, avec le message `"Hi there!"`.

# --hints--

Tu devrais avoir une instruction `console`.

```js
assert.match(code, /console\.log(.*)/);
```

Ton instruction console devrait afficher le message `"Hi there!"`.

```js
assert.match(code, /console\.log\((['"])(Hi\s+there!)\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
