---
id: 66b59b12a745e10011158f7b
title: Étape 4
challengeType: 1
dashedName: step-4
---

# --description--

L'étape suivante consiste à créer quelques variables supplémentaires qui serviront dans les futurs messages du bot.

Crée une variable appelée `subject` et affecte-lui la valeur de type chaîne `"JavaScript"`.

Crée ensuite une variable appelée `topic` et affecte-lui la valeur de type chaîne `"strings"`.

# --hints--

Tu devrais avoir une variable appelée `subject`.

```js
assert.isNotNull(subject);
```

Tu devrais affecter une chaîne à ta variable `subject`.

```js
assert.isString(subject);
```

Tu devrais affecter la chaîne `"JavaScript"` à ta variable `subject`.

```js
assert.equal(subject, "JavaScript");
```

Tu devrais avoir une variable appelée `topic`.

```js
assert.isNotNull(topic);
```

Tu devrais affecter une chaîne à ta variable `topic`.

```js
assert.isString(topic);
```

Tu devrais affecter la chaîne `"strings"` à ta variable `topic`.

```js
assert.equal(topic, "strings");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

--fcc-editable-region--

--fcc-editable-region--
```
