---
id: 66b59be6ab830800c4df9146
title: Étape 5
challengeType: 1
dashedName: step-5
---

# --description--

Maintenant, il est temps d'utiliser les variables que tu as créées à l'étape précédente.

Dans les leçons précédentes, tu as appris à travailler avec les littéraux de modèle comme ceci :

```js
const name = "John";
const age = 30;

// My name is John and I am 30 years old.
`My name is ${name} and I am ${age} years old.`;
```

Commence par créer une variable `sentence`.

Avec la syntaxe des littéraux de modèle, affecte la chaîne `Today, you will learn about [topic variable goes here] in [subject variable goes here].` à la variable `sentence`.

Tu remplaceras les espaces réservés `[topic variable goes here]` et `[subject variable goes here]` par les variables `topic` et `subject` et la syntaxe `${}`.

Enfin, affiche la variable `sentence` dans la console.

# --hints--

Tu devrais avoir une variable `sentence`.

```js
assert.isNotNull(sentence);
```

Ta variable `sentence` devrait être une chaîne.

```js
assert.isString(sentence);
```

Tu devrais utiliser la syntaxe des littéraux de modèle pour affecter la chaîne `Today, you will learn about [topic variable goes here] in [subject variable goes here].` à la variable `sentence`. Reporte-toi à l'exemple si tu as besoin d'aide.

```js
assert.match(code, /Today,\s*you\s+will\s+learn\s+about\s+\$\{topic\}\s+in\s+\$\{subject\}\./);
```

Tu devrais afficher la variable `sentence` dans la console.

```js
assert.match(code,/console\.log\(\s*sentence\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

--fcc-editable-region--

--fcc-editable-region--
```
