---
id: 66b59829dba144ff1351220f
title: Étape 3
challengeType: 1
dashedName: step-3
---

# --description--

Il est maintenant temps de créer une salutation à l'aide de la variable `botName`.

Dans les leçons précédentes, tu as appris à concaténer des chaînes à l'aide des littéraux de modèle comme ceci :

```js
const name = "John";

// "Hello, John!"
`Hello, ${name}!`;
```

Commence par créer une variable appelée `greeting`.

Ensuite, avec la syntaxe des littéraux de modèle, affecte une chaîne qui dit `My name is`, suivie de la variable `botName`, et se terminant par un point (`.`).

Enfin, affiche la variable `greeting` dans la console.

# --hints--

Tu devrais avoir une variable appelée `greeting`.

```js
assert.isNotNull(greeting);
```

Tu devrais affecter une chaîne à ta variable `greeting`.

```js
assert.isString(greeting);
```

Tu devrais utiliser les littéraux de modèle pour concaténer la chaîne `My name is` avec la variable `botName` suivie d'un point (`.`).

```js
assert.equal(greeting, "My name is teacherBot.");
```

Tu devrais avoir une instruction `console` qui affiche la variable `greeting` dans la console.

```js
assert.match(code, /console\.log\(.*greeting.*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
