---
id: 66b5928b86e907fdfbf34e56
title: Étape 2
challengeType: 1
dashedName: step-2
---

# --description--

Il est maintenant temps de définir le nom du bot.

Crée une variable appelée `botName` et affecte-lui la valeur de type chaîne `"teacherBot"`.

# --hints--

Tu devrais avoir une variable appelée `botName`.

```js
assert.isNotNull(botName);
```

Ta variable `botName` devrait être une chaîne.

```js
assert.isString(botName);
```

Ta variable `botName` devrait avoir la valeur `"teacherBot"`.

```js
assert.strictEqual(botName, "teacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

--fcc-editable-region--

--fcc-editable-region--
```
