---
id: 66b6e39031393e30f2c48d0f
title: Étape 7
challengeType: 1
dashedName: step-7
---

# --description--

Pour obtenir la longueur d'une chaîne, tu peux utiliser la propriété <dfn>length</dfn>. Cette propriété renvoie le nombre de caractères d'une chaîne.

Voici un exemple :

```js
const greeting = "Hello, world!";

// Output: 13
console.log(greeting.length);
```

Commence par utiliser la propriété `length` pour obtenir la longueur de la chaîne `subject` et affiche cette valeur dans la console.

# --hints--


Tu devrais utiliser la propriété `length` pour obtenir la longueur de la chaîne `subject`.

```js
assert.match(code, /subject\.length/);
```

Tu devrais afficher la longueur de la chaîne `subject` dans la console.

```js
assert.match(code, /console\.log\(\s*subject\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

--fcc-editable-region--

--fcc-editable-region--
```
