---
id: 66b6fdb76441c738719039fa
title: Étape 13
challengeType: 1
dashedName: step-13
---

# --description--

Dans les leçons précédentes, tu as appris à accéder au dernier caractère d'une chaîne comme ceci :

```js
const firstName = "Jessica";

// returns "a"
firstName[firstName.length - 1];
```

`string.length - 1` te donnera toujours le dernier indice d'une chaîne.

Crée une nouvelle variable appelée `lastCharacter` et affecte-lui la valeur du dernier caractère de la variable `subject`.

Ensuite, affiche la valeur de la variable `lastCharacter` dans la console.

# --hints--

Tu devrais avoir une variable appelée `lastCharacter`.

```js
assert.isNotNull(lastCharacter);
```

Tu devrais affecter la valeur du dernier caractère de la variable `subject` à la variable `lastCharacter`. Reporte-toi à l'exemple si tu as besoin d'aide.

```js
assert.strictEqual(lastCharacter, subject[subject.length - 1]);
```

Tu devrais afficher la valeur de la variable `lastCharacter` dans la `console`.

```js
assert.match(code, /console\.log\(lastCharacter\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

--fcc-editable-region--

--fcc-editable-region--
```
