---
id: 66b6f586767a1534f3097353
title: Étape 11
challengeType: 1
dashedName: step-11
---

# --description--

Il est maintenant temps d'accéder à la deuxième lettre de la variable `subject`.

Commence par ajouter une instruction `console` qui affiche le message `Here is an example of accessing the second letter in the word [subject].` Pense à remplacer `[subject]` par la valeur réelle de la variable `subject` et à utiliser la bonne syntaxe des littéraux de modèle.

Ajoute ensuite une autre instruction `console` qui affiche la deuxième lettre de la variable `subject` à l'aide de la notation entre crochets et du bon indice.

# --hints--

Ton instruction `console` devrait afficher le message `Here is an example of accessing the second letter in the word [subject].`

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+accessing\s+the\s+second\s+letter\s+in\s+the\s+word\s+\$\{subject\}\.`\);?/);
```

Tu devrais avoir une autre instruction `console` qui affiche la deuxième lettre de la variable `subject` à l'aide de la notation entre crochets et de l'indice `1`.

```js
assert.match(code, /console\.log\(subject\[1\]\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

--fcc-editable-region--

--fcc-editable-region--
```
