---
id: 66b6ffb42b88e33943788abf
title: Étape 14
challengeType: 1
dashedName: step-14
---

# --description--

Pour la dernière partie du workshop, tu vas revoir comment trouver la position d'une sous-chaîne dans une chaîne. Souviens-toi qu'une sous-chaîne est une partie d'une chaîne.

Commence par créer une variable appelée `learningIsFunSentence` et affecte-lui la valeur de type chaîne `"Learning is fun."`.

# --hints--

Tu devrais avoir une variable appelée `learningIsFunSentence`.

```js
assert.isNotNull(learningIsFunSentence);
```

Ta variable `learningIsFunSentence` devrait contenir une valeur de type chaîne.

```js
assert.isString(learningIsFunSentence);
```

Tu devrais affecter la valeur `"Learning is fun."` à la variable `learningIsFunSentence`.

```js
assert.equal(learningIsFunSentence, "Learning is fun.");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

--fcc-editable-region--

--fcc-editable-region--
```
