---
id: 66b6e62423e8a031d6c1f03d
title: Étape 8
challengeType: 1
dashedName: step-8
---

# --description--

Il est maintenant temps d'obtenir la longueur de la chaîne `topic`.

Tu peux utiliser des littéraux de modèle à l'intérieur des instructions `console` comme ceci :

```js
const developer = "Jessica";
console.log(`Hello, my name is ${developer}.`);
```

Commence par afficher le message `Here is an example of using the length property on the word [topic].` dans la console.

Pense à remplacer `[topic]` par la variable `topic`, et à utiliser la bonne syntaxe des littéraux de modèle comme aux étapes précédentes.

Ensuite, ajoute une deuxième instruction `console.log` qui affiche la longueur de la chaîne `topic` dans la console.

# --hints--

Tu devrais utiliser la syntaxe des littéraux de modèle pour afficher le message `Here is an example of using the length property on the word [topic].` dans la `console`.

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+using\s+the\s+length\s+property\s+on\s+the\s+word\s+\$\{topic\}\.`\)/);
```

Tu devrais avoir une deuxième instruction `console` pour afficher la longueur de la variable `topic` dans la `console`.

```js
assert.match(code, /console\.log\(\s*topic\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);
console.log(subject.length);

--fcc-editable-region--

--fcc-editable-region--
```
