---
id: 66b6efddeca35833cd6f0b03
title: Étape 10
challengeType: 1
dashedName: step-10
---

# --description--

Dans les leçons précédentes, tu as appris à accéder aux caractères d'une chaîne comme ceci :

```js
const firstName = "Jessica";
// returns "J"
firstName[0];
```

Souviens-toi que les indices commencent à `0`, donc la première lettre d'une chaîne sera toujours à l'indice `0`.

Commence par ajouter une autre instruction `console`.

À l'intérieur de l'instruction `console`, affiche la première lettre de la variable `subject` à l'aide de la notation entre crochets et du bon indice.

# --hints--

Ton instruction `console` devrait afficher la première lettre de la variable `subject` à l'aide de la notation entre crochets et de l'indice `0`.

```js
assert.match(code, /console\.log\(subject\[0\]\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

--fcc-editable-region--

--fcc-editable-region--
```
