---
id: 66b6d482bbb9e12f2e5ee1ae
title: Étape 6
challengeType: 1
dashedName: step-6
---

# --description--

Pour cette prochaine partie du projet, le bot va apprendre à travailler avec la propriété `length` des chaînes.

Commence par créer une nouvelle variable appelée `strLengthIntro`.

Ensuite, avec la syntaxe des littéraux de modèle, affecte la chaîne `Here is an example of using the length property on the word [subject].` à la variable `strLengthIntro`.

Remplace `[subject]` par la variable `subject` comme tu l'as fait précédemment.

Enfin, affiche la variable `strLengthIntro` dans la console.

# --hints--

Tu devrais avoir une variable appelée `strLengthIntro`.

```js
assert.isNotNull(strLengthIntro);
```

Ta variable `strLengthIntro` devrait être une chaîne.

```js
assert.isString(strLengthIntro);
```

Tu devrais utiliser la syntaxe des littéraux de modèle pour affecter la chaîne `Here is an example of using the length property on the word [subject goes here].` à la variable `strLengthIntro`.

```js
assert.match(code,/`Here\s+is\s+an\s+example\s+of\s+using\s+the\s+length\s+property\s+on\s+the\s+word\s+\$\{subject\}\.`/);
```

Tu devrais afficher la variable `strLengthIntro` dans la console.

```js
assert.match(code,/console\.log\(\s*strLengthIntro\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

--fcc-editable-region--

--fcc-editable-region--
```
