---
id: 66b70b8a611cbf3bcc5c6c5f
title: Étape 16
challengeType: 1
dashedName: step-16
---

# --description--

Dans les leçons précédentes, tu as appris à travailler avec la méthode `indexOf` comme ceci :

```js
const sentence = "I love to learn.";

// returns index 2
console.log(sentence.indexOf("love"));

// returns -1
console.log(sentence.indexOf("hate"));
```

Souviens-toi que la méthode `indexOf` renvoie la position de la première occurrence d'une sous-chaîne dans une chaîne. Si la sous-chaîne n'est pas trouvée, elle renvoie `-1`.

Ajoute une nouvelle instruction `console` qui affiche le résultat de l'utilisation de la méthode `indexOf` sur la variable `learningIsFunSentence` pour trouver la position de la sous-chaîne `"Learning"`.

# --hints--

Ta `console` devrait utiliser la méthode `indexOf` sur la variable `learningIsFunSentence` pour trouver la position de la sous-chaîne `"Learning"`.

```js
assert.match(code, /console\.log\(\s*learningIsFunSentence\.indexOf\(\s*(['"])(Learning)\1\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

const learningIsFunSentence = "Learning is fun.";

console.log("Here are examples of finding the positions of substrings in the sentence.");

--fcc-editable-region--

--fcc-editable-region--
```
