---
id: 69a0102e761a245bb1cbd8e7
title: Étape 5
challengeType: 1
dashedName: step-5
---

# --description--

À ce stade, tu devrais tester la liste chaînée pour mieux la comprendre.

Crée une variable `myList` et définis-la à `initList()`. En dessous, affiche dans la console un appel à `isEmpty()` en passant `myList` en paramètre. Ensuite, appelle `add` avec `myList` et un nombre de ton choix, puis affiche `myList` dans la console. En dessous, affiche à nouveau un appel à `isEmpty()` dans la console en passant encore `myList` en paramètre.

# --hints--

Tu dois créer une variable `myList` définie à `initList()`.

```js
assert.deepEqual(myList, initList());
```

Tu dois afficher un appel à `isEmpty(myList)` dans la console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(isEmpty\(myList\)\);?/)
```

Tu dois appeler `add` avec `myList` et un nombre de ton choix comme paramètres.

```js
assert.match(__helpers.removeJSComments(code), /add\s*\(\s*myList\s*,\s*\d+\s*\)/);
```

Tu dois afficher ta variable `myList` dans la console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(myList\);?/)
```

Tu dois afficher à nouveau `isEmpty(myList)` dans la console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(myList\);?\s*console\.log\(isEmpty\(myList\)\);?/)
```


# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };
}

--fcc-editable-region--

--fcc-editable-region--
```
