---
id: 69a00a051b9a18694d00c16b
title: Étape 1
challengeType: 1
dashedName: step-1
---

# --description--

Dans cet atelier, tu vas implémenter une liste chaînée.

Pour commencer, crée une fonction nommée `initList`. À l'intérieur, renvoie un objet avec une propriété `head` définie à `null`, et une propriété `length` égale à `0`.

C'est la fonction que tu appelleras pour initialiser une liste chaînée. Le `head` représente le premier élément de la liste ; au départ, il vaut `null` car il n'y a aucun élément dans la liste. La propriété `length` suit le nombre de nœuds (éléments) dans la liste, que tu incrémenteras ou décrémenteras quand tu ajoutes ou retires un nœud de la liste.

# --hints--

Tu dois créer une fonction `initList`.

```js
assert.isFunction(initList)
```

Ta fonction `initList` doit renvoyer un objet.

```js
assert.isObject(initList());
```

Ta fonction `initList` doit renvoyer un objet avec un `head` défini à `null`.

```js
assert.propertyVal(initList(), "head", null);
```

Ta fonction `initList` doit renvoyer un objet avec un `length` défini à `0`.

```js
assert.propertyVal(initList(), "length", 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
