---
id: 69a04466917c66b132323a1b
title: Étape 18
challengeType: 1
dashedName: step-18
---

# --description--

Si `previous` n'est pas `null`, cela signifie que l'élément à retirer n'est pas à la tête mais quelque part dans la liste, donc le nœud actuel doit être contourné en reliant le nœud précédent directement au nœud suivant.

Pour cela, crée une autre instruction `if` qui vérifie si previous n'est pas strictement `null`. Si c'est le cas, définis `previous.next` à `current.next`.

# --hints--

Tu dois créer une autre instruction `if` avec la condition `previous !== null`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(previous\s*!==\s*null\s*\)\s*\{\s*/);
```

Tu dois définir `previous.next` à `current.next`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(previous\s*!==\s*null\s*\)\s*\{\s*previous\.next\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next;
    }

    if (current === null) {
        return;
    }
--fcc-editable-region--

--fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
