---
id: 69a012c6224fac85832247dd
title: Étape 6
challengeType: 1
dashedName: step-6
---

# --description--

Tu remarques les résultats `true`, `{ head: null, length: 0 }` et `true` dans la console ? Cela se produit parce que tu n'as rien assigné à la liste `myList`, et que tu n'as pas incrémenté la propriété `length`.

Assigne `node` à `list.head`, et en dessous, incrémente la propriété `length` de `list`.

# --hints--

Tu dois assigner `node` à `list.head`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /list\.head\s*=\s*node;?/);
```

Tu dois incrémenter la propriété `length` de `list`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /list\.length\+\+;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    --fcc-editable-region--

    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
