---
id: 69a04088e600bd8074a9852f
title: Étape 17
challengeType: 1
dashedName: step-17
---

# --description--

Si `current` est `null`, cela signifie que c'est la fin de la liste, donc la fonction doit cesser de s'exécuter.

Crée une instruction `if` qui vérifie si `current` est strictement égal à `null`. À l'intérieur de l'instruction `if`, utilise `return` pour quitter la fonction.

# --hints--

Tu dois créer une instruction `if` avec la condition `current === null`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(current\s*===\s*null\s*\)\s*\{\s*/);
```

Tu dois utiliser le mot-clé `return` à l'intérieur de ton instruction `if`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(current\s*===\s*null\s*\)\s*\{\s*return;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next;
    }
    --fcc-editable-region--

    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
