---
id: 69a045fd236d0d973a59f103
title: Étape 19
challengeType: 1
dashedName: step-19
---

# --description--

Si `previous` est `null`, cela signifie que l'élément à retirer est à la tête de la liste, donc tu dois mettre à jour le `head` pour qu'il pointe vers le nœud suivant afin de retirer le premier élément.

Pour cela, ajoute une clause `else`. À l'intérieur, définis `list.head` à `current.next`.

# --hints--

Tu dois ajouter une clause `else` à ton instruction `if`.

```js
assert.match(__helpers.removeJSComments(remove.toString().split("previous.next = current.next;")), /\}\s*else\s*\{\s*/);
```

Tu dois définir `list.head` à `current.next` à l'intérieur de la clause `else`.

```js
assert.match(__helpers.removeJSComments(remove.toString().split("previous.next = current.next;")), /\}\s*else\s*\{\s*list\.head\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next;
    }

    if (current === null) {
        return;
    }

    if(previous !== null) {
        previous.next = current.next;
    --fcc-editable-region--
    }

    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
