---
id: 69a00eaa43458e618331873c
title: Étape 3
challengeType: 1
dashedName: step-3
---

# --description--

Maintenant, tu vas travailler sur la fonction qui ajoutera un nœud à la liste.

Crée une fonction `add` avec les paramètres `list` et `element`.

# --hints--

Tu dois créer une fonction `add`.

```js
assert.isFunction(add)
```

Ta fonction `add` doit avoir les paramètres `list` et `element`.

```js
const addParams = __helpers.getFunctionParams(add.toString());
const addParamsNames = addParams.map(param => param.name)
assert.include(addParamsNames, "list");
assert.include(addParamsNames, "element");
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

--fcc-editable-region--

--fcc-editable-region--
```
