---
id: 69a013a86c97dac3784b2636
title: Étape 7
challengeType: 1
dashedName: step-7
---

# --description--

Maintenant, le log devrait afficher quelque chose comme `{ head: { element: 42, next: null }, length: 1 }` dans la console.

Cependant, ce n'est pas ainsi que tu gérerais le cas où le nœud ajouté est le premier élément de la liste, c'est-à-dire le head.

Pour le faire correctement, utilise une instruction if pour vérifier si la liste est vide en appelant la fonction `isEmpty` avec `list`. Si c'est vrai, alors déplace ta ligne `list.head = node` dans le corps de l'instruction `if`.

# --hints--

Tu dois avoir une instruction `if` avec la condition `isEmpty(list)`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /if\s*\(isEmpty\(list\)\)\s*\{\s*/);
```

Tu dois déplacer ta ligne `list.head = node` dans l'instruction `if`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /if\s*\(isEmpty\(list\)\)\s*\{\s*list\.head\s*=\s*node;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

--fcc-editable-region--
    list.head = node;
--fcc-editable-region--
    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
