---
id: 69a014bd9b1730e94136279f
title: Étape 8
challengeType: 1
dashedName: step-8
---

# --description--

Si la liste n'est pas vide, tu veux trouver sa fin, y ajouter le nouvel élément, et le faire pointer vers `null`. Cela signifie que tu dois parcourir la liste à partir du `head`.

Pour commencer, ouvre une clause `else`. À l'intérieur, utilise `let` pour créer une variable `current` et définis-la à `list.head`.

# --hints--

Tu dois avoir une clause `else`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /else\s*{/);
```

Tu dois créer une variable `current` définie à `list.head` à l'intérieur de la clause `else`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /else\s*{\s*let|var\s+current\s*=\s*list\.head;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    --fcc-editable-region--
    }
    --fcc-editable-region--

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
