---
id: 69a00cd4ec513345d40e95ec
title: Étape 2
challengeType: 1
dashedName: step-2
---

# --description--

Ensuite, tu vas travailler sur une fonction qui te permet de vérifier si la liste est vide.

Pour cela, crée une fonction nommée `isEmpty` qui prend un paramètre `list`. À l'intérieur de la fonction, renvoie une vérification pour voir si la propriété `length` de `list` est strictement égale à `0`.

# --hints--

Tu dois créer une fonction `isEmpty`.

```js
assert.isFunction(isEmpty)
```

Ta fonction `isEmpty` doit avoir un paramètre `list`.

```js
const isEmptyParams = __helpers.getFunctionParams(isEmpty.toString());
assert.deepEqual(isEmptyParams[0].name, "list");
```

Ta fonction `isEmpty` doit renvoyer `list.length === 0`.

```js
const emptyList = { head: null, length: 0 };
const nonEmptyList = { head: {}, length: 1 };

assert.isTrue(isEmpty(emptyList));
assert.isFalse(isEmpty(nonEmptyList));
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

--fcc-editable-region--

--fcc-editable-region--
```
