---
id: 69a00f673f171e29159e3ad8
title: Étape 4
challengeType: 1
dashedName: step-4
---

# --description--

Ce nœud aurait besoin de deux propriétés : l'élément à ajouter, et un pointeur vers le nœud suivant.

À l'intérieur de la fonction add, crée un objet nommé `node` ; il doit avoir une propriété nommée `element`, et une propriété `next` définie à `null`, qui est le pointeur vers l'élément suivant de la liste.

# --hints--

Tu dois créer un objet nommé `node` à l'intérieur de ta fonction `add`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /const|var|let\s+node\s*=\s*{/);
```

Ton objet `node` doit avoir une propriété nommée `element` définie à `element`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /element\s*:\s*element/);
```

Ton objet `node` doit avoir une propriété nommée `next` définie à `null`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /next\s*:\s*null/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    --fcc-editable-region--

    --fcc-editable-region--
}
```
