---
id: 69a0179a85b00b6127e0af45
title: Étape 10
challengeType: 1
dashedName: step-10
---

# --description--

À l'intérieur de la boucle `while`, mets à jour `current` à `current.next` pour que la boucle avance dans la liste un nœud à la fois, et s'arrête quand elle atteint le nœud dont la propriété `next` vaut `null`, c'est-à-dire la fin de la liste.

# --hints--

Tu dois mettre à jour la valeur de ta variable `current` à `current.next` à l'intérieur de la boucle `while`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /current\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            --fcc-editable-region--

            --fcc-editable-region--
        }
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
