---
id: 69a03aaa1e6fe8152788ce6a
title: Étape 14
challengeType: 1
dashedName: step-14
---

# --description--

À l'intérieur de la fonction `remove`, utilise `let` pour créer une variable `previous` définie à `null`, et `let` à nouveau pour créer une variable `current` définie à `list.head`. Ce sont les variables que tu utiliseras dans la boucle qui parcourra la liste.

`current` suit le nœud actuellement vérifié dans la liste, en commençant à la tête et en avançant d'un nœud à la fois. `previous`, de son côté, garde la trace du nœud juste avant le nœud actuel.

C'est important lors de la suppression d'un nœud, car tu dois mettre à jour la valeur next du nœud précédent pour qu'il pointe vers le nœud après celui qui est retiré, afin de garder la liste connectée.

# --hints--

Tu dois utiliser `let` pour créer une variable `previous` définie à `null` à l'intérieur de ta fonction `remove`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /let|var\s+previous\s*=\s*null;?/);
```

Tu dois utiliser `let` pour créer une variable `current` définie à `list.head` à l'intérieur de ta fonction `remove`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /let|var\s+current\s*=\s*list\.head;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    --fcc-editable-region--

    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
