---
id: 69a01864d343dc2f2484d131
title: Étape 11
challengeType: 1
dashedName: step-11
---

# --description--

Après avoir trouvé le dernier nœud, tu dois y attacher le nouveau nœud. Pour cela, définis `current.next` à `node` sous la boucle while mais toujours dans la clause `else`.

# --hints--

Tu dois définir `current.next` à `node` sous ta boucle `while`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /current\.next\s*=\s*node;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        --fcc-editable-region--

        --fcc-editable-region--
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
