---
id: 69a03ddce82b0626f329b001
title: Étape 15
challengeType: 1
dashedName: step-15
---

# --description--

Maintenant, crée une boucle `while` qui vérifie si `current` n'est pas strictement `null` et si `current.element` n'est pas strictement égal à element.

Cela signifie que la boucle doit continuer à parcourir la liste jusqu'à ce que la fin de la liste soit atteinte ou que l'élément à retirer soit trouvé.

# --hints--

Tu dois créer une boucle `while` avec la condition `current !== null && current.element !== element`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /while\s+\(current\s*!==\s*null\s*&&\s*current\.element\s*!==\s*element\)\s*\{\s*/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    --fcc-editable-region--

    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
