---
id: 69b279bae0f20eda63163cf6
title: Étape 12
challengeType: 1
dashedName: step-12
---

# --description--

Dans une étape précédente, tu as assigné l'élément de `leftPart` à la liste `array` parce qu'il était plus petit. Mais ce ne sera pas toujours le cas. Dans certaines comparaisons, l'élément de droite pourrait être plus petit.

Crée une clause `else` à exécuter quand l'élément de `leftPart` n'est pas inférieur à l'élément de `rightPart`.

À l'intérieur du bloc `else`, assigne `rightPart[rightArrayIndex]` à `array[sortedIndex]`.

# --hints--

Tu dois créer un bloc `else` qui contient `array[sortedIndex] = rightPart[rightArrayIndex]`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /else\s*\{[\s\S]*array\s*\[\s*sortedIndex\s*\]\s*=\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*;?[\s\S]*\}/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);

  mergeSort(leftPart);
  mergeSort(rightPart);

  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
--fcc-editable-region--
    }
--fcc-editable-region--
  }
}
```
