---
id: 69b279bae0f20eda63163cf2
title: Étape 8
challengeType: 1
dashedName: step-8
---

# --description--

Ensuite, tu dois créer une boucle qui continue tant qu'il reste des éléments dans `leftPart` et dans `rightPart`.

Pour cela, crée une boucle `while` avec deux conditions : l'une qui vérifie si `leftArrayIndex` est inférieur à la longueur de `leftPart` et une autre qui vérifie si `rightArrayIndex` est inférieur à la longueur de `rightPart`.

# --hints--

Tu dois créer une boucle `while` qui vérifie si `leftArrayIndex` est inférieur à la longueur de `leftPart` et si `rightArrayIndex` est inférieur à la longueur de `rightPart`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*&&\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);

  mergeSort(leftPart);
  mergeSort(rightPart);

  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

--fcc-editable-region--

--fcc-editable-region--
}
```
