---
id: 69b279bae0f20eda63163cf1
title: Étape 7
challengeType: 1
dashedName: step-7
---

# --description--

Il est maintenant temps de trier et de fusionner les listes (`leftPart` et `rightPart`) dans le `array` d'origine.

Tu peux le faire en comparant les éléments des deux listes, et en fusionnant le plus petit élément dans la liste principale. Tu vas faire cette comparaison pour tous les indices de `leftPart` et `rightPart`.

Crée trois variables : `leftArrayIndex`, `rightArrayIndex` et `sortedIndex`, et définis leurs valeurs à `0`. Ces variables t'aideront à suivre chaque indice pendant le processus de tri.

# --hints--

Tu dois déclarer une variable `leftArrayIndex` initialisée à `0`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+leftArrayIndex\s*=\s*0\s*;?/
);
```

Tu dois déclarer une variable `rightArrayIndex` initialisée à `0`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+rightArrayIndex\s*=\s*0\s*;?/
);
```

Tu dois déclarer une variable `sortedIndex` initialisée à `0`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+sortedIndex\s*=\s*0\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);

  mergeSort(leftPart);
  mergeSort(rightPart);

--fcc-editable-region--

--fcc-editable-region--
}
```
