---
id: 69b278d82f0e514f32acb724
title: Étape 4
challengeType: 1
dashedName: step-4
---

# --description--

Après avoir découpé la moitié gauche, tu as encore besoin du côté droit pour compléter l'étape de fusion.

Utilise la syntaxe slice pour extraire la moitié droite de `array` et assigne-la à une variable nommée `rightPart`.

# --hints--

Tu dois avoir une variable `rightPart` à l'intérieur de la fonction `mergeSort`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+rightPart\s*=/
);
```

Tu dois utiliser la syntaxe slice pour extraire la moitié droite du tableau et l'assigner à la variable `rightPart`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /rightPart\s*=\s*array\.slice\(\s*middlePoint\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
--fcc-editable-region--

--fcc-editable-region--
}
```
