---
id: 69b279bae0f20eda63163cf3
title: Étape 9
challengeType: 1
dashedName: step-9
---

# --description--

À l'intérieur de la boucle, commence par comparer les éléments courants de chaque moitié afin de fusionner d'abord la plus petite valeur.

Ajoute un bloc `if` qui vérifie si l'élément courant dans `leftPart` est inférieur à l'élément courant dans `rightPart`.
# --hints--

La condition `if` doit vérifier si `leftPart[leftArrayIndex]` est inférieur à `rightPart[rightArrayIndex]`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*leftPart\s*\[\s*leftArrayIndex\s*\]\s*<\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);

  mergeSort(leftPart);
  mergeSort(rightPart);

  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
