---
id: 69b279bae0f20eda63163cf8
title: Étape 14
challengeType: 1
dashedName: step-14
---

# --description--

Les instructions `if` et `else` que tu as créées dans les étapes précédentes assignent des éléments au tableau trié.

Chaque élément assigné au tableau trié occupe un indice dans la liste. Tu dois donc passer à l'indice suivant du tableau trié après chaque assignation.

Sous le bloc `if`/`else`, mais toujours dans la boucle `while`, incrémente `sortedIndex` de `1`.

# --hints--

Tu dois incrémenter la valeur actuelle de `sortedIndex` de `1` après ta clause `else`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*&&\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{[\s\S]*sortedIndex\s*(\+\+|\+=\s*1)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);

  mergeSort(leftPart);
  mergeSort(rightPart);

  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
