---
id: 69b279bae0f20eda63163cfc
title: Étape 18
challengeType: 1
dashedName: step-18
---

# --description--

La dernière chose à faire pour la boucle `while` est de passer à l'indice suivant du tableau trié.

Ajoute `1` à la valeur de `sortedIndex`.

Note que l'incrémentation doit se faire après avoir assigné l'élément de la partie gauche au tableau trié.

# --hints--

Tu dois incrémenter la valeur actuelle de `sortedIndex` d'un à l'intérieur de la boucle `while`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*\)\s*\{[\s\S]*sortedIndex\s*(\+\+|\+=\s*1)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {

  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);

  mergeSort(leftPart);
  mergeSort(rightPart);

  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

  while (leftArrayIndex < leftPart.length) {
    array[sortedIndex] = leftPart[leftArrayIndex];
    leftArrayIndex += 1;
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
