---
id: 69b2713ff6c21fed2d3ad88b
title: Étape 1
challengeType: 1
dashedName: step-1
---

# --description--

Le tri fusion (merge sort) est un algorithme de tri qui utilise le principe diviser pour régner pour trier des collections de données. Autrement dit, il « divise » une collection en sous-parties plus petites, et « conquiert » ces sous-parties en les triant indépendamment, puis fusionne les sous-parties triées.

Commence par créer une fonction appelée `mergeSort` qui prend le paramètre `array`.

# --hints--

Tu dois déclarer une fonction nommée `mergeSort` avec un seul paramètre `array`.

```js
assert.isFunction(mergeSort);
const params = __helpers.getFunctionParams(mergeSort.toString());
assert.lengthOf(params, 1);
assert.deepEqual(params[0].name, "array");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
