---
id: 69b279bae0f20eda63163cf9
title: Étape 15
challengeType: 1
dashedName: step-15
---

# --description--

La boucle `while` que tu as créée compare un élément de `leftPart` avec un autre de `rightPart`, puis ajoute le plus petit élément à la liste principale `array`.

Elle continue cette opération jusqu'à ce qu'il ne reste plus d'éléments à comparer. Mais `leftPart` peut encore avoir des éléments alors que `rightPart` n'en a plus, et inversement.

Crée une autre boucle `while` qui s'exécute quand `leftArrayIndex` est inférieur à la longueur de `leftPart`. Dans les prochaines étapes, tu l'utiliseras pour copier les éléments restants de `leftPart` dans la liste `array`.

# --hints--

Tu dois créer une boucle `while` avec la condition `leftArrayIndex < leftPart.length`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);

  mergeSort(leftPart);
  mergeSort(rightPart);

  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

--fcc-editable-region--

--fcc-editable-region--
}
```
