---
id: 69b2762b17d23fbee40a6c00
title: Étape 2
challengeType: 1
dashedName: step-2
---

# --description--

L'algorithme de tri fusion effectue trois actions :

- Diviser une séquence non triée d'éléments en sous-parties
- Trier les éléments dans les sous-parties
- Fusionner les sous-parties triées

Ce qui précède se produit de façon récursive jusqu'à ce que les sous-parties soient fusionnées en la séquence triée complète. Commençons par diviser la séquence.

D'abord, à l'intérieur de la fonction `mergeSort`, déclare une variable `middlePoint` avec `const` et assigne-lui `Math.floor(array.length / 2)` afin de pouvoir couper la séquence en deux avant de récursiver.

# --hints--

Tu dois déclarer une variable `middlePoint` et lui assigner `Math.floor(array.length / 2)`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(const|let|var)\s+middlePoint\s*=\s*Math\.floor\(\s*array\.length\s*\/\s*2\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--

--fcc-editable-region--
}
```
