---
id: 69b279bae0f20eda63163cf4
title: Étape 10
challengeType: 1
dashedName: step-10
---

# --description--

Quand la condition `if` s'évalue à `true`, cela signifie que l'élément de la liste `leftPart` est plus petit que l'élément auquel il est comparé dans la liste `rightPart`.

Dans ce cas, tu peux placer l'élément trouvé à `leftArrayIndex` dans `leftPart` dans le tableau trié.

À l'intérieur du bloc `if`, assigne `leftPart[leftArrayIndex]` à `array[sortedIndex]`.

# --hints--

Tu dois assigner `leftPart[leftArrayIndex]` à `array[sortedIndex]` à l'intérieur du `if`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /array\s*\[\s*sortedIndex\s*\]\s*=\s*leftPart\s*\[\s*leftArrayIndex\s*\]\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);

  mergeSort(leftPart);
  mergeSort(rightPart);

  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
--fcc-editable-region--

--fcc-editable-region--
    }
  }
}
```
