---
id: 69b278519ef46b2fd8bf1546
title: Étape 3
challengeType: 1
dashedName: step-3
---

# --description--

Tu as déjà une valeur `middlePoint`. Utilise-la pour couper le tableau en deux moitiés et garder la portion gauche prête pour le tri.

Utilise la syntaxe slice pour extraire la moitié gauche de `array` et assigne-la à une variable nommée `leftPart`.

# --hints--

Tu dois déclarer une variable `leftPart` à l'intérieur de la fonction `mergeSort`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+leftPart\s*=/
);
```

Tu dois assigner `array.slice(0, middlePoint)` à `leftPart`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /leftPart\s*=\s*array\.slice\(\s*0\s*,\s*middlePoint\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
--fcc-editable-region--

--fcc-editable-region--
}
```
