---
id: 69b279bae0f20eda63163cf0
title: Étape 6
challengeType: 1
dashedName: step-6
---

# --description--

En bas du corps de ta fonction, appelle à nouveau la fonction `mergeSort()`. Cette fois, passe `rightPart` comme argument à l'appel de fonction.

# --hints--

Tu dois appeler `mergeSort()` avec l'argument `rightPart` en bas du corps de ta fonction.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /mergeSort\s*\(\s*rightPart\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);

  mergeSort(leftPart);
--fcc-editable-region--

--fcc-editable-region--
}
```
