---
id: 69b279bae0f20eda63163cef
title: Étape 5
challengeType: 1
dashedName: step-5
---

# --description--

Maintenant que tu as divisé la liste `array` en deux listes séparées, tu vas continuer à diviser chaque liste jusqu'à ce que chaque élément soit seul dans sa propre liste. Une liste avec un seul nombre est toujours triée.

Pour cela, appelle `mergeSort` de façon récursive à l'intérieur de ta fonction et passe `leftPart` comme argument à l'appel.

# --hints--

Tu dois appeler `mergeSort(leftPart)` à l'intérieur de la même fonction `mergeSort`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /mergeSort\s*\(\s*leftPart\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);

--fcc-editable-region--

--fcc-editable-region--
}
```
