---
id: 69b279bae0f20eda63163cfe
title: Étape 20
challengeType: 1
dashedName: step-20
---

# --description--

Avant de tester la fonction `mergeSort()`, tu dois créer un cas de base qui arrête l'exécution de la fonction quand la longueur de `array` est inférieure ou égale à `1`.

Ce cas de base arrêtera l'appel récursif. Sans lui, l'opération de tri fusion continuerait même quand la liste a été triée ou n'a plus d'éléments.

Juste après la déclaration de la fonction, crée une instruction `if` qui vérifie si `array.length` est inférieur ou égal à 1. Dans le corps de l'instruction `if`, ajoute une instruction `return` pour arrêter l'exécution de la fonction.

# --hints--

Tu dois créer une instruction `if` avec la condition `array.length <= 1`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*array\.length\s*<=\s*1\s*\)\s*\{/
);
```

Tu dois avoir une instruction `return` dans le corps de l'instruction `if`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*array\.length\s*<=\s*1\s*\)\s*\{[\s\S]*return\s*;?[\s\S]*\}/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--

--fcc-editable-region--
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);

  mergeSort(leftPart);
  mergeSort(rightPart);

  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

  while (leftArrayIndex < leftPart.length) {
    array[sortedIndex] = leftPart[leftArrayIndex];
    leftArrayIndex += 1;
    sortedIndex += 1;
  }

  while (rightArrayIndex < rightPart.length) {
    array[sortedIndex] = rightPart[rightArrayIndex];
    rightArrayIndex += 1;
    sortedIndex += 1;
  }
}
```
