---
id: a202eed8fc186c8434cb6d61
title: Crée un inverseur de chaîne
challengeType: 26
dashedName: build-a-string-inverter
---

# --description--

Dans ce lab, tu vas créer un simple inverseur de chaîne qui inverse les caractères d'une chaîne donnée.

Par exemple, `"hello"` doit devenir `"olleh"`.

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois créer une fonction nommée `reverseString` qui prend une chaîne en argument.
2. La fonction doit renvoyer la chaîne inversée.

# --hints--

Tu dois avoir une fonction nommée `reverseString`.

```js
assert.isFunction(reverseString);
```

`reverseString` doit prendre une chaîne en argument.

```js
assert.match(reverseString.toString(), /\s*function\s+reverseString\s*\(\s*\w+\s*\)/);
```

`reverseString("hello")` doit renvoyer une chaîne.

```js
assert.isString(reverseString('hello'));
```

`reverseString("hello")` doit renvoyer la chaîne `olleh`.

```js
assert.strictEqual(reverseString('hello'), 'olleh');
```

`reverseString("Howdy")` doit renvoyer la chaîne `ydwoH`.

```js
assert.strictEqual(reverseString('Howdy'), 'ydwoH');
```

`reverseString("Greetings from Earth")` doit renvoyer la chaîne `htraE morf sgniteerG`.

```js
assert.strictEqual(
  reverseString('Greetings from Earth'),
  'htraE morf sgniteerG'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function reverseString(str) {
  return str.split('').reverse().join('');
}

reverseString('hello');
```
