---
id: 66db57ad34c7089b9b41bfd6
title: Construire une visionneuse Lightbox
challengeType: 25
dashedName: build-a-lightbox-viewer
demoType: onClick
---

# --description--

Une Lightbox est utilisée sur les sites web pour présenter du contenu multimédia de façon plus attrayante au moyen d'une fenêtre contextuelle ou modale qui s'affiche au-dessus du contenu principal de la page.

Dans ce lab, tu vas créer une galerie Lightbox qui affiche les images en taille réelle lorsque l'utilisateur clique sur une miniature. Deux versions de chaque image sont fournies : une miniature et une image en taille réelle. L'image en taille réelle est identique à la miniature, mais sans le suffixe `-thumbnail`.

**Objectif :** Réponds aux user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois avoir un élément `div` avec une classe `gallery` dans ton élément `body`.
1. Dans l'élément `.gallery`, tu dois avoir trois miniatures, chacune avec une classe `gallery-item`. Tu dois utiliser les liens suivants pour les miniatures :

    - `https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg`

1. Tu dois avoir un élément `div` avec une classe `lightbox` dans ton élément `body`.
1. Tu dois avoir un élément `button` avec l'`id` `close-btn` dans ton élément `.lightbox`. Tu peux utiliser `&times;` comme texte si tu le souhaites.
1. Tu dois avoir un élément `img` avec l'`id` `lightbox-image` dans ton élément `.lightbox`.
1. Ton élément `.lightbox` doit avoir une position fixe afin que l'aperçu s'ouvre au-dessus des images actuelles.
1. Ton élément `.lightbox` doit couvrir toute la zone d'affichage en définissant sa hauteur et sa largeur sur 100 % du conteneur. Tu dois veiller à ce que l'élément `.lightbox` commence dans l'angle supérieur gauche du conteneur.
1. `.lightbox` doit avoir une couleur d'arrière-plan. Initialement, sa propriété `display` doit être définie sur `none` pour le masquer.
1. Lorsque tu cliques sur l'un de tes éléments `.gallery-item`, la propriété `display` de l'élément `.lightbox` doit être définie sur `flex` afin de rendre visibles l'élément `.lightbox` et les deux éléments qu'il contient.
1. Lorsque tu cliques sur l'un de tes éléments `.gallery-item`, l'attribut `src` de l'élément `#lightbox-image` doit être défini sur la version en taille réelle de l'image sélectionnée en retirant `-thumbnail` de l'attribut `src` de l'image. Les images en taille réelle se trouvent aux liens suivants :

    - `https://cdn.freecodecamp.org/curriculum/labs/stonehenge.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/storm.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/trees.jpg`

1. Lorsque ton élément `.lightbox` est visible et que tu cliques sur `#close-btn` ou sur l'élément `.lightbox`, la propriété `display` de l'élément `.lightbox` doit de nouveau être définie sur `none`.

**Remarque :** Veille à lier ta feuille de style et ton fichier JavaScript dans ton HTML.

# --hints--

Tu dois avoir un élément `div` avec la classe `gallery` dans ton élément `body`.

```js
assert.equal(document.querySelector("body .gallery")?.tagName, "DIV");
```

Dans l'élément `.gallery`, tu dois avoir trois éléments `img` avec la classe `gallery-item`.

```js
assert.lengthOf(document.querySelectorAll(".gallery .gallery-item"), 3)
```

Dans l'élément `.gallery`, tu dois avoir un élément `img` dont l'attribut `src` est défini sur `https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg`.

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0;
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) {
    srcCount++;
  }
}

assert.strictEqual(srcCount, 1);
```

Dans l'élément `.gallery`, tu dois avoir un élément `img` dont l'attribut `src` est défini sur `https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg`.

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0;
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) {
    srcCount++;
  }
}

assert.strictEqual(srcCount, 1);
```

Dans l'élément `.gallery`, tu dois avoir un élément `img` dont l'attribut `src` est défini sur `https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg`.

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0;
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) {
    srcCount++;
  }
}

assert.strictEqual(srcCount, 1);
```

Tu dois avoir un élément `div` avec la classe `lightbox` dans ton élément `body`.

```js
const lightboxEl = document.querySelector("body .lightbox");
assert.exists(lightboxEl);
assert.equal(lightboxEl?.tagName, "DIV");
```

Dans ton élément `.lightbox`, tu dois avoir un élément `button` dont l'`id` est défini sur `close-btn`.

```js
assert.equal(document.querySelector(".lightbox #close-btn")?.tagName, "BUTTON");
```

Dans ton élément `.lightbox`, tu dois avoir un élément `img` dont l'`id` est défini sur `lightbox-image`.

```js
assert.equal(document.querySelector(".lightbox #lightbox-image")?.tagName, "IMG");
```

Ton élément `.lightbox` doit avoir une position fixe.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.position, "fixed");
```

Ton élément `.lightbox` doit couvrir toute la zone d'affichage.

```js
const style = new __helpers.CSSHelp(document).getStyle(".lightbox");

const width = style?.width;
const height = style?.height;

assert.oneOf(
  width,
  ["100%", "100vw"]
);

assert.oneOf(
  height,
  ["100%", "100vh"]
);
```

Ton élément `.lightbox` doit être aligné avec l'angle supérieur gauche du conteneur.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.top, "0px");
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.left, "0px");
```

Ton élément `.lightbox` doit avoir une couleur d'arrière-plan.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".lightbox")?.backgroundColor);
```

Ton élément `.lightbox` doit être masqué initialement.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.display, "none");
```

Lorsque tu cliques sur l'un de tes éléments `.gallery-item`, la propriété `display` de l'élément `.lightbox` doit être définie sur `flex` afin de rendre visibles `.lightbox` et les deux éléments qu'il contient.

```js
const lightbox = document.querySelector(".lightbox");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

assert.strictEqual(getComputedDisplay(lightbox), "none");

const galleryItem = document.querySelector(".gallery-item");
galleryItem.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "flex");
```

Lorsque tu cliques sur l'un de tes éléments `.gallery-item`, l'attribut `src` de l'élément `#lightbox-image` doit être défini sur la version en taille réelle de l'image sélectionnée en retirant `-thumbnail` de l'attribut `src` de l'image.

```js
const galleryItems = document.querySelectorAll(".gallery-item");
const lightboxImage = document.getElementById("lightbox-image");

assert.isAbove(galleryItems.length, 0);

function getFullSizeImageUrl(thumbnailUrl) {
  return thumbnailUrl.replace("-thumbnail", "");
}

galleryItems.forEach((item) => {
  const thumbnailUrl = item.src;
  const expectedFullSizeUrl = getFullSizeImageUrl(thumbnailUrl);

  item.dispatchEvent(new MouseEvent("click", { bubbles: true }));

  assert.strictEqual(lightboxImage.src, expectedFullSizeUrl);
});
```

Lorsque ton élément `.lightbox` est visible et que tu cliques sur le bouton `#close-btn`, la propriété `display` de l'élément `.lightbox` doit de nouveau être définie sur `none`.

```js
const lightbox = document.querySelector(".lightbox");
const closeBtn = document.getElementById("close-btn");
const galleryItem = document.querySelector(".gallery-item");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

galleryItem.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "flex");

closeBtn.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "none");
```

Lorsque ton élément `.lightbox` est visible et que tu cliques sur l'élément `.lightbox`, la propriété `display` de l'élément `.lightbox` doit de nouveau être définie sur `none`.

```js
const lightbox = document.querySelector(".lightbox");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

const galleryItem = document.querySelector(".gallery-item");
galleryItem.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "flex");

lightbox.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "none");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Lightbox Viewer</title>
  </head>

  <body>

  </body>

</html>
```

```css

```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Lightbox Viewer</title>
    <link rel="stylesheet" href="styles.css">
  </head>

  <body>
    <h1 class="title">Museum Gallery</h1>

    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg" alt="Stonehenge landscape with large stone structures." class="gallery-item">
      <img src="https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg" alt="A peaceful forest with tall green trees." class="gallery-item">
      <img src="https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg" alt="A dramatic stormy sky with dark clouds and lightning." class="gallery-item">
    </div>
    <div id="lightbox" class="lightbox">
      <button id="close-btn" class="close-btn">&times;</button>
      <img id="lightbox-image" src="" alt="Lightbox Image">
    </div>
    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  background-color: #f0f0f0;
  color: #333;
}

.title {
  margin-top: 100px;
  text-align: center;
  background-color: rgb(211, 221, 224);
}

.gallery {
  display: flex;
  flex-wrap: wrap;
  gap: 10px;
  justify-content: center;
  padding: 20px;
}

.gallery-item {
  width: 250px;
  height: 300px;
  object-fit: cover;
  cursor: pointer;
  transition: transform 0.2s;
}

.gallery-item:hover {
  transform: scale(1.05);
}

.lightbox {
  display: none;
  position: fixed;
  top: 0;
  left: 0;
  width: 100%;
  height: 100%;
  background-color: rgba(0, 0, 0, 0.8);
  justify-content: center;
  align-items: center;
  z-index: 1000;
}

.lightbox img {
  max-width: 90%;
  max-height: 90%;
}

.close-btn {
  position: absolute;
  top: 20px;
  right: 10px;
  font-size: 2em;
  color: white;
  cursor: pointer;
  border: none;
  background-color: transparent;
}
```

```js
const lightbox = document.getElementById("lightbox");
const lightboxImage = document.getElementById("lightbox-image");

function openLightbox(src) {
  lightboxImage.src = src;
  lightbox.style.display = "flex";
}

function closeLightbox() {
  lightbox.style.display = "none";
}

const closeButton = document.getElementById("close-btn");

closeButton.addEventListener("click", closeLightbox);

lightbox.addEventListener("click", (e) => {
  if (e.target === lightbox) {
    closeLightbox();
  }
});

const galleryItems = document.querySelectorAll(".gallery-item");

galleryItems.forEach((item) => {
  item.addEventListener("click", () => {
    openLightbox(item.src.replace("-thumbnail", ""));
  });
});
```
