---
id: 67ef8f695db52583424bd118
title: Concevoir une page de critique de film
challengeType: 25
dashedName: design-a-movie-review-page
demoType: onClick
---

# --description--

**Objectif :** Réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab. 

**User Stories :**

1. Tu devrais avoir un élément `main`.
2. À l'intérieur de l'élément `main`, tu devrais avoir un élément `h1` pour le titre du film.
3. Sous l'élément `h1`, tu devrais avoir un élément `img` qui affiche la jaquette du film. Ton élément `img` doit avoir un texte `alt` descriptif qui décrit l'image. Tu peux utiliser l'image suivante si tu veux : `https://cdn.freecodecamp.org/curriculum/labs/rise-beyond-2.png`.
4. Tu devrais avoir un élément `p` qui contient une brève description du film.
5. Tu devrais avoir un autre élément `p` pour afficher la note du film. À l'intérieur, tu devrais avoir ces éléments dans l'ordre indiqué :
   - Un élément `strong` avec le texte `Movie Rating:`.
   - Un élément `span` avec un attribut `aria-hidden` mis à `true` qui contient une représentation visuelle de la note avec des étoiles `⭐⭐⭐⭐⭐⭐⭐⭐⭐☆`.
   - Une valeur numérique, qui représente la note, entre parenthèses (ex. `9.2/10`) après le span.
6. Tu devrais avoir un élément `h2` avec le texte `Cast Members`.
7. Tu devrais avoir un élément `ul`.
8. À l'intérieur de l'élément `ul`, tu devrais avoir plusieurs éléments `li` qui contiennent chacun un élément `strong` pour le nom de l'acteur, suivi du nom du personnage correspondant précédé du texte `as`. (ex. `James Holloway as Ethan Carter`).

# --hints--

Tu devrais avoir un élément `main`.

```js
assert.isNotNull(document.querySelector("main"));
```

À l'intérieur de l'élément `main`, tu devrais avoir un élément `h1` qui contient le titre du film.

```js
const h1 = document.querySelector("main > h1");
assert.isAbove(h1?.innerText.trim().length, 0);
```

Sous l'élément `h1`, tu devrais avoir un élément `img` avec un attribut `alt` descriptif.
  
```js
const img = document.querySelector("main > h1 + img");
assert.isNotNull(img);
assert.isNotEmpty(img.getAttribute("alt"));
```

Tu devrais avoir un élément `p` après l'image qui contient la description du film.

```js
const description = document.querySelector("main img + p");
assert.isNotNull(description);
assert.isNotEmpty(description.innerText.trim());
```

Tu devrais avoir un autre élément `p` qui contient la note.

```js
const paragraphs = document.querySelectorAll("main p");
assert.isAtLeast(paragraphs.length, 2);
```

À l'intérieur du second élément `p`, tu devrais avoir un élément `strong` avec le texte `Movie Rating: `.

```js
const strongEl = document.querySelector("main p:nth-of-type(2) strong");
assert.match(strongEl?.innerText.trim(), /^Movie Rating:?$/);
```

Il devrait y avoir un élément `span` à l'intérieur du paragraphe de note.

```js
const spanEl = document.querySelector("main p:nth-of-type(2) span");
assert.isNotNull(spanEl);
```

L'élément `span` à l'intérieur du paragraphe de note doit contenir dix étoiles, soit pleines (⭐) soit vides (☆), suivies d'une valeur numérique entre parenthèses après le `span`.

```js  
const spanEl = document.querySelector("main p:nth-of-type(2) span");
const fullRatingText = document.querySelector("main p:nth-of-type(2)")?.textContent.replace(/\s+/g, ' ').trim();

assert.match(spanEl?.innerText.trim(), /^[⭐☆]{10}$/);
assert.match(fullRatingText, /Movie Rating:?\s*[⭐☆]{10}\s*\(\s*\d+(?:\.\d+)?\s*\/\s*10\s*\)$/);
```

L'élément `span` à l'intérieur du paragraphe de note devrait avoir un attribut `aria-hidden` mis à `true`.

```js
const spanEl = document.querySelector("main p:nth-of-type(2) span");
assert.strictEqual(spanEl?.getAttribute("aria-hidden"), "true");
```

Tu devrais avoir un élément `h2`.

```js
assert.exists(document.querySelector("h2"));
```

Ton `h2` doit venir après ton second élément `p`.

```js
const allParagraphs = document.querySelectorAll("p");
assert.isAtLeast(allParagraphs.length, 2);

const h2 = document.querySelector("h2");
const main = document.querySelector("main");

assert.exists(h2);
assert.exists(main);

const elements = Array.from(main.querySelectorAll("*"));
const lastPIndex = elements.indexOf(allParagraphs[1]);
const h2Index = elements.indexOf(h2);

assert.isAbove(h2Index, lastPIndex);
```

Ton `h2` devrait avoir le texte `Cast Members`.

```js
const h2 = document.querySelector("h2");
assert.equal(h2?.textContent.trim(), "Cast Members");
```

Tu devrais avoir un élément `ul` après le `h2`.

```js
const ul = document.querySelector("main h2 + ul");
assert.isNotNull(ul);
```

L'élément `ul` doit contenir plusieurs éléments `li`.

```js
const liItems = document.querySelectorAll("main ul li");
assert.isAbove(liItems.length, 1);
```

Chaque `li` doit contenir un élément `strong` pour le nom de l'acteur, suivi du nom du personnage correspondant précédé du texte `as`. (ex. `James Holloway as Ethan Carter`).

```js
const liItems = document.querySelectorAll("main ul li");  
assert.isNotEmpty(liItems);
liItems.forEach(li => {
  const strong = li.querySelector("strong");
  assert.isNotNull(strong);
  const fullText = li.textContent.trim();
  const actorName = strong.textContent.trim();
  assert.isTrue(fullText.startsWith(actorName));
  const characterPart = fullText.slice(actorName.length).trim();
  assert.match(characterPart, /^as\s.+/);
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
 <html lang="en">
 <head>
     <meta charset="UTF-8">
     <meta name="viewport" content="width=device-width, initial-scale=1.0">
     <title>Movie Review</title>
 </head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Movie Review</title>
  </head>
  <body>
    <main class="review">
      <h1>Movie Review: Rise Beyond</h1>
      <img
        src="https://cdn.freecodecamp.org/curriculum/labs/rise-beyond-2.png"
        alt="The Rise Beyond movie poster, starring Ethan Carter. A portrait of a young boy standing on a stage with his back toward the audience, holding a guitar at his side, staring at a backdrop of a cityscape with the sun rising behind its tall buildings."
      />
      <p>
        Rise Beyond is an uplifting drama that tells the story of Ethan Carter,
        a young man from a small town who dreams of becoming a world-class
        musician. Facing financial struggles and self-doubt, Ethan embarks on a
        journey of resilience, guided by his mentor Professor Adams and
        childhood friend Lena Mitchell. Through setbacks, failures, and moments
        of triumph, Rise Beyond reminds us that our greatest potential lies
        beyond our fears.
      </p>
      <p>
        <strong>Movie Rating:</strong>
        <span aria-hidden="true">⭐⭐⭐⭐⭐⭐⭐⭐⭐☆</span> (9.2/10)
      </p>

      <h2>Cast Members</h2>
      <ul>
        <li><strong>James Holloway</strong> as Ethan Carter</li>
        <li><strong>Olivia Sterling</strong> as Lena Mitchell</li>
        <li><strong>William Lancaster</strong> as Professor Adams</li>
        <li>
          <strong>Maria Collins</strong> as Evelyn Carter (Ethan’s Mother)
        </li>
      </ul>
    </main>
  </body>
</html>
```
