# Handoff Traductions FR вЂ” freeCodeCamp Local Fork

Ce fichier contient toutes les informations nГ©cessaires pour continuer le travail de traduction du curriculum freeCodeCamp en franГ§ais dans une nouvelle session Claude.

## Contexte Projet

- **Repo local** : `c:\Users\Erazer\.vscode\code\Nouveau dossier\freeCodeCamp`
- **Remote** : `https://github.com/MATTEO12SA/freecodecamp-fr-local.git` (alias `standalone`, branche `main`)
- **Objectif** : Traduire le superblock `responsive-web-design-v9` (cert RWD v9) du curriculum freeCodeCamp en franГ§ais.
- **Source EN** : `curriculum/challenges/english/blocks/<bloc>/<id>.md`
- **Cible FR** : `curriculum/i18n-curriculum/curriculum/challenges/french/blocks/<bloc>/<id>.md` (mГЄme `id`, mГЄme nom de fichier)

## Г‰tat Actuel вЂ” Ce Qui Est Fait

### Chapitre HTML

- вњ… 100% traduit (modules basic-html, semantic-html, html-forms-and-tables, lab-survey-form, html-and-accessibility, review-html)

### Chapitre Computers

- вњ… Module `computer-basics` вЂ” 16 fichiers (5 blocs : 3 lectures + review + quiz)

### Chapitre CSS (en cours)

- вњ… Module `basic-css` вЂ” 40 fichiers complets + 2 cafe-menu manquants
- вњ… Module `design-for-developers` вЂ” 23 fichiers complets
- вњ… Module `absolute-and-relative-units` вЂ” 8 fichiers complets
- вњ… Module `pseudo-classes-and-elements` вЂ” 10 fichiers pГ©dagogiques (lectures + lab + review + quiz). **Workshops non traduits** : `workshop-greeting-card` (27) + `workshop-parent-teacher-conference-form` (37) = 64 fichiers restants
- вњ… Module `css-colors` вЂ” 9 fichiers pГ©dagogiques. **Workshop non traduit** : `workshop-colored-markers` (89 fichiers)
- вњ… Module `styling-forms` вЂ” 7 fichiers pГ©dagogiques. **Workshops non traduits** : `workshop-registration-form` (61) + `workshop-game-settings-panel` (16) = 77 fichiers
- вњ… Module `css-box-model` вЂ” 10 fichiers pГ©dagogiques complets (7 lectures + lab + review + quiz). **Workshop non traduit** : `workshop-rothko-painting` (44)
- вњ… Module `css-flexbox` вЂ” 5 fichiers pГ©dagogiques complets (2 lectures + lab + review + quiz). **Workshops non traduits** : `workshop-flexbox-photo-gallery` (22) + `workshop-colorful-boxes` (43)
- вњ… Module `lab-page-of-playing-cards` вЂ” 1 lab complet
- вњ… Module `css-typography` вЂ” 10 fichiers pГ©dagogiques complets (7 lectures + lab + review + quiz). **Workshop non traduit** : `workshop-nutritional-label` (68)
- вњ… Module `css-and-accessibility` вЂ” 5 fichiers pГ©dagogiques complets (2 lectures + lab + review + quiz). **Workshop non traduit** : `workshop-accessibility-quiz` (67)
- вњ… Module `css-positioning` вЂ” 8 fichiers pГ©dagogiques complets (5 lectures + lab + review + quiz). **Workshop non traduit** : `workshop-cat-painting` (80)
- вњ… Module `attribute-selectors` вЂ” 5 fichiers pГ©dagogiques complets (3 lectures + review + quiz). **Workshop non traduit** : `workshop-balance-sheet` (66)

### Infrastructure

- вњ… Fix watcher `fs.watchFile` (commit `6163ae508f`) вЂ” hot-reload des `.md` FR sur Windows
- вњ… Fix watcher `fs.watch` recursive (commit `227a48aada`) вЂ” dГ©tecte les nouveaux `.md` crГ©Г©s aprГЁs dГ©marrage Gatsby
- вњ… Fix rГ©gression `writeFileSync` dans `build-external-curricula-data-v2.ts`
- вњ… Fix doublons titres `Г‰tape 60`/`Г‰tape 76` dans `workshop-cafe-menu`
- вњ… Fix titres `workshop-cafe-menu` alignГ©s sur les `dashedName: step-X` (supprime les doublons `Г‰tape 52`/`Г‰tape 70` pendant `create:external-curriculum`)
- вњ… Fix `create:external-curriculum` stable avec Gatsby dev server : le gГ©nГ©rateur ne rГ©Г©crit plus les JSON inchangГ©s, ce qui Г©vite les crashs `ENOENT ... chmod client/public/curriculum-data/...`
- вњ… Logs `intro.integrating` / `intro.integrated` dans `dev-logs/latest.log` pour confirmer l'intГ©gration de `intro.json` dans les `curriculum-data`
- вњ… Watcher `intro.json` dans `dev.ps1` : un edit direct Г©crit `intro.changed` puis `intro.integrated` dans `latest.log` avec `sourceJson`, `curriculumData` et `serverPath=/learn/responsive-web-design-v9/`
- вњ… Fix hot-reload des titres `/learn` : `client/i18n/config.js` garde les fichiers i18n FR dans le graphe Webpack au lieu de les figer via `preval`
- вњ… `client/i18n/locales/french/intro.json` Г  jour pour tous les modules traduits (titres + intros des blocs ET du module)
- вњ… Curriculum-data rГ©gГ©nГ©rГ© (Г  refaire aprГЁs chaque modif d'`intro.json`)
- вњ… `DOCS-FR.md` et `QUICKSTART.md` documentent le fix `fs.watch` recursive

## Г‰tat Actuel вЂ” Ce Qui Reste ГЂ Faire

### Chapitre CSS вЂ” Modules restants (dans l'ordre du superblock)

| #                                         | Module                                     | Lectures                       | Workshops Г  skip                                          | Lab/Review/Quiz            | Total pГ©dagogique |
| ----------------------------------------- | ------------------------------------------ | ------------------------------ | --------------------------------------------------------- | -------------------------- | ----------------- |
| 1                                         | `lab-book-inventory-app`                   | вЂ”                              | вЂ”                                                         | 1 lab                      | 1                 |
| 2                                         | `responsive-design`                        | 4                              | workshop-piano (31)                                       | 1 review + 1 quiz          | 6                 |
| 3                                         | `lab-technical-documentation-page`         | вЂ”                              | вЂ”                                                         | 1 lab                      | 1                 |
| 4                                         | `css-variables`                            | 2                              | workshop-city-skyline (115)                               | 1 lab + 1 review + 1 quiz  | 5                 |
| 5                                         | `css-grid`                                 | 8 (dont lecture-debugging-css) | workshop-magazine (79)                                    | 1 lab + 1 review + 1 quiz  | 11                |
| 6                                         | `lab-product-landing-page`                 | вЂ”                              | вЂ”                                                         | 1 lab                      | 1                 |
| 7                                         | `css-animations`                           | 2                              | workshop-ferris-wheel (29), workshop-flappy-penguin (104) | 2 labs + 1 review + 1 quiz | 6                 |
| 8                                         | `review-css`                               | вЂ”                              | вЂ”                                                         | 1 review                   | 1                 |
| 9                                         | `exam-responsive-web-design-certification` | вЂ”                              | вЂ”                                                         | 1 exam                     | 1                 |
| **TOTAL contenu pГ©dagogique**             |                                            |                                |                                                           | **33 fichiers**            |
| **TOTAL workshops Г  skip pour l'instant** |                                            |                                |                                                           | **~407 fichiers**          |

**PrioritГ© : faire d'abord les 33 fichiers de contenu pГ©dagogique restants, puis attaquer les workshops si l'utilisateur le demande.**

## Pattern De Traduction (RГЁgles Strictes)

### ГЂ traduire (prose franГ§aise)

- `title:` (frontmatter)
- Sections `# --description--`, `# --interactive--`, `# --questions--`, `# --assignment--`
- `## --text--`, `## --answers--`, `### --feedback--`, `#### --text--`, `#### --distractors--`, `#### --answer--`
- Messages d'erreur dans les blocs ` ```js ` qui sont des chaГ®nes texte affichГ©es Г  l'utilisateur final (rare, vГ©rifier le contexte)

### ГЂ NE JAMAIS toucher (verbatim, copier-coller du EN)

- `id:`, `challengeType:`, `dashedName:`, `videoId:`, `demoType:`, `blockType:`, etc. (frontmatter technique)
- Marqueurs de section : `# --description--`, `## --quiz--`, `## --video-solution--`, `# --hints--`, `# --seed--`, `## --seed-contents--`, `# --solutions--`, etc.
- Tout code dans des blocs ` ```html `, ` ```js `, ` ```css ` (sauf strings UI)
- SГ©lecteurs CSS, `assert(...)`, `document.querySelector(...)`, regex
- Backticks inline contenant du code/attributs (ex: `` `<h1>` ``, `` `class` ``, `` `:hover` ``, `` `300px` ``)
- URLs (sauf si la doc l'autorise pour `cdn.freecodecamp.org/curriculum/lecture-transcripts/<slug>-fr.png` mais pour l'instant on garde les URLs EN)

### Style FR

- **Tutoiement** systГ©matique (В« tu В» jamais В« vous В»)
- Lexique technique cohГ©rent avec les blocs prГ©cГ©dents :
  - HTML в†’ В« HTML В», CSS в†’ В« CSS В»
  - element в†’ В« Г©lГ©ment В»
  - selector в†’ В« sГ©lecteur В»
  - property в†’ В« propriГ©tГ© В»
  - value в†’ В« valeur В»
  - browser в†’ В« navigateur В»
  - file в†’ В« fichier В»
  - background-color, margin, padding, etc. в†’ garder en anglais avec backticks
  - Hue в†’ В« teinte В», Saturation в†’ В« saturation В», Lightness в†’ В« luminositГ© В»
  - light в†’ В« clair В», dark в†’ В« sombre В»
  - hover в†’ В« survol В» ou garder `:hover`
- Conserver les guillemets `В«` `В»` franГ§ais pour les citations longues, et `"` pour le code

### Cas particulier вЂ” Frontmatter avec deux-points

Si le titre EN contient `:` (ex: `title: When Should You Use appearance: none...`), il faut entourer le titre FR de guillemets doubles :

```yaml
title: 'Quand devrais-tu utiliser appearance: none pour...'
```

Sinon le parser YAML casse.

## Workflow Type Par Module

```bash
# 1. Lister les fichiers du module
ls curriculum/challenges/english/blocks/<bloc>/

# 2. Lire les fichiers EN (4 max en parallГЁle, pour prГ©server le contexte)
# (Read tool sur 4 fichiers en parallГЁle)

# 3. Г‰crire les versions FR (Write tool sur 4 fichiers en parallГЁle)
# Cible : curriculum/i18n-curriculum/curriculum/challenges/french/blocks/<bloc>/<id>.md

# 4. Commit + push Г  la fin du module
git add curriculum/i18n-curriculum/curriculum/challenges/french/blocks/<bloc>/
git commit -m "translate <bloc-slug> module"
git push standalone main
```

**Note** : si tu modifies plusieurs blocs dans le mГЄme module, fais un seul commit groupГ© Г  la fin.

## PiГЁges Connus (Ne Pas Refaire)

1. **Doublons de titres** : avant d'Г©crire un nouveau `Г‰tape X` dans un workshop, vГ©rifier que le titre n'existe pas dГ©jГ  (`grep -l "^title: Г‰tape X$" curriculum/i18n-curriculum/.../blocks/<workshop>/*.md`). Si oui, c'est probablement un dashedName diffГ©rent et il faut corriger l'autre fichier (titre doit matcher le `dashedName: step-X`).

2. **RГ©gГ©nГ©ration curriculum-data** : aprГЁs chaque modif d'`intro.json`, exГ©cuter :

   ```powershell
   $env:CURRICULUM_LOCALE='french'; $env:CLIENT_LOCALE='french'
   pnpm -C curriculum build
   pnpm -C client create:external-curriculum
   ```

   Le gГ©nГ©rateur Г©crit `intro.integrating` puis `intro.integrated` dans `dev-logs/latest.log`. Sinon les titres des blocs sur `/cours-fr` restent en anglais (ils sont mis en cache dans `client/static/curriculum-data/v2/responsive-web-design-v9.json`). Sur `/learn`, `client/i18n/config.js` permet maintenant Г  Gatsby de voir les changements `intro.json` et `dev.ps1` Г©crit `intro.changed` puis `intro.integrated` quand le fichier est sauvegardГ© ; si un onglet garde l'ancien texte, faire `Ctrl + Shift + R`.

3. **Pre-push hook prettier** : peut bloquer le push si un fichier `.md` a un problГЁme de formatage. Corriger avec :

   ```powershell
   npx prettier --write <file>.md
   git add <file>.md
   git commit --amend --no-edit
   git push standalone main
   ```

4. **DOCS-FR.md sans accent dans le code** : ne pas mettre d'accents dans les noms de variables/commandes des exemples de doc.

5. **PowerShell quoting** : les paths contiennent un espace (`Nouveau dossier`). Toujours utiliser des chemins absolus entre guillemets ou se positionner dans le repo d'abord.

## Commandes Utiles

### Dev server

```powershell
.\dev.ps1 -Fast            # relance rapide Gatsby
.\dev.ps1                  # relance complГЁte avec turbo setup
.\dev.ps1 -Clean           # avec wipe cache Gatsby
```

### Statut serveur

```powershell
Get-Content dev-logs\status.json   # voir si UP/STARTING/DOWN/ERROR
Get-Content dev-logs\latest.log -Tail 50
Get-Content dev-logs\latest.log -Wait | Select-String -Pattern "status.up|status.error|watcher.|challenge.integrating|challenge.integrated|challenge.error|intro.changed|intro.integrating|intro.integrated"
```

### VГ©rifier un push

```bash
git log --oneline -5
```

### Lister ce qui manque dans un module

```bash
for b in <bloc1> <bloc2> ...; do
  en=$(ls "curriculum/challenges/english/blocks/$b" 2>/dev/null | wc -l)
  fr=$(ls "curriculum/i18n-curriculum/curriculum/challenges/french/blocks/$b" 2>/dev/null | wc -l)
  echo "$b: EN=$en FR=$fr MISSING=$((en - fr))"
done
```

## MГ©moire Utilisateur (Important)

- **В« Dis oui tout le temps В»** : enchaГ®ner les opГ©rations sans demander confirmation.
- L'utilisateur veut le maximum de fichiers traduits par session, en priorisant le contenu pГ©dagogique (lectures + reviews + quizzes + labs) sur les workshops massifs.
- Toujours commit + push Г  la fin de chaque module (pas juste commit local).
- Tutoiement systГ©matique dans les traductions.

## Comment DГ©marrer La Prochaine Session

1. Lire ce fichier (`HANDOFF-TRADUCTIONS.md`) en premier
2. Continuer par le module `lab-book-inventory-app` (1 lab)
3. Puis enchaГ®ner les modules dans l'ordre du tableau ci-dessus
4. ГЂ la fin de chaque module : commit + push + mettre Г  jour `intro.json` si le titre du bloc/module est encore en anglais
5. ГЂ la fin de tous les modules pГ©dagogiques : rГ©gГ©nГ©rer la curriculum-data (commande ci-dessus)

## Fichier De Structure Du Superblock

Pour vГ©rifier l'ordre exact des blocs/modules :
`curriculum/structure/superblocks/responsive-web-design-v9.json`

## Contact Avec Le Hot-Reload

Tu peux modifier n'importe quel `.md` FR et il sera hot-reloadГ© en ~5s dans le navigateur (Ctrl + Shift + R). Si tu crГ©es un nouveau `.md` (nouveau bloc ou nouveau fichier dans un bloc existant), le `fs.watch` recursive le dГ©tecte automatiquement sans redГ©marrer le serveur.

Pour vГ©rifier qu'un edit a bien Г©tГ© pris en compte :

```powershell
Select-String -Path dev-logs\latest.log -Pattern "watcher.|challenge.integrating|challenge.integrated|challenge.error|intro.changed|intro.integrating|intro.integrated" | Select-Object -Last 10
```

Pour surveiller en direct : `Get-Content dev-logs\latest.log -Wait | Select-String -Pattern "status.up|status.error|watcher.|challenge.integrating|challenge.integrated|challenge.error|intro.changed|intro.integrating|intro.integrated"`. `status.up` = serveur prГЄt ; `watcher.changed` / `watcher.added` = modification ou nouveau `.md` dГ©tectГ© ; `challenge.integrating` / `challenge.integrated` = Gatsby relance puis termine l'intГ©gration de la page ; `intro.changed` / `intro.integrated` = edit direct de `intro.json` vu par Gatsby ; `intro.integrating` / `intro.integrated` = `intro.json` repris dans les donnГ©es statiques.

---

**DerniГЁre session** : module pГ©dagogique `attribute-selectors` traduit (3 lectures + review + quiz), `intro.json` mis Г  jour, curriculum-data rГ©gГ©nГ©rГ©, logs `intro.*` ajoutГ©s et affichage `/learn` vГ©rifiГ© en franГ§ais (`SГ©lecteurs d'attributs`). Prochain module : `lab-book-inventory-app`. Total fichiers FR actuellement dans le repo : ~728+ sur ~1700 dans la cert RWD v9.
