---
id: 68e4276483b0c71658aab1ca
title: Étape 1
challengeType: 1
dashedName: step-1
---

# --description--

Dans ce workshop, tu vas t'entraîner à utiliser différentes méthodes de chaîne servant à formater des chaînes.

Commence par créer une variable appelée `userInput` et affecte-lui la valeur suivante : `"   Hello World!   "`.

# --hints--

Tu devrais avoir une variable `userInput`.

```js
assert.exists(userInput);
```

Ta variable `userInput` devrait être une chaîne.

```js
assert.isString(userInput);
```

Tu devrais affecter `"   Hello World!   "` à ta variable `userInput`. Assure-toi qu'il y a trois espaces avant le `H` et trois espaces après le `!`.

```js
assert.equal(userInput, "   Hello World!   ");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
