---
id: 68e4276483b0c71658aab1ce
title: Étape 6
challengeType: 1
dashedName: step-6
---

# --description--

Il y aura des moments où tu voudras supprimer les espaces uniquement au début de la chaîne. C'est là que la méthode `trimStart()` est utile.

Voici un exemple :

```js
const sentence = "  I love learning.  ";
sentence.trimStart(); // "I love learning.  "
```

Souviens-toi que les chaînes sont immutables. Cette méthode crée une nouvelle chaîne avec le résultat. Elle ne modifie pas l'originale.

Crée une variable `trimmedStart` et affecte-lui le résultat de l'utilisation de la méthode `trimStart()` sur la variable `userInput`.

# --hints--

Tu devrais avoir une variable `trimmedStart`.

```js
assert.exists(trimmedStart);
```

Ta variable `trimmedStart` devrait être une chaîne.

```js
assert.isString(trimmedStart);
```

Tu devrais affecter `userInput.trimStart()` à ta variable `trimmedStart`.

```js
assert.equal(trimmedStart, userInput.trimStart());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);
--fcc-editable-region--

--fcc-editable-region--
```
