---
id: 68e4276483b0c71658aab1d2
title: Étape 13
challengeType: 1
dashedName: step-13
---

# --description--

Il y aura des moments où tu auras besoin d'utiliser plusieurs méthodes pour créer une nouvelle chaîne. Au cours des prochaines étapes, tu utiliseras certaines des méthodes que tu as apprises jusqu'à présent pour créer une nouvelle chaîne qui suit la convention de nommage camelCase.

Commence par créer une variable appelée `lowercaseWord` et affecte-lui la chaîne `"camelcase"`.

Sous cette variable `lowercaseWord`, crée une autre variable appelée `camelCasedVersion` et affecte-lui une chaîne vide pour l'instant.

Sous cette variable `camelCasedVersion`, ajoute un `console.log()` avec la chaîne `"Camel cased version:"` suivi d'un autre `console.log()` pour la variable `camelCasedVersion`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Tu devrais avoir une variable appelée `lowercaseWord`.

```js
assert.exists(lowercaseWord);
```

Ta variable `lowercaseWord` devrait être une chaîne.

```js
assert.isString(lowercaseWord);
```

Tu devrais affecter la chaîne `"camelcase"` à ta variable `lowercaseWord`.

```js
assert.equal(lowercaseWord, "camelcase");
```

Tu devrais avoir une variable `camelCasedVersion`.

```js
assert.exists(camelCasedVersion);
```

Ta variable `camelCasedVersion` devrait être une chaîne.

```js
assert.isString(camelCasedVersion);
```

Tu devrais affecter une chaîne vide à ta variable `camelCasedVersion`.

```js
assert.isEmpty(camelCasedVersion);
```

Tu devrais afficher la chaîne `"Camel cased version:"` dans la console.

```js
assert.equal(getLogs()[12], "Camel cased version:");
```

Tu devrais afficher la variable `camelCasedVersion` dans la console.

```js
assert.equal(getLogs()[13], camelCasedVersion);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

--fcc-editable-region--

--fcc-editable-region--
```
