---
id: 68e4276483b0c71658aab1d3
title: Étape 14
challengeType: 1
dashedName: step-14
---

# --description--

Comme tu t'en souviens des leçons et workshops précédents, tu peux extraire des portions d'une chaîne à l'aide de la méthode `slice()`.

Voici un rappel :

```js
const str = "abcd";
str.slice(0, 3); // "abc"
```

Tu devrais modifier l'affectation existante de `camelCasedVersion` en remplaçant la chaîne vide par `lowercaseWord.slice(0, 5)`.

Tu devrais maintenant voir le mot `"camel"` dans la console.

# --hints--

Ta variable `camelCasedVersion` ne devrait pas être une chaîne vide.

```js
assert.isNotEmpty(camelCasedVersion);
```

Tu devrais affecter `lowercaseWord.slice(0, 5)` à la variable `camelCasedVersion`.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5));
```

Tu ne devrais PAS affecter la chaîne littérale `"camel"` à `camelCasedVersion`. Relis les consignes pour la bonne approche.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camel(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = ""
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```
