---
id: 68e4276483b0c71658aab1d0
title: Étape 10
challengeType: 1
dashedName: step-10
---

# --description--

Il y aura des moments où tu auras besoin de créer une nouvelle chaîne avec une casse différente.

Pour créer une nouvelle chaîne dont toutes les lettres sont converties en majuscules, tu peux utiliser la méthode `toUpperCase()`.

Voici un exemple :

```js
const sentence = "I love to code!";
sentence.toUpperCase(); // "I LOVE TO CODE!"
```

Crée une variable appelée `upperCaseInput` et affecte-lui le résultat de l'utilisation de la méthode `toUpperCase()` sur la variable `cleanedInput`.

# --hints--

Tu devrais avoir une variable `upperCaseInput`.

```js
assert.exists(upperCaseInput);
```

Ta variable `upperCaseInput` devrait être une chaîne.

```js
assert.isString(upperCaseInput);
```

Tu devrais affecter `cleanedInput.toUpperCase()` à ta variable `upperCaseInput`.

```js
assert.equal(upperCaseInput, cleanedInput.toUpperCase());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

--fcc-editable-region--

--fcc-editable-region--
```
