---
id: 68e4276483b0c71658aab1d1
title: Étape 12
challengeType: 1
dashedName: step-12
---

# --description--

Si tu as besoin de créer une nouvelle chaîne dont toutes les lettres sont en minuscules, tu peux utiliser la méthode `toLowerCase()`.

Commence par créer une variable appelée `lowerCaseInput` et affecte-lui le résultat de l'utilisation de la méthode `toLowerCase()` sur la variable `cleanedInput`.

Puis affiche la chaîne `"Result of using the toLowerCase() method:"`.

Enfin, ajoute un autre `console.log()` pour la variable `lowerCaseInput`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Tu devrais avoir une variable `lowerCaseInput`.

```js
assert.exists(lowerCaseInput);
```

Ta variable `lowerCaseInput` devrait être une chaîne.

```js
assert.isString(lowerCaseInput);
```

Tu devrais affecter `cleanedInput.toLowerCase()` à ta variable `lowerCaseInput`.

```js
assert.equal(lowerCaseInput, cleanedInput.toLowerCase());
```

Tu devrais afficher la chaîne `"Result of using the toLowerCase() method:"` dans la console.

```js
assert.equal(getLogs()[10], "Result of using the toLowerCase() method:");
```

Tu devrais afficher la variable `lowerCaseInput` dans la console.

```js
assert.equal(getLogs()[11], lowerCaseInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);


const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

--fcc-editable-region--

--fcc-editable-region--
```
