---
id: 68e4276483b0c71658aab1cc
title: Étape 4
challengeType: 1
dashedName: step-4
---

# --description--

Ta valeur `userInput` a quelques espaces en trop aux deux extrémités de la chaîne. Pour créer une nouvelle chaîne sans ces espaces, tu peux utiliser la méthode `trim()`.

Voici un exemple :

```js
const sentence = "  Learning about the trim method.  ";
sentence.trim(); // "Learning about the trim method."
```

Crée une variable nommée `cleanedInput` et affecte-lui le résultat de l'utilisation de la méthode `trim()` sur la variable `userInput`.

# --hints--

Tu devrais avoir une variable `cleanedInput`.

```js
assert.exists(cleanedInput);
```

Ta variable `cleanedInput` devrait être une chaîne.

```js
assert.isString(cleanedInput);
```

Tu devrais affecter `userInput.trim()` à ta variable `cleanedInput`.

```js
assert.equal(cleanedInput, userInput.trim());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);
--fcc-editable-region--

--fcc-editable-region--
```
