---
id: 68e4276483b0c71658aab1cd
title: Étape 5
challengeType: 1
dashedName: step-5
---

# --description--

Il est maintenant temps de voir le résultat.

Commence par ajouter un `console.log()` avec la chaîne `"Result of trimming whitespace from both ends:"`.

Puis sous ce `console.log()`, ajoute un autre `console.log()` avec la variable `cleanedInput`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Tu devrais afficher la chaîne `"Result of trimming whitespace from both ends:"` dans la console.

```js
assert.equal(getLogs()[2], "Result of trimming whitespace from both ends:");
```

Tu devrais afficher la variable `cleanedInput` dans la console.

```js
assert.equal(getLogs()[3], cleanedInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
--fcc-editable-region--

--fcc-editable-region--
```
