---
id: 68ed6079fe0e3e98b94a1b75
title: Étape 2
challengeType: 1
dashedName: step-2
---

# --description--

Ensuite, ajoute un `console.log()` avec la chaîne `"Original input:"`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Tu devrais afficher la chaîne `"Original input:"` dans la console.

```js
assert.equal(getLogs()[0], "Original input:");
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
--fcc-editable-region--

--fcc-editable-region--
```
