---
id: 68f1179eaea3bb59ba434386
title: Étape 8
challengeType: 1
dashedName: step-8
---

# --description--

Pour supprimer les espaces à la fin d'une chaîne, tu peux utiliser la méthode `trimEnd()`.

Voici un exemple :

```js
const sentence = "  I love learning.  ";
sentence.trimEnd(); // "  I love learning."
```

Crée une variable appelée `trimmedEnd` et affecte-lui le résultat de l'utilisation de la méthode `trimEnd()` sur la variable `userInput`.

# --hints--

Tu devrais avoir une variable `trimmedEnd`.

```js
assert.exists(trimmedEnd);
```

Ta variable `trimmedEnd` devrait être une chaîne.

```js
assert.isString(trimmedEnd);
```

Tu devrais affecter `userInput.trimEnd()` à ta variable `trimmedEnd`.

```js
assert.equal(trimmedEnd, userInput.trimEnd());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

--fcc-editable-region--

--fcc-editable-region--
```
