---
id: 68f3f5ca1d2a490c97f14f03
title: Étape 9
challengeType: 1
dashedName: step-9
---

# --description--

Il est maintenant temps de voir le résultat de la suppression des espaces à la fin de la chaîne.

Sous ta variable `trimmedEnd`, ajoute un `console.log()` avec la chaîne `"After using the trimEnd() method, trailing spaces removed:"`.

Puis sous ce `console.log()`, ajoute un autre `console.log()` avec la variable `trimmedEnd`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Tu devrais afficher la chaîne `"After using the trimEnd() method, trailing spaces removed:"` dans la console.

```js
assert.equal(getLogs()[6], "After using the trimEnd() method, trailing spaces removed:");
```

Tu devrais afficher la variable `trimmedEnd` dans la console.

```js
assert.equal(getLogs()[7], trimmedEnd);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
--fcc-editable-region--

--fcc-editable-region--
```
