---
id: 68f3f890010c870e17bbe754
title: Étape 11
challengeType: 1
dashedName: step-11
---

# --description--

Il est maintenant temps de voir le résultat.

Sous ta variable `upperCaseInput`, ajoute un `console.log()` avec la chaîne `"Result of using the toUpperCase() method:"`.

Puis sous ce `console.log()`, ajoute un autre `console.log()` avec la variable `upperCaseInput`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Tu devrais afficher la chaîne `"Result of using the toUpperCase() method:"` dans la console.

```js
assert.equal(getLogs()[8], "Result of using the toUpperCase() method:");
```

Tu devrais afficher la variable `upperCaseInput` dans la console.

```js
assert.equal(getLogs()[9], upperCaseInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
--fcc-editable-region--

--fcc-editable-region--
```
