---
id: 68f3f3d65d51d80ac0b412fc
title: Étape 7
challengeType: 1
dashedName: step-7
---

# --description--

Il est maintenant temps de voir le résultat de la suppression des espaces au début de la chaîne.

Sous ta variable `trimmedStart`, ajoute un `console.log()` avec la chaîne `"After using the trimStart() method, leading spaces removed:"`.

Puis sous ce `console.log()`, ajoute un autre `console.log()` avec la variable `trimmedStart`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Tu devrais afficher la chaîne `"After using the trimStart() method, leading spaces removed:"` dans la console.

```js
assert.equal(getLogs()[4], "After using the trimStart() method, leading spaces removed:");
```

Tu devrais afficher la variable `trimmedStart` dans la console.

```js
assert.equal(getLogs()[5], trimmedStart);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
--fcc-editable-region--

--fcc-editable-region--
```
