---
id: 68f40925d810ae14b2d5a169
title: Étape 16
challengeType: 1
dashedName: step-16
---

# --description--

Pour l'instant, ta variable `camelCasedVersion` contient la valeur `camelC`. Mais le résultat souhaité est `camelCase`.

Pour obtenir les lettres `ase` restantes, tu peux utiliser la méthode `slice()` avec une valeur d'indice négative. Souviens-toi qu'utiliser des valeurs d'indice négatives indique à la méthode `slice()` d'extraire depuis la fin.

Voici un exemple de rappel :

```js
const str = "abcd";
str.slice(-2); // cd
```

Mets à jour l'affectation de `camelCasedVersion` pour concaténer aussi `lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()` avec le résultat de l'utilisation de la méthode `slice()` sur `lowercaseWord`. Passe le nombre `-3` à la méthode `slice()`.

Et avec ce dernier changement, ton workshop est terminé !

# --hints--

Tu devrais concaténer `lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()` avec `lowercaseWord.slice(-3)`.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase() + lowercaseWord.slice(-3));
```

Tu ne devrais PAS affecter la chaîne littérale `"camelCase"` à `camelCasedVersion`. Relis les consignes pour la bonne approche.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camelCase(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```

# --solutions--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";

const camelCasedVersion =
  lowercaseWord.slice(0, 5) +
  lowercaseWord[5].toUpperCase() +
  lowercaseWord.slice(-3);

console.log("Camel cased version:");
console.log(camelCasedVersion);
```
