---
id: 68f405c8c9df44132c8fe1bb
title: Étape 15
challengeType: 1
dashedName: step-15
---

# --description--

Comme tu t'en souviens des leçons précédentes, la convention de nommage camelCase commence par une lettre minuscule pour le premier mot, et chaque mot suivant commence par une lettre majuscule suivie de lettres minuscules.

Voici quelques exemples de rappel :

```md
freeCodeCamp
camelCase
learningIsFun
```

Le deuxième mot dans la variable `lowercaseWord` est `"case"`. Pour accéder au `c` de ce mot, tu peux utiliser `lowercaseWord[5]`.

Mets à jour l'affectation de `camelCasedVersion` pour utiliser l'opérateur `+` afin de concaténer `lowercaseWord.slice(0, 5)` avec le résultat de l'utilisation de la bonne méthode pour convertir les chaînes en majuscules sur `lowercaseWord[5]`.

Tu devrais maintenant voir `camelC` dans la console.

# --hints--

Tu devrais concaténer `lowercaseWord.slice(0, 5)` avec `lowercaseWord[5].toUpperCase()`.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase());
```

Tu ne devrais PAS affecter la chaîne littérale `"camelC"` à `camelCasedVersion`. Relis les consignes pour la bonne approche.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camelC(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = lowercaseWord.slice(0, 5)
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```
