---
id: 685239f0cad43c7a909cbd98
title: Déboguer une page d'adoption d'animaux
challengeType: 25
dashedName: lab-debug-pet-adoption-page
# Le demoType est volontairement retiré pour ne pas afficher d'exemple dans ce projet de débogage.
---

# --description--

Sally, propriétaire d'un magasin d'adoption d'animaux, a construit sa première page web, mais il y a quelques problèmes.

Ton objectif est de corriger toutes les erreurs pour que Sally puisse continuer à construire sa page.

**Objectif :** respecte les user stories ci-dessous et fais passer tous les tests pour terminer le labo.

**User Stories :**

1. Sally veut utiliser une image de chats, mais elle ne s'affiche pas correctement. Tu dois corriger les éléments suivants dans l'élément `img` :
   - Remplace l'attribut `href` par le bon attribut pour indiquer la source de l'image.
   - Remplace l'attribut `att` par le bon attribut qui représente un court texte descriptif pour les images.
   - Supprime la balise fermante `</img>`, car les éléments `img` sont des éléments vides et n'ont pas de balise fermante.
2. Sally veut utiliser des liens pour diriger les utilisateurs vers les pages des chiens et des chats. Mais les liens ne fonctionnent pas correctement. Tu dois corriger les éléments suivants dans les éléments `a` :
   - Remplace les deux attributs `src` par les bons attributs utilisés pour indiquer des URL.

# --hints--

Ton élément `img` doit avoir un attribut `src` au lieu de l'attribut `href`.

```js
const imgEl = document.querySelector("img");
assert.isTrue(imgEl.hasAttribute("src"));
```

Ton élément `img` doit avoir un attribut `alt` au lieu de l'attribut inexistant `att`.

```js
const imgEl = document.querySelector("img");
assert.isTrue(imgEl.hasAttribute("alt"));
```

Ton élément `img` ne doit pas avoir de balise fermante `</img>`.

```js
assert.notMatch(code, /<\/img>/);
```

Ton élément `a` avec le texte `Visit cats page` doit avoir un attribut `href` au lieu d'un attribut `src`.

```js
const anchors = document.querySelectorAll("a");
const catLink = Array.from(anchors).find(a => a.textContent.trim() === "Visit cats page");
assert.isDefined(catLink);
assert.isTrue(catLink.hasAttribute("href"));
```

Ton élément `a` avec le texte `Visit dogs page` doit avoir un attribut `href` au lieu d'un attribut `src`.

```js
const anchors = document.querySelectorAll("a");
const dogLink = Array.from(anchors).find(a => a.textContent.trim() === "Visit dogs page");
assert.isDefined(dogLink);
assert.isTrue(dogLink.hasAttribute("href"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome XYZ Pet Adoption!</h1>
<p>Consider adopting a pet today. We have cats, dogs, rabbits and more.</p>

<h2>See our cats!</h2>
<img href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" att="Two tabby kittens sleeping together on a couch."></img>

<h2>Adopt a cat!</h2>
<a src="/cats">Visit cats page</a>

<h2>Adopt a dog!</h2>
<a src="/dogs">Visit dogs page</a>
```

# --solutions--

```html
<h1>Welcome XYZ Pet Adoption!</h1>
<p>Consider adopting a pet today. We have cats, dogs, rabbits and more.</p>

<h2>See our cats!</h2>
<img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">

<h2>Adopt a cat!</h2>
<a href="/cats">Visit cats page</a>

<h2>Adopt a dog!</h2>
<a href="/dogs">Visit dogs page</a>
```
