---
id: acda2fb1324d9b0fa741e6b5
title: Crée un outil de vérification de fin de chaîne
challengeType: 26
dashedName: implement-a-string-ending-checker-function
---

# --description--

Dans ce lab, tu vas implémenter une fonction qui vérifie si une chaîne se termine par la chaîne cible donnée.

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois créer une fonction nommée `confirmEnding` qui prend deux paramètres : la chaîne à vérifier et la chaîne de référence.
2. La fonction doit retourner `true` si la première chaîne se termine par la deuxième chaîne, et `false` sinon.
3. Tu ne dois pas utiliser la méthode `.endsWith()` ; utilise plutôt l'une des méthodes de sous-chaîne de JavaScript pour y parvenir.

# --hints--

Tu dois créer une fonction nommée `confirmEnding`.

```js
assert.isFunction(confirmEnding);
```

`confirmEnding` doit prendre 2 paramètres.

```js
assert.lengthOf(confirmEnding, 2);
```

`confirmEnding("Bastian", "n")` doit retourner `true`.

```js
assert.isTrue(confirmEnding('Bastian', 'n'));
```

`confirmEnding("Congratulation", "on")` doit retourner `true`.

```js
assert.isTrue(confirmEnding('Congratulation', 'on'));
```

`confirmEnding("Connor", "n")` doit retourner `false`.

```js
assert.isFalse(confirmEnding('Connor', 'n'));
```

`confirmEnding("Walking on water and developing software from a specification are easy if both are frozen", "specification")` doit retourner `false`.

```js
assert.isFalse(
  confirmEnding(
    'Walking on water and developing software from a specification are easy if both are frozen',
    'specification'
  )
);
```

`confirmEnding("He has to give me a new name", "name")` doit retourner `true`.

```js
assert.isTrue(confirmEnding('He has to give me a new name', 'name'));
```

`confirmEnding("Open sesame", "same")` doit retourner `true`.

```js
assert.isTrue(confirmEnding('Open sesame', 'same'));
```

`confirmEnding("Open sesame", "sage")` doit retourner `false`.

```js
assert.isFalse(confirmEnding('Open sesame', 'sage'));
```

`confirmEnding("Open sesame", "game")` doit retourner `false`.

```js
assert.isFalse(confirmEnding('Open sesame', 'game'));
```

`confirmEnding("If you want to save our world, you must hurry. We don't know how much longer we can withstand the nothing", "mountain")` doit retourner `false`.

```js
assert.isFalse(
  confirmEnding(
    "If you want to save our world, you must hurry. We don't know how much longer we can withstand the nothing",
    'mountain'
  )
);
```

`confirmEnding("Abstraction", "action")` doit retourner `true`.

```js
assert.isTrue(confirmEnding('Abstraction', 'action'));
```

Ton code ne doit pas utiliser la méthode intégrée `.endsWith()` pour résoudre le lab.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.endsWith\(.*?\)\s*?;?/);
assert.notMatch(__helpers.removeJSComments(code), /\['endsWith'\]/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function confirmEnding(str, target) {
  return str.substring(str.length - target.length) === target;
}

confirmEnding('Bastian', 'n');
```
