---
id: a77dbc43c33f39daa4429b4f
title: Crée une fonction de vérification de booléen
challengeType: 26
dashedName: build-a-boolean-check-function
---

# --description--

Dans ce lab, tu vas créer une fonction qui vérifie si une valeur est classée comme un booléen primitif.

Les booléens primitifs sont `true` et `false`.

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois avoir une fonction appelée `booWho` qui reçoit un argument.
1. Si l'argument reçu est un booléen primitif, la fonction doit retourner `true`.
1. Si l'argument est toute autre valeur, la fonction doit retourner `false`.

# --hints--

Tu dois avoir une fonction `booWho`.

```js
assert.isFunction(booWho);
```

`booWho(true)` doit retourner `true`.

```js
assert.isTrue(booWho(true));
```

`booWho(false)` doit retourner `true`.

```js
assert.isTrue(booWho(false));
```

`booWho([1, 2, 3])` doit retourner `false`.

```js
assert.isFalse(booWho([1, 2, 3]));
```

`booWho([].slice)` doit retourner `false`.

```js
assert.isFalse(booWho([].slice));
```

`booWho({ "a": 1 })` doit retourner `false`.

```js
assert.isFalse(booWho({ a: 1 }));
```

`booWho(1)` doit retourner `false`.

```js
assert.isFalse(booWho(1));
```

`booWho(NaN)` doit retourner `false`.

```js
assert.isFalse(booWho(NaN));
```

`booWho("a")` doit retourner `false`.

```js
assert.isFalse(booWho('a'));
```

`booWho("true")` doit retourner `false`.

```js
assert.isFalse(booWho('true'));
```

`booWho("false")` doit retourner `false`.

```js
assert.isFalse(booWho('false'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function booWho(bool) {
  return typeof bool === 'boolean';
}

```
