---
id: a5deed1811a43193f9f1c841
title: Implémente un sauteur d'éléments
challengeType: 26
dashedName: implement-an-element-skipper
---

# --description--

Dans ce lab, tu vas créer une fonction qui saute des éléments dans un tableau jusqu'à trouver un élément acceptable selon une fonction de test spécifique.

Par exemple, pour un tableau comme `[1, 1, 1, 2, 1, 1, 1]` et une fonction de test `function(n) {return n === 2}`, le premier élément acceptable est celui à l'indice 3, donc tous les éléments avant doivent être supprimés, et la sortie doit être les éléments restants `[2, 1, 1, 1]`.

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories**

1. Tu dois avoir une fonction `dropElements` qui accepte un tableau (`arr`) et une fonction (`func`) en arguments.
1. La fonction `dropElements` doit parcourir le tableau et supprimer les éléments à partir du premier jusqu'à ce que `func` renvoie `true` pour un élément.
1. La fonction `dropElements` doit renvoyer les éléments restants du tableau si la condition est remplie.
1. Si la condition n'est jamais satisfaite, elle doit renvoyer un tableau vide.

# --hints--

Tu dois avoir une fonction `dropElements`.

```js
assert.isFunction(dropElements);
```

`dropElements([1, 2, 3, 4], function(n) {return n >= 3;})` doit renvoyer `[3, 4]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 4], function(n) {return n >= 3;}), [3, 4]);
```

`dropElements([0, 1, 0, 1], function(n) {return n === 1;})` doit renvoyer `[1, 0, 1]`.

```js
assert.deepEqual(dropElements([0, 1, 0, 1], function(n) {return n === 1;}), [1, 0, 1]);
```

`dropElements([1, 2, 3], function(n) {return n > 0;})` doit renvoyer `[1, 2, 3]`.

```js
assert.deepEqual(dropElements([1, 2, 3], function(n) {return n > 0;}), [1, 2, 3]);
```

`dropElements([1, 2, 3, 4], function(n) {return n > 5;})` doit renvoyer `[]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 4], function(n) {return n > 5;}), []);
```

`dropElements([1, 2, 3, 7, 4], function(n) {return n > 3;})` doit renvoyer `[7, 4]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 7, 4], function(n) {return n > 3;}), [7, 4]);
```

`dropElements([1, 2, 3, 9, 2], function(n) {return n > 2;})` doit renvoyer `[3, 9, 2]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 9, 2], function(n) {return n > 2;}), [3, 9, 2]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dropElements(arr, func) {
  while (arr.length > 0 && !func(arr[0])) {
    arr.shift();
  }
  return arr;
}
```
