---
id: 69757cc0faae0152c1418aad
title: Étape 1
challengeType: 20
dashedName: step-1
---

# --description--

Dans cet atelier, tu vas t'entraîner à travailler avec les nombres et les opérations mathématiques en construisant un répartiteur d'addition. Cet outil calculera ce que chaque personne doit après avoir ajouté le coût des plats et un pourboire.

Pour commencer, tu as besoin d'un moyen de suivre le montant total au fur et à mesure que les coûts sont ajoutés.

Crée une variable nommée `running_total` et assigne-lui la valeur `0`.

# --hints--

Tu devrais avoir une variable nommée `running_total`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_variable('running_total')
    `))
})
```

Tu devrais assigner l'entier `0` à ta variable `running_total`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('running_total').is_equivalent('running_total = 0')
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
