---
id: 6976d8a7cd2a1a5ffcd7f9c4
title: Étape 3
challengeType: 20
dashedName: step-3
---

# --description--

Chaque plat a un coût. Tu dois stocker ces montants dans des variables pour les utiliser plus tard. Comme ces montants incluent des centimes, tu vas utiliser le type float, qui sert à représenter les nombres décimaux. Voici un exemple de variable avec une valeur float :

```py
change = 2.35
```

Crée quatre variables : `appetizers` égale à `37.89`, `main_courses` égale à `57.34`, `desserts` égale à `39.39`, et `drinks` égale à `64.21`.

# --hints--

Tu devrais définir une variable nommée `appetizers` et lui assigner la valeur `37.89`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('appetizers').is_equivalent('appetizers = 37.89')
    `))
})
```

Tu devrais définir une variable nommée `main_courses` et lui assigner la valeur `57.34`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('main_courses').is_equivalent('main_courses = 57.34')
    `))
})
```

Tu devrais définir une variable nommée `desserts` et lui assigner la valeur `39.39`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('desserts').is_equivalent('desserts = 39.39')
    `))
})
```

Tu devrais définir une variable nommée `drinks` et lui assigner la valeur `64.21`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('drinks').is_equivalent('drinks = 64.21')
    `))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

--fcc-editable-region--

--fcc-editable-region--
```
