---
id: 6976cc0b0c135686a4620b94
title: Étape 2
challengeType: 20
dashedName: step-2
---

# --description--

Ensuite, tu dois prendre en compte le nombre de personnes qui partagent l'addition. Stocke cette valeur dans une variable, comme à l'étape précédente.

Crée une variable nommée `num_of_friends` et assigne-lui la valeur `4`. Elle servira plus tard dans l'atelier à calculer la répartition finale.

# --hints--

Tu devrais avoir une variable nommée `num_of_friends`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_variable('num_of_friends')
    `))
})
```

Tu devrais assigner l'entier `4` à ta variable `num_of_friends`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('num_of_friends').is_equivalent('num_of_friends = 4')
    `))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

--fcc-editable-region--

--fcc-editable-region--
```
