---
id: 6977a4306bb3e856690a4890
title: Étape 6
challengeType: 20
dashedName: step-6
---

# --description--

Maintenant que tu as calculé le pourboire, tu dois l'ajouter à ton `running_total` pour obtenir le montant final de l'addition.

Utilise l'opérateur `+=` pour ajouter la valeur de `tip` à ton `running_total`. Enfin, utilise `print()` pour afficher la chaîne `Total with tip:` suivie d'un espace et de la valeur de `running_total`.

# --hints--

Tu devrais utiliser l'opérateur `+=` pour ajouter la valeur de `tip` à ton `running_total`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_stmt('running_total += tip')
    `))
})
```

Tu devrais afficher la chaîne `Total with tip:` suivie d'un espace et de la variable `running_total`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Total with tip:', running_total)") or _Node(_code).has_call("print(f'Total with tip: {running_total}')")`))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

tip = running_total * 0.25
print('Tip amount:', tip)

--fcc-editable-region--

--fcc-editable-region--
```
