---
id: 69779dcb44016eccc05c15a4
title: Étape 5
challengeType: 20
dashedName: step-5
---

# --description--

Le service était excellent, donc le groupe décide de laisser un pourboire de 25 %. Pour calculer un pourcentage en Python, tu peux multiplier le total par l'équivalent décimal du pourcentage.

Par exemple, pour trouver 10 % d'une valeur, tu la multiplierais par `0.10` avec l'opérateur `*` :

```py
tax = total * 0.10
```

Crée une variable nommée `tip` et assigne-lui le résultat de la multiplication de `running_total` par `0.25`.

Enfin, utilise `print()` pour afficher la chaîne `Tip amount:` suivie d'un espace et de la valeur de ta variable `tip`.

# --hints--

Tu devrais avoir une variable nommée `tip`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_variable('tip')
    `))
})
```

Ta variable `tip` devrait être le résultat de `running_total * 0.25`.

```js
({
    test: () => runPython(`
    t = _Node(_code).find_variable('tip')
    assert t.is_equivalent('tip = running_total * 0.25') or t.is_equivalent('tip = 0.25 * running_total')`)
})
```

Tu devrais afficher la chaîne `Tip amount:` suivie d'un espace et de la variable `tip`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Tip amount:', tip)") or _Node(_code).has_call("print(f'Tip amount: {tip}')")`))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

--fcc-editable-region--

--fcc-editable-region--
```
