---
id: 6977a8bbcbdd557886bd0718
title: Étape 7
challengeType: 20
dashedName: step-7
---

# --description--

Avec le pourboire maintenant inclus, tu as le montant final pour tout le groupe. Tu dois déterminer ce que chaque personne doit en divisant l'addition totale par le nombre d'ami(e)s.

En Python, tu utilises la barre oblique `/` pour effectuer une division. Par exemple :

```py
half = 10 / 2
```

Crée une variable nommée `final_bill` et assigne-lui le résultat de la division de `running_total` par `num_of_friends`.

Enfin, utilise la fonction `print()` pour afficher la chaîne `Bill per person:` suivie d'un espace et de la valeur de `final_bill`.

# --hints--

Tu devrais définir une variable nommée `final_bill`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_variable('final_bill')
    `))
})
```

Tu devrais utiliser l'opérateur `/` pour diviser `running_total` par `num_of_friends` et assigner le résultat à ta variable `final_bill`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('final_bill').is_equivalent('final_bill = running_total / num_of_friends')
    `))
})
```

Tu devrais afficher la chaîne `Bill per person:` suivie d'un espace et de la variable `final_bill`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Bill per person:', final_bill)") or _Node(_code).has_call("print(f'Bill per person: {final_bill}')")`))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

tip = running_total * 0.25
print('Tip amount:', tip)

running_total += tip
print('Total with tip:', running_total)

--fcc-editable-region--

--fcc-editable-region--
```
