---
id: 6982684f3a25f379e195a5fc
title: Étape 4
challengeType: 20
dashedName: step-4
---

# --description--

Maintenant que tu as stocké les coûts individuels, tu peux calculer le total.

Rappelle-toi que l'opérateur `+=` ajoute une valeur à une variable existante et la met à jour en même temps. Par exemple :

```py
total = 10
total += 2 + 2 + 1
print(total)  # total is now 15
```

Utilise l'opérateur `+=` une seule fois pour ajouter `appetizers`, `main_courses`, `desserts` et `drinks` à `running_total`.

Enfin, utilise `print()` pour afficher la chaîne `Total bill so far:` suivie d'un espace et de la valeur de `running_total`.

**Note :** tu remarqueras peut-être que la sortie a plus de chiffres décimaux que prévu. Comme tu l'as appris dans une leçon précédente, cela arrive parce que les nombres sont stockés en binaire, et de nombreuses valeurs décimales ne peuvent pas être représentées exactement dans ce format, ce qui entraîne des erreurs d'arrondi.

# --hints--

Tu devrais mettre à jour `running_total` avec l'opérateur `+=` une seule fois pour ajouter les quatre coûts de plats.

```js
({
       test: () => runPython(`
import itertools
perms = itertools.permutations(['+ appetizers', '+ main_courses', '+ desserts', '+ drinks'])
values = (' '.join(perm).lstrip('+') for perm in perms)
solutions = (f'running_total += {v}' for v in values)
assert any(_Node(_code).has_stmt(s) for s in solutions)`)
})
```

Tu devrais afficher la chaîne `Total bill so far:` suivie d'un espace et de la valeur de `running_total`.

```js
({
    test: () => runPython(`
    sol1 = "print('Total bill so far:', running_total)"
    sol2 = "print(f'Total bill so far: {running_total}')"
    assert _Node(_code).has_call(sol1) or _Node(_code).has_call(sol2)`)
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

--fcc-editable-region--

--fcc-editable-region--
```
