---
id: 68a47dcdbce12193caf77012
title: Implémente un chercheur d'indice trié
challengeType: 26
dashedName: lab-sorted-index-finder
---

# --description--

Dans ce lab, tu vas créer une fonction qui renvoie l'indice le plus bas auquel une valeur doit être insérée dans un tableau une fois celui-ci trié dans l'ordre croissant.

**Objectif** : réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois avoir une fonction `getIndexToIns` qui prend deux arguments : un tableau et un nombre.
2. Tu dois utiliser la méthode `sort` pour trier le tableau dans l'ordre croissant.
3. Ta fonction `getIndexToIns` doit renvoyer l'indice le plus bas auquel le nombre doit être inséré en utilisant la méthode `findIndex`.
4. Ta fonction `getIndexToIns` doit toujours renvoyer un nombre.

**Indice :**
La méthode `findIndex` est une méthode de tableau intégrée à JavaScript. Elle prend une fonction de rappel et renvoie l'indice du premier élément qui satisfait la condition. `findIndex` et `sort` sont toutes deux des fonctions d'ordre supérieur.

**Exemples :**

- `getIndexToIns([1, 2, 3, 4], 1.5)` doit renvoyer `1` parce que `1.5` est supérieur à `1` (indice `0`) et inférieur à `2` (indice `1`).
- `getIndexToIns([20, 3, 5], 19)` doit renvoyer `2` parce qu'après le tri en `[3, 5, 20]`, `19` est inférieur à `20` (indice `2`) et supérieur à `5` (indice `1`).


# --hints--

Tu dois avoir une fonction `getIndexToIns`.

```js
assert.isFunction(getIndexToIns);
```

`getIndexToIns` doit toujours renvoyer un nombre.

```js
assert.isNumber(getIndexToIns([10, 20, 30, 40, 50], 35));
```

Ta fonction doit utiliser la méthode `sort`.

```js
assert.match(getIndexToIns.toString(), /\.sort\(/);
```


Ta fonction doit utiliser la méthode `findIndex`.

```js
assert.match(getIndexToIns.toString(), /\.findIndex\(/);
```

`getIndexToIns([10, 20, 30, 40, 50], 35)` doit renvoyer `3`.

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 35), 3);
```

`getIndexToIns([10, 20, 30, 40, 50], 30)` doit renvoyer `2`.

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 30), 2);
```

`getIndexToIns([40, 60], 50)` doit renvoyer `1`.

```js
assert.strictEqual(getIndexToIns([40, 60], 50), 1);
```

`getIndexToIns([3, 10, 5], 3)` doit renvoyer `0`.

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 3), 0);
```

`getIndexToIns([5, 3, 20, 3], 5)` doit renvoyer `2`.

```js
assert.strictEqual(getIndexToIns([5, 3, 20, 3], 5), 2);
```

`getIndexToIns([2, 20, 10], 19)` doit renvoyer `2`.

```js
assert.strictEqual(getIndexToIns([2, 20, 10], 19), 2);
```

`getIndexToIns([3, 10, 5], 11)` doit renvoyer `3`

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 11), 3);
```

`getIndexToIns([], 5)` doit renvoyer `0`

```js
assert.strictEqual(getIndexToIns([], 5), 0);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function getIndexToIns(arr, num) {
  const sorted = [...arr].sort((a, b) => a - b);
  const index = sorted.findIndex(el => el >= num);

  return index === -1 ? sorted.length : index;
}
```
