---
id: 68de40f53fcaf106f51cc20f
title: Construire un lecteur audio et vidéo HTML
challengeType: 25
dashedName: build-an-html-audio-and-video-player
demoType: onClick
---

# --description--

**Objectif :** réponds aux user stories ci-dessous et fais passer tous les tests pour terminer le labo.

**User Stories :**

1. Tu dois avoir un élément `h1` pour le titre principal de la page.
2. Tu dois avoir deux éléments `section`.
3. Dans le premier élément `section`, tu dois avoir un élément `h2` pour le titre de la vidéo en cours de lecture.
4. Sous l'élément `h2`, tu dois avoir un élément `video` avec les attributs `controls` et `width`. L'attribut `width` doit être défini sur `640`.
5. Dans l'élément `video`, tu dois avoir un élément `source` avec un attribut `src` qui pointe vers un fichier vidéo et un attribut `type`.
   - Tu peux utiliser `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4`.
6. Dans le deuxième élément `section`, tu dois avoir un élément `h2` pour le titre de la chanson en cours de lecture.
7. Sous l'élément `h2`, tu dois avoir un élément `audio` avec les attributs `controls` et `loop`, ainsi qu'un attribut `src` qui pointe vers un fichier audio.
   - Tu peux utiliser `https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3`.
   - Ou `https://cdn.freecodecamp.org/curriculum/js-music-player/we-are-going-to-make-it.mp3`.

# --hints--

Tu dois avoir un élément `h1` pour le titre principal de la page.

```js
const h1 = document.querySelector('h1');
assert.exists(h1);
```

Ton élément `h1` ne doit pas être vide.

```js
const h1 = document.querySelector('h1');
assert.isNotEmpty(h1.textContent.trim());
```

Tu ne dois avoir qu'un seul élément `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Tu dois avoir deux éléments `section`.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);

```

Dans le premier élément `section`, tu dois avoir un élément `h2` pour le titre de la vidéo en cours de lecture.

```js
const section = document.querySelector('section');
const h2 = section?.querySelector('h2');
assert.exists(h2);
```

Sous l'élément `h2`, tu dois avoir un élément `video`.

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
assert.exists(video);
```

L'élément `video` doit avoir un attribut `controls`.

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
const controlsAttr = video.controls;
assert.isTrue(controlsAttr);
```

L'élément `video` doit avoir un attribut `width` défini sur `640`.

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
const widthAttr = video.getAttribute('width');
assert.equal(widthAttr, '640');
```

Dans l'élément `video`, tu dois avoir un élément `source`.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
assert.exists(source);
```

L'élément `source` doit avoir un attribut `src` qui pointe vers un fichier vidéo.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const srcAttr = source.getAttribute('src');
assert.match(srcAttr, /\.(mp4|webm|ogg|avi|mov)$/i);
```

L'élément `source` doit avoir un attribut `type`.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const typeAttr = source.getAttribute('type');
assert.isNotEmpty(typeAttr);
```

L'attribut `type` doit correspondre au type de média de l'extension de fichier de l'attribut `src`.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const typeAttr = source.getAttribute('type');
const srcAttr = source.getAttribute('src');
const srcFileExt = srcAttr.match(/[^\.]+$/g)[0];
const fileExtToMediaTypes = {
  mp4: ['video/mp4'],
  webm: ['video/webm'],
  ogg: ['video/ogg'],
  avi: ['video/avi', 'video/vnd.avi', 'video/msvideo', 'video/x-msvideo'],
  mov: ['video/quicktime']
};
const srcMediaTypes = fileExtToMediaTypes[srcFileExt];
assert.oneOf(typeAttr, srcMediaTypes);
```

Dans le deuxième élément `section`, tu dois avoir un élément `h2` pour le titre de la chanson en cours de lecture.

```js
const section = document.querySelector('section:nth-of-type(2)');
const h2 = section?.querySelector('h2');
assert.exists(h2);
```

Tous tes éléments `h2` doivent contenir du texte.

```js
const h2s = document.querySelectorAll('h2');
assert.isNotEmpty(h2s);
h2s.forEach((h2) => assert.isNotEmpty(h2.textContent.trim()));
```

Sous l'élément `h2`, tu dois avoir un élément `audio`.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
assert.exists(audio);
```

L'élément `audio` doit avoir un attribut `controls`.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const controlsAttr = audio.controls;
assert.isTrue(controlsAttr);
```

L'élément `audio` doit avoir un attribut `loop`.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const loopAttr = audio.loop;
assert.isTrue(loopAttr);
```

L'élément `audio` doit avoir un attribut `src` qui pointe vers un fichier audio.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const srcAttr = audio.getAttribute('src');
assert.match(srcAttr, /\.(mp3|wav|ogg|aac|flac)$/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0"/>
  <title>Labo audio et vidéo HTML</title>
</head>
<body>

</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0"/>
  <title>Labo audio et vidéo HTML</title>
</head>
<body>
  <main>
    <h1>Labo audio et vidéo HTML</h1>

    <section>
      <h2>What is the map method and how does it work?</h2>
      <video controls width="640">
        <source src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4" type="video/mp4" />
      </video>
    </section>

    <section>
      <h2>Sailing Away</h2>
      <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3" controls loop>
      </audio>
    </section>

  </main>
</body>
</html>
```
