---
id: a26cbbe9ad8655a977e1ceb5
title: Crée une application de recherche du mot le plus long
challengeType: 26
dashedName: build-a-longest-word-finder-app
---

# --description--

Dans ce lab, tu vas créer une fonction qui retourne la longueur du mot le plus long dans la phrase fournie.

Par exemple, dans la phrase `"The quick brown fox jumped over the lazy dog"`, le mot le plus long est `"jumped"`, qui a une longueur de 6.

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois créer une fonction nommée `findLongestWordLength` qui prend une chaîne en argument.
2. La fonction doit retourner la longueur du mot le plus long de la chaîne.

# --hints--

Tu dois créer une fonction nommée `findLongestWordLength`.

```js
assert.isFunction(findLongestWordLength);
```

`findLongestWordLength` doit avoir un seul paramètre.

```js
assert.lengthOf(findLongestWordLength, 1);
```

`findLongestWordLength("The quick brown fox jumped over the lazy dog")` doit retourner un nombre.

```js
assert.isNumber(
  findLongestWordLength('The quick brown fox jumped over the lazy dog')
);
```

`findLongestWordLength("The quick brown fox jumped over the lazy dog")` doit retourner `6`.

```js
assert.strictEqual(
  findLongestWordLength('The quick brown fox jumped over the lazy dog'),
  6
);
```

`findLongestWordLength("May the force be with you")` doit retourner `5`.

```js
assert.strictEqual(findLongestWordLength('May the force be with you'), 5);
```

`findLongestWordLength("Google do a barrel roll")` doit retourner `6`.

```js
assert.strictEqual(findLongestWordLength('Google do a barrel roll'), 6);
```

`findLongestWordLength("Googling do a barrel roll")` doit retourner `8`.

```js
assert.strictEqual(findLongestWordLength('Googling do a barrel roll'), 8);
```

`findLongestWordLength("What is the average airspeed velocity of an unladen swallow")` doit retourner `8`.

```js
assert.strictEqual(
  findLongestWordLength(
    'What is the average airspeed velocity of an unladen swallow'
  ),
  8
);
```

`findLongestWordLength("What if we try a super-long word such as otorhinolaryngology")` doit retourner `19`.

```js
assert.strictEqual(
  findLongestWordLength(
    'What if we try a super-long word such as otorhinolaryngology'
  ),
  19
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function findLongestWordLength(str) {
  return str.split(' ').sort((a, b) => b.length - a.length)[0].length;
}

findLongestWordLength('The quick brown fox jumped over the lazy dog');
```
