---
id: 587d8259367417b2b2512c86
title: Implémenter le tri par insertion
challengeType: 26
dashedName: implement-insertion-sort
---

# --description--

Dans ce lab, tu vas implémenter l'algorithme de tri par insertion. Cette méthode fonctionne en construisant un tableau trié au début de la liste. Elle commence le tableau trié avec le premier élément. Ensuite, elle inspecte l'élément suivant et l'échange vers l'arrière dans le tableau trié jusqu'à ce qu'il soit en position triée. Elle continue à parcourir la liste et à échanger les nouveaux éléments vers l'arrière dans la portion triée jusqu'à atteindre la fin. Cet algorithme a une complexité temporelle quadratique dans les cas moyen et pire.

**Objectif :** Remplir les user stories ci-dessous et faire passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu dois créer une fonction `insertionSort`.
1. La fonction `insertionSort` doit prendre un tableau d'entiers et renvoyer un tableau avec les mêmes entiers triés du plus petit au plus grand.

# --hints--

`insertionSort` doit être une fonction.

```js
assert.isFunction(insertionSort);
```

`insertionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` doit renvoyer un tableau inchangé sauf pour l'ordre.

```js
assert.sameMembers(
  insertionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`insertionSort` doit renvoyer un tableau trié (du plus petit au plus grand).

```js

assert.sameOrderedMembers(
  insertionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
);
```

`insertionSort([5, 4, 33, 2, 8])` doit renvoyer `[2, 4, 5, 8, 33]`.

```js
assert.sameOrderedMembers(insertionSort([5, 4, 33, 2, 8]), [2, 4, 5, 8, 33])
```

`insertionSort` ne doit pas utiliser la méthode intégrée `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    insertionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function insertionSort (array) {
  for (let currentIndex = 0; currentIndex < array.length; currentIndex++) {
    let current = array[currentIndex];
    let j = currentIndex - 1;
    while (j > -1 && array[j] > current) {
      array[j + 1] = array[j];
      j--;
    }
    array[j + 1] = current;
  }
  return array;
}
```
