---
id: 6735a7370e0ae93a4577c689
title: Concevoir une mise en page de journal
challengeType: 25
dashedName: design-a-newspaper-layout
demoType: onClick
---

# --description--

Dans ce lab, tu vas concevoir une mise en page de journal avec CSS Grid, en utilisant des concepts comme les lignes et les colonnes de grille.

**Objectif :** Réponds aux user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User Stories :**

1. Ta page doit contenir un élément `main` avec la classe `newspaper-layout`.
1. `.newspaper-layout` doit inclure un `header` avec la classe `title`, contenant un `h1` pour afficher le nom du journal.
1. `.newspaper-layout` doit inclure un `article` avec la classe `feature-article` pour l'article principal.
1. `.feature-article` doit inclure un élément `h2` pour le titre de l'article, suivi d'un élément `p` pour le contenu de l'article.
1. Tu dois ajouter trois autres éléments `article` pour les articles plus courts, avec les classes `small-article1`, `small-article2` et `small-article3`.
1. Chacun des articles plus courts doit inclure un élément `h3` pour le titre de l'article, suivi d'un élément `p` pour le contenu de l'article.
1. `.newspaper-layout` doit inclure un élément `article` avec la classe `secondary-article` pour une section d'actualité supplémentaire.
1. `.secondary-article` doit inclure un élément `h2` pour le titre de l'article, suivi d'un élément `p` pour le contenu de l'article.
1. `.newspaper-layout` doit inclure un élément `figure` avec la classe `cover-image` pour afficher une image qui représente le contenu du journal.
1. Les éléments avec les classes `title`, `feature-article`, `secondary-article`, `cover-image`, `small-article1`, `small-article2` et `small-article3` doivent avoir une propriété `grid-area` définie sur le même nom que leur classe.
1. Ton élément `.newspaper-layout` doit utiliser CSS Grid avec une propriété `grid-template-areas` pour définir l'organisation des sections :
   - `.title` doit s'étendre sur toute la ligne du haut.
   - `.feature-article` doit s'étendre sur deux colonnes dans la deuxième ligne, avec `.cover-image` dans la troisième colonne.
   - `.secondary-article` doit s'étendre sur deux colonnes dans la troisième ligne, avec `.cover-image` dans la troisième colonne.
   - Les trois petits articles doivent apparaître sur la quatrième ligne.
1. `.newspaper-layout` doit avoir une propriété `grid-template-columns` qui divise l'espace en trois colonnes égales.
1. Tu dois définir la propriété `grid-template-rows` de `.newspaper-layout` sur `auto` pour la première ligne et diviser l'espace restant en parts égales pour les autres lignes.
1. Tu dois ajouter un espace entre les éléments de grille.
1. Assure-toi que l'image dans `.cover-image` tient dans l'espace prévu en définissant son `max-width` sur `100%`.

**Note :** Assure-toi de lier ta feuille de style dans ton HTML et d'appliquer ton CSS.

# --hints--

Tu dois avoir un élément `main`.

```js
assert.exists(document.querySelector('main'));
```

Ton élément `main` doit avoir une `class` de `newspaper-layout`.

```js
assert.exists(document.querySelector('main.newspaper-layout'));
```

Tu dois avoir un élément `header` avec la `class` `title` dans ton élément `.newspaper-layout`.

```js
const header = document.querySelector(".newspaper-layout > header");
assert.exists(header);
assert(header.classList.contains("title"));
```

Ton élément `header` doit contenir un élément `h1`.

```js
const h1 = document.querySelector("header h1");
assert.exists(h1);
```

Ton élément `.title` doit contenir du texte.

```js
assert.isAbove(document.querySelector(".title").innerText.length, 0);
```

Tu dois avoir un élément `article` pour chaque article, avec respectivement les classes `feature-article`, `small-article1`, `small-article2`, `small-article3` et `secondary-article`.

```js
const featureArticle = document.querySelector('article.feature-article');
const smallArticle1 = document.querySelector('article.small-article1');
const smallArticle2 = document.querySelector('article.small-article2');
const smallArticle3 = document.querySelector('article.small-article3');
const secondaryArticle = document.querySelector('article.secondary-article');

assert.exists(featureArticle);
assert.exists(smallArticle1);
assert.exists(smallArticle2);
assert.exists(smallArticle3);
assert.exists(secondaryArticle);
```

`.feature-article` doit inclure un élément `h2` suivi d'un élément `p`.

```js
const h2 = document.querySelector(".feature-article h2");
const p = document.querySelector(".feature-article h2 + p");
assert.exists(h2);
assert.exists(p);
```

`.small-article1`, `.small-article2` et `.small-article3` doivent inclure un élément `h3` suivi d'un élément `p`.

```js
for (let i = 1; i < 4; i++) {
    let h3 = document.querySelector(`.small-article${i} h3`);
    let p = document.querySelector(`.small-article${i} h3 + p`);
    assert.exists(h3);
    assert.exists(p);
}
```

`.secondary-article` doit inclure un élément `h2` suivi d'un élément `p`.

```js
const h2 = document.querySelector(".secondary-article h2");
const p = document.querySelector(".secondary-article h2 + p");
assert.exists(h2);
assert.exists(p);
```

Tu dois avoir un élément `figure` avec la classe `cover-image` dans ton élément `.newspaper-layout`.

```js
assert.exists(document.querySelector(".newspaper-layout figure.cover-image"));
```

Ton élément `.cover-image` doit contenir un élément `img`.

```js
assert.exists(document.querySelector(".cover-image img"));
```

La propriété `display` de l'élément `.newspaper-layout` doit être définie sur `grid`.

```js
assert.equal(getComputedStyle(document.querySelector('.newspaper-layout')).display, 'grid');
```

Ton élément `.newspaper-layout` doit avoir une propriété `grid-template-areas` où `title` s'étend sur les trois colonnes de la première ligne.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
assert.match(newspaperLayout.gridTemplateAreas, /^("|')title title title\1/);
```

La deuxième ligne du modèle de grille doit contenir un `feature-article` qui s'étend sur deux colonnes et une `cover-image`.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromSecondRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").trimStart();
assert.match(fromSecondRow, /^("|')feature-article feature-article cover-image\1/);
```

La troisième ligne du modèle de grille doit contenir un `secondary-article` qui s'étend sur deux colonnes et une `cover-image`.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromThirdRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").replace(rowRegex, "").trimStart();
assert.match(fromThirdRow, /^("|')secondary-article secondary-article cover-image\1/);
```

La quatrième ligne du modèle de grille doit contenir `small-article1`, `small-article2` et `small-article3`.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromFourthRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").replace(rowRegex, "").replace(rowRegex, "").trimStart();
assert.match(fromFourthRow, /^("|')small-article1 small-article2 small-article3\1$/);
```

Ton sélecteur `.newspaper-layout` doit utiliser la propriété `grid-template-columns` pour diviser l'espace en trois parts égales. Souviens-toi que tu peux utiliser la fonction `repeat` ou définir manuellement chaque unité fractionnaire.

```js
assert.oneOf(
    new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gridTemplateColumns,
    ['1fr 1fr 1fr', 'repeat(3, 1fr)'] 
);
```

Ton sélecteur `.newspaper-layout` doit avoir une propriété `grid-template-rows` définie sur `auto 1fr 1fr 1fr`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gridTemplateRows, ['auto 1fr 1fr 1fr', 'auto repeat(3, 1fr)']);
```

Tu dois ajouter un espace entre les éléments de grille.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gap);
```

`.title` doit avoir une `grid-area` de `title`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.title')?.gridArea, 'title');
```

`.feature-article` doit avoir une `grid-area` de `feature-article`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.feature-article')?.gridArea, 'feature-article');
```

`.cover-image` doit avoir une `grid-area` de `cover-image`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cover-image')?.gridArea, 'cover-image');
```

`.secondary-article` doit avoir une `grid-area` de `secondary-article`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.secondary-article')?.gridArea, 'secondary-article');
```

`.small-article1` doit avoir une `grid-area` de `small-article1`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article1')?.gridArea, 'small-article1');
```

`.small-article2` doit avoir une `grid-area` de `small-article2`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article2')?.gridArea, 'small-article2');
```

`.small-article3` doit avoir une `grid-area` de `small-article3`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article3')?.gridArea, 'small-article3');
```

L'élément `img` dans `.cover-image` doit avoir un `max-width` de `100%`.

```js
const coverImage = document.querySelector(".cover-image");
const img = document.querySelector(".cover-image img");
assert.equal(getComputedStyle(img).maxWidth, "100%");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
<meta charset="utf-8">
<title>Newspaper Layout</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Newspaper Layout</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main class="newspaper-layout">
        <header class="title">
            <h1>The Daily Local News</h1>
        </header>

        <article class="feature-article">
            <h2>Breaking News: Volcano Eruption Disrupts Tourism</h2>
            <p>Recently, a volcano erupted in a popular tourist destination. The eruption has caused widespread panic
                and has disrupted tourism in the area. The volcano has been spewing lava and ash for several days now,
                and authorities are urging residents and tourists to evacuate the area immediately. The eruption has
                also caused a number of flights to be cancelled, leaving many tourists stranded. The situation is still
                developing, and authorities are working to contain the eruption and ensure the safety of everyone in the
                area.</p>
        </article>

        <article class="small-article1">
            <h3>Sports: Local Team Wins Championship</h3>
            <p>Hockey fans are celebrating today as the local team has won the championship. The team, which has been on
                a winning streak all season, clinched the title in a thrilling final match. Fans took to the streets to
                celebrate the victory, waving flags and chanting the team's name.</p>
        </article>

        <article class="small-article2">
            <h3>Health: Tips for a Balanced Diet</h3>
            <p>A diet high in calories, sugar, and unhealthy fats can lead to a variety of health problems, including
                obesity, diabetes, and heart disease. To maintain a healthy weight and reduce your risk of chronic
                diseases, it's important to eat a balanced diet that includes a variety of nutrient-rich foods. Here are
                some tips for eating a balanced diet:</p>
        </article>

        <article class="small-article3">
            <h3>Travel: Top 10 Destinations for 2025</h3>
            <p>Traveling is one of the best ways to experience new cultures, meet new people, and see the world. If
                you're looking for inspiration for your next trip, here are the top 10 destinations for 2025:</p>
        </article>

        <article class="secondary-article">
            <h2>Technology: The Rise of AI</h2>
            <p>Artificial intelligence (AI) is changing the way we live and work. From self-driving cars to virtual
                assistants, AI is revolutionizing the world around us. But what exactly is AI, and how does it work? In
                this article, we'll explore the rise of AI and its impact on society.</p>
        </article>

        <figure class="cover-image">
            <img src="https://cdn.freecodecamp.org/curriculum/labs/volcano.jpg" alt="Cover Image">
        </figure>
    </main>
</body>

</html>
```

```css
body {
    font-family: 'Georgia', 'Times New Roman', Times, serif;
    padding: 0;
    margin: 0;
    background-color: #f0f0f0;
}

.newspaper-layout {
    display: grid;
    grid-template-areas:
        "title title title"
        "feature-article feature-article cover-image"
        "secondary-article secondary-article cover-image"
        "small-article1 small-article2 small-article3";
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: auto 1fr 1fr 1fr;
    gap: 15px;
    background-color: #fff;
    padding: 15px;
    box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
    box-sizing: border-box;
    text-align: justify;
}

.title {
    grid-area: title;
    background-color: #1B1B32;
    color: white;
    padding: 15px;
    text-align: center;
}

.title h1 {
    font-size: 2.5rem;
    line-height: 1.3;
    margin: 0;
}

.feature-article,
.secondary-article {
    background-color: #f9f9f9;
    padding: 15px;
    font-size: 1rem;
}

.feature-article {
    grid-area: feature-article;
}

.secondary-article {
    grid-area: secondary-article;
}

.small-article1,
.small-article2,
.small-article3 {
    background-color: #e9e9e9;
    padding: 15px;
    font-size: 0.9rem;
}

.small-article1 {
    grid-area: small-article1;
}

.small-article2 {
    grid-area: small-article2;
}

.small-article3 {
    grid-area: small-article3;
}

.cover-image {
    grid-area: cover-image;
    display: flex;
    justify-content: center;
    align-items: center;
    background-color: #ddd;
    padding: 15px;
    margin: 0;
}

.cover-image img {
    max-width: 100%;
    height: 350px;
    width: 350px;
    border-radius: 5px;
}
```
