---
id: 56533eb9ac21ba0edf2244cf
title: Crée une collection de disques
challengeType: 26
dashedName: build-a-record-collection
---

# --description--

Tu crées une fonction qui aide à la maintenance d'une collection d'albums de musique. La collection est organisée sous forme d'objet qui contient plusieurs albums, qui sont eux-mêmes des objets. Chaque album est représenté dans la collection avec un `id` unique comme nom de propriété. À l'intérieur de chaque objet album, il y a diverses propriétés décrivant des informations sur l'album. Tous les albums n'ont pas des informations complètes.

La fonction `updateRecords` prend 4 arguments représentés par les paramètres de fonction suivants :

-   `records` - un objet contenant plusieurs albums individuels
-   `id` - un nombre représentant un album précis dans l'objet `records`
-   `prop` - une chaîne représentant le nom de la propriété de l'album à mettre à jour
-   `value` - une chaîne contenant l'information utilisée pour mettre à jour la propriété de l'album

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Ta fonction doit toujours retourner l'objet `records` entier.
2. Si `value` est une chaîne vide, supprime la propriété `prop` donnée de l'album.
3. Si `prop` n'est pas `tracks` et que `value` n'est pas une chaîne vide, affecte la `value` à la propriété `prop` de cet album.
4. Si `prop` est `tracks` et que `value` n'est pas une chaîne vide, mais que l'album n'a pas de propriété `tracks`, crée un tableau vide et ajoute-lui `value`.
5. Si `prop` est `tracks` et que `value` n'est pas une chaîne vide, ajoute `value` à la fin du tableau `tracks` existant de l'album.

**Note :** une copie de l'objet `recordCollection` est utilisée pour les tests. Ta fonction ne doit pas se référer directement à l'objet `recordCollection`, seulement au paramètre de la fonction.

# --before-each--

```js
const _recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};
```

# --hints--

Tu dois avoir une fonction `updateRecords`.

```js
assert.isFunction(updateRecords);
```

Après `updateRecords(recordCollection, 5439, "artist", "ABBA")`, `artist` doit être la chaîne `ABBA`

```js
assert.equal(
  updateRecords(_recordCollection, 5439, 'artist', 'ABBA')[5439]['artist'],
    'ABBA'
);
```

Après `updateRecords(recordCollection, 5439, "tracks", "Take a Chance on Me")`, `tracks` doit avoir la chaîne `Take a Chance on Me` comme dernier et unique élément.

```js
updateRecords(_recordCollection, 5439, 'tracks', 'Take a Chance on Me');
assert.lengthOf(_recordCollection[5439]['tracks'], 1);
assert.equal(_recordCollection[5439]['tracks'].pop(), 'Take a Chance on Me');
```

Après `updateRecords(recordCollection, 2548, "artist", "")`, `artist` ne doit pas être défini

```js
updateRecords(_recordCollection, 2548, 'artist', '');
assert.isFalse(_recordCollection[2548].hasOwnProperty('artist'));
```

Après `updateRecords(recordCollection, 1245, "tracks", "Addicted to Love")`, `tracks` doit avoir la chaîne `Addicted to Love` comme dernier élément.

```js
assert.equal(
  updateRecords(_recordCollection, 1245, 'tracks', 'Addicted to Love')[1245][
    'tracks'
  ].pop(), 'Addicted to Love'
);
```

Après `updateRecords(recordCollection, 2468, "tracks", "Free")`, `tracks` doit avoir la chaîne `1999` comme premier élément.

```js
assert.equal(
  updateRecords(_recordCollection, 2468, 'tracks', 'Free')[2468][
    'tracks'
  ][0], '1999'
);
```

Après `updateRecords(recordCollection, 2548, "tracks", "")`, `tracks` ne doit pas être défini

```js
updateRecords(_recordCollection, 2548, 'tracks', '');
assert.isFalse(_recordCollection[2548].hasOwnProperty('tracks'));
```

Après `updateRecords(recordCollection, 1245, "albumTitle", "Riptide")`, `albumTitle` doit être la chaîne `Riptide`

```js
assert.equal(
  updateRecords(_recordCollection, 1245, 'albumTitle', 'Riptide')[1245][
    'albumTitle'
  ], 'Riptide'
);
```

# --seed--

## --seed-contents--

```js
const recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};



```

# --solutions--

```js
const recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};

function updateRecords(records, id, prop, value) {
  if (value === '') delete records[id][prop];
  else if (prop === 'tracks') {
    records[id][prop] = records[id][prop] || [];
    records[id][prop].push(value);
  } else {
    records[id][prop] = value;
  }

  return records;
}
```
