---
id: 687ad7be1ff45032ccf1d92f
title: Construire un sélecteur de thème
challengeType: 25
dashedName: lab-theme-switcher
demoType: onClick
---

# --description--

Dans ce lab, tu vas construire une application qui bascule entre différents thèmes. Quand tu changes de thème, la couleur d'arrière-plan change et un message différent s'affiche.

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois avoir un élément `button` avec le texte `Switch Theme`.
2. Ton élément `button` doit avoir les attributs suivants :
   - Un attribut `id` défini à `"theme-switcher-button"`.
   - Un attribut `aria-haspopup` défini à `"true"`.
   - Un attribut `aria-expanded` défini à `"false"`.
   - Un attribut `aria-controls` défini à `"theme-dropdown"`.
3. Tu dois avoir un élément `ul` avec les attributs suivants :
   - Un attribut `id` défini à `"theme-dropdown"`.
   - Un attribut `role` défini à `"menu"`.
   - Un attribut `aria-labelledby` défini à `"theme-switcher-button"`.
   - Un attribut `hidden`.
4. Ton élément `ul` doit avoir au moins deux éléments `li` avec un attribut `role` défini à `"menuitem"` et un texte de ton choix représentant un thème différent.
5. Chacun de tes éléments `li` doit avoir un attribut `id` qui commence par `theme-` et se termine par le thème que tu as défini pour le texte de l'élément `li`. Par exemple, si l'un de tes thèmes est `Light`, alors ton `id` doit être `theme-light`.
6. Tu dois avoir un élément avec un attribut `aria-live` défini à `"polite"`.
7. Tu dois avoir un tableau `themes` avec au moins deux objets qui contiennent chacun une propriété `name` et `message`. La propriété `name` représentera un thème différent et la propriété `message` s'affichera quand le thème change. Tu es libre de choisir les valeurs de `name` et `message` que tu veux, mais les valeurs de `name` doivent correspondre au contenu textuel de tes éléments `li`.
8. Quand un utilisateur clique sur `#theme-switcher-button`, l'élément `#theme-dropdown` doit afficher les thèmes disponibles.
9. Quand l'élément `#theme-dropdown` est affiché, l'attribut `aria-expanded` doit être défini à `"true"` pour l'élément `button`.
10. Quand l'utilisateur clique sur `#theme-switcher-button` alors que l'élément `#theme-dropdown` est affiché, le menu doit être fermé, l'attribut `hidden` doit être ajouté et l'attribut `aria-expanded` doit être défini à `"false"`.
11. Quand un utilisateur clique sur `#theme-switcher-button` et sélectionne un thème, la classe `theme-<name>` correspondante doit être ajoutée à l'élément `body` et le message de thème associé défini dans le tableau `themes` doit s'afficher dans l'élément `aria-live="polite"`.
12. Quand un utilisateur clique sur `#theme-switcher-button` et sélectionne un thème différent du thème actuel, la classe de thème sur `body` et le message d'état doivent se mettre à jour pour correspondre à la nouvelle sélection.

# --before-all--

```js
const resetThemeSwitch = () => {
  const themeDropdownBtn = document.getElementById("theme-switcher-button");
  const themeDropdownEl = document.getElementById("theme-dropdown");

  // Reset attributes due to earlier button clicks interfering with these tests
  themeDropdownEl.setAttribute("hidden", "");
  themeDropdownBtn.setAttribute("aria-expanded", "false");
}
```

# --hints--

Tu dois avoir un élément `button`.

```js
assert.exists(document.querySelector("button"));
```

Ton élément `button` doit avoir le texte `Switch Theme`.

```js
const btn = document.querySelector("button");
assert.equal(btn?.textContent.trim(), "Switch Theme");
```

Ton élément `button` doit avoir un attribut `id` défini à `"theme-switcher-button"`.

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("id"), "theme-switcher-button");
```

Ton élément `button` doit avoir un attribut `aria-haspopup` défini à `"true"`.

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-haspopup"), "true");
```

Ton élément `button` doit avoir un attribut `aria-expanded` défini à `"false"`.

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-expanded"), "false");
```

Ton élément `button` doit avoir un attribut `aria-controls` défini à `"theme-dropdown"`.

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-controls"), "theme-dropdown");
```

Tu dois avoir un élément `ul`.

```js
assert.exists(document.querySelector("ul"));
```

Ton élément `ul` doit avoir un attribut `id` défini à `"theme-dropdown"`.

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("id"), "theme-dropdown");
```

Ton élément `ul` doit avoir un attribut `role` défini à `"menu"`.

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("role"), "menu");
```

Ton élément `ul` doit avoir un attribut `aria-labelledby` défini à `"theme-switcher-button"`.

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("aria-labelledby"), "theme-switcher-button");
```

Ton élément `ul` doit avoir un attribut `hidden`.

```js
const ul = document.querySelector("ul");
assert.isTrue(ul?.hasAttribute("hidden"));
```

Ton élément `ul` doit avoir au moins deux éléments `li`.

```js
const listItems = document.querySelectorAll("ul li");
assert.isAtLeast(listItems?.length, 2);
```

Chacun de tes éléments `li` doit avoir un attribut `role` défini à `"menuitem"`.

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
listItems.forEach(li => {
  assert.equal(li?.getAttribute("role"), "menuitem");
})
```

Chacun de tes éléments `li` doit avoir un texte qui représente un thème.

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
listItems.forEach(li => {
  const text = li?.textContent.trim();
  assert.isTrue(text.length > 0);
})
```

Chacun de tes éléments `li` doit avoir un attribut `id` qui commence par `theme-` et se termine par le thème correspondant à cet élément. Assure-toi que la valeur de ton `id` est entièrement en minuscules.

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
const listTitles = [...listItems].map(li => li?.textContent.trim().toLowerCase());

listItems.forEach((li, index) =>{
  assert.equal(li.getAttribute("id"), `theme-${listTitles[index]}`);
})
```

Tu dois avoir un élément avec un attribut `aria-live` défini à `"polite"`.

```js
assert.exists(document.querySelector('[aria-live="polite"]'));
```

Tu dois avoir un tableau `themes`.

```js
assert.isArray(themes);
```

Ton tableau `themes` doit avoir au moins deux objets.

```js
assert.isAtLeast(themes?.length, 2);
themes.forEach(theme => {
  assert.isObject(theme);
});
```

Ton tableau `themes` doit avoir au moins deux objets, chacun avec une propriété `name` et `message`.

```js
assert.isNotEmpty(themes);
themes?.forEach(theme => {
  assert.isObject(theme);
  assert.property(theme, "name");
  assert.property(theme, "message");
})
```

Tes valeurs de propriétés `name` et `message` doivent être des chaînes non vides.

```js
assert.isNotEmpty(themes);
themes?.forEach(theme => {
  assert.isObject(theme);
  assert.isString(theme["name"]);
  assert.isNotEmpty(theme["name"]);
  assert.isString(theme["message"]);
  assert.isNotEmpty(theme["message"]);
})
```

Chacune de tes valeurs de propriété `name` doit correspondre au contenu textuel de l'un de tes éléments `li`. Assure-toi que ces valeurs sont entièrement en minuscules.

```js
const listItems = [...document.querySelectorAll("ul li")]?.map(li => li.textContent.toLowerCase().trim());
const namesArr = themes?.map(({name}) => name);

assert.equal(listItems.length, namesArr.length);
assert.isTrue(namesArr?.every(name => listItems.includes(name)));
```

Quand un utilisateur clique sur `#theme-switcher-button` pour afficher les options de thème, l'attribut `hidden` doit être retiré de l'élément `#theme-dropdown`.

```js
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const themeDropdownEl = document.getElementById("theme-dropdown");

assert.isTrue(themeDropdownEl?.hasAttribute("hidden"));
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isFalse(themeDropdownEl?.hasAttribute("hidden"));
```

Quand l'élément `#theme-dropdown` est affiché, l'attribut `aria-expanded` doit être défini à `"true"` pour l'élément `button`.

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "true");
```

Quand l'utilisateur clique sur `#theme-switcher-button` alors que l'élément `#theme-dropdown` est affiché, le menu déroulant doit être remis à `hidden`.

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const themeDropdownEl = document.getElementById("theme-dropdown");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isFalse(themeDropdownEl?.hasAttribute("hidden"));
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isTrue(themeDropdownEl?.hasAttribute("hidden"));
```

Quand l'utilisateur clique sur `#theme-switcher-button` alors que l'élément `#theme-dropdown` est affiché, l'attribut `aria-expanded` doit être défini à `"false"`.

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "true");
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "false");
```

Quand un utilisateur clique sur `#theme-switcher-button` et sélectionne un thème, la classe `theme-<name>` correspondante doit être ajoutée à l'élément `body`.

```js
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const listItems = document.querySelectorAll("#theme-dropdown li");

themeDropdownBtn.dispatchEvent(new Event("click"));

const clickedItem = listItems[1];
clickedItem?.dispatchEvent(new Event("click", {bubbles: true}));

const expectedClass = clickedItem?.id;
const actualClass = document.body.className;

assert.equal(actualClass, expectedClass);
```

Quand un utilisateur clique sur `#theme-switcher-button` et sélectionne un thème, un `message` lié au `theme` sélectionné depuis le tableau `themes` doit s'afficher dans l'élément `aria-live="polite"`.


```js
const liveRegion = document.querySelector('[aria-live="polite"]');
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const listItems = document.querySelectorAll("ul li");

themeDropdownBtn?.dispatchEvent(new Event("click"));
listItems[1]?.dispatchEvent(new Event("click", {bubbles: true}));

const selectedThemeName = listItems[1]?.textContent.toLowerCase().trim();
const selectedTheme = themes?.find(theme => theme.name === selectedThemeName);

assert.equal(liveRegion?.textContent, selectedTheme?.message);
```

Quand un utilisateur clique sur `#theme-switcher-button` et sélectionne deux thèmes différents, la classe de `body` et le message d'état doivent se mettre à jour pour correspondre à chaque thème sélectionné.

```js
const liveRegion = document.querySelector('[aria-live="polite"]');
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const listItems = document.querySelectorAll("#theme-dropdown li");

themeDropdownBtn?.dispatchEvent(new Event("click"));
listItems[0]?.dispatchEvent(new Event("click", { bubbles: true }));
const firstClass = document.body.className;
const firstMessage = liveRegion?.textContent;

themeDropdownBtn?.dispatchEvent(new Event("click"));
listItems[1]?.dispatchEvent(new Event("click", { bubbles: true }));
const secondClass = document.body.className;
const secondMessage = liveRegion?.textContent;

assert.equal(firstClass, listItems[0]?.id);
assert.equal(secondClass, listItems[1]?.id);
assert.notEqual(firstClass, secondClass);
assert.notEqual(firstMessage, secondMessage);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Theme Switcher</title>
  <link rel="stylesheet" href="./styles.css">
</head>
<body>

  <script src="./script.js"></script>
</body>
</html>
```

```css
body {
  margin: 0;
  font-family: sans-serif;
  transition: background 0.3s, color 0.3s;
}

ul {
  margin: 0;
  padding: 0;
}

li {
  list-style-type: none;
}

#status {
  text-align: center;
  min-height: 20px;
}
```

```js
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Theme Switcher</title>
  <link rel="stylesheet" href="./styles.css">
</head>
<body class="theme-light">
  <button
    class="dropdown-btn"
    id="theme-switcher-button"
    aria-haspopup="true"
    aria-expanded="false"
    aria-controls="theme-dropdown"
  >
    Switch Theme
  </button>

  <ul
    class="dropdown-content"
    id="theme-dropdown"
    role="menu"
    aria-labelledby="theme-switcher-button"
    hidden
  >
    <li class="dropdown-item" id="theme-light" role="menuitem">Light</li>
    <li class="dropdown-item" id="theme-dark" role="menuitem">Dark</li>
    <li class="dropdown-item" id="theme-ocean" role="menuitem">Ocean</li>
    <li class="dropdown-item" id="theme-nord" role="menuitem">Nord</li>
  </ul>

  <div id="status" aria-live="polite"></div>

  <script src="./script.js"></script>
</body>
</html>
```

```css
body {
  margin: 0;
  font-family: sans-serif;
  transition: background 0.3s, color 0.3s;
}

ul {
  margin: 0;
  padding: 0;
}

li {
  list-style-type: none;
}

.dropdown-btn {
  position: relative;
  padding: 10px 20px;
  border: none;
  cursor: pointer;
}

.dropdown-content {
  position: absolute;
  background-color: #ffffff;
  min-width: 160px;
  box-shadow: 0px 4px 8px rgba(0,0,0,0.2);
  z-index: 1;
}

.dropdown-item {
  width: 100%;
  padding: 12px;
  color: #000000;
  background-color: #ffffff;
  cursor: pointer;
}

.dropdown-item:hover {
  background-color: #f1f1f1;
}

#status {
  text-align: center;
  min-height: 20px;
}

/* Themes */
.theme-light {
  background-color: #ffffff;
  color: #000000;
}

.theme-light .dropdown-btn {
  background-color: #d3b3aa;
  color: #000000;
}

.theme-dark {
  background-color: #121212;
  color: #ffffff;
}

.theme-dark .dropdown-btn {
  background-color: #4f4f4f;
  color: #ffffff;
}

.theme-ocean {
  background-color: #0077be;
  color: #e0f7fa;
}

.theme-ocean .dropdown-btn {
  background-color: #193e51;
  color: #ffffff;
}

.theme-nord {
  background-color: #2e3440;
  color: #d8dee9;
}

.theme-nord .dropdown-btn {
  background-color: #596680;
  color: #ffffff;
}
```

```js
const themes = [
  {
    name: "light",
    message: "Hello sunshine — Light theme is on!"
  },
  {
    name: "dark",
    message: "The night is yours — Dark theme is on!"
  },
  {
    name: "nord",
    message: "The frost has settled - Nord theme is on!"
  },
  {
    name: "ocean",
    message: "Blue skies and high tides — Ocean theme is on!"
  },
];

const setTheme = (theme) => {
  document.body.className = "";
  document.body.classList.add(`theme-${theme.name}`);
};

const updateStatus = (theme) => {
  const status = document.getElementById("status");
  status.textContent = theme.message;
};

document.addEventListener("DOMContentLoaded", () => {
  themes.forEach((theme) => {
    const button = document.getElementById(`theme-${theme.name}`);

    if (button) {
      button.addEventListener("click", () => {
        setTheme(theme);
        updateStatus(theme);
        closeMenu();
      });
    }
  });

  const toggleButton = document.getElementById("theme-switcher-button");
  const menu = document.getElementById("theme-dropdown");

  toggleButton.addEventListener("click", () => {
    const isExpanded = toggleButton.getAttribute("aria-expanded") === "true";
    if (isExpanded) {
      closeMenu();
    } else {
      openMenu();
    }
  });

  function closeMenu() {
    menu.setAttribute("hidden", "");
    toggleButton.setAttribute("aria-expanded", "false");
  }

  function openMenu() {
    menu.removeAttribute("hidden");
    toggleButton.setAttribute("aria-expanded", "true");
  }

  // Close the menu when clicking outside
  document.addEventListener("click", (e) => {
    if (!toggleButton.contains(e.target) && !menu.contains(e.target)) {
      closeMenu();
    }
  });
});
```
