---
id: 69aaafeaa4037197c75587cd
title: Implémenter une file
challengeType: 26
dashedName: implement-a-queue
---

# --description--

Dans ce lab, tu vas implémenter une structure de données file à l'aide de fonctions. Une file est une structure de données First-In-First-Out (FIFO) où les éléments sont ajoutés à l'arrière de la file et retirés de l'avant.

**Objectif :** Remplir les user stories ci-dessous et faire passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu dois avoir une fonction `enqueue` qui ajoute un élément à l'arrière de la file.
1. Tu dois avoir une fonction `dequeue` qui retire et renvoie l'élément à l'avant de la file.
1. Tu dois avoir une fonction `front` qui renvoie l'élément à l'avant de la file sans le retirer.
1. Tu dois avoir une fonction `size` qui renvoie le nombre d'éléments dans la file.
1. Tu dois avoir une fonction `isEmpty` qui renvoie `true` si la file est vide et `false` sinon.

**Note :** La plupart des tests dépendent de la fonction `enqueue`. Implémente celle-là d'abord, car les tests pour `dequeue`, `front`, `size` et `isEmpty` nécessitent d'ajouter des éléments à la file.

# --hints--

Tu dois avoir une fonction `enqueue`.

```js
assert.isFunction(enqueue);
```

La fonction `enqueue` doit ajouter un élément à l'arrière de la file.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(size(queue), 2);
```

Tu dois avoir une fonction `dequeue`.

```js
assert.isFunction(dequeue);
```

La fonction `dequeue` doit retirer et renvoyer l'élément à l'avant de la file.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
enqueue(queue, 'third');
assert.strictEqual(dequeue(queue), 'first');
assert.strictEqual(dequeue(queue), 'second');
assert.strictEqual(size(queue), 1);
```

`dequeue` doit renvoyer `undefined` quand elle est appelée sur une file vide.

```js
const queue = initQueue();
assert.isUndefined(dequeue(queue));
```

`enqueue` doit continuer à ajouter des éléments à l'arrière de la file après des dequeues.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(dequeue(queue), 'first');
enqueue(queue, 'third');
assert.strictEqual(front(queue), 'second');
assert.strictEqual(dequeue(queue), 'second');
assert.strictEqual(dequeue(queue), 'third');
```

Différentes instances de file ne doivent pas partager d'état.

```js
const queueOne = initQueue();
const queueTwo = initQueue();
enqueue(queueOne, 'first');
enqueue(queueTwo, 'second');
assert.strictEqual(front(queueOne), 'first');
assert.strictEqual(front(queueTwo), 'second');
assert.strictEqual(size(queueOne), 1);
assert.strictEqual(size(queueTwo), 1);
```

Tu dois avoir une fonction `front`.

```js
assert.isFunction(front);
```

La fonction `front` doit renvoyer l'élément à l'avant de la file sans le retirer.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(size(queue), 2);
```

`front` doit renvoyer `undefined` quand elle est appelée sur une file vide.

```js
const queue = initQueue();
assert.isUndefined(front(queue));
```

Tu dois avoir une fonction `size`.

```js
assert.isFunction(size);
```

La fonction `size` doit renvoyer le nombre correct d'éléments dans la file.

```js
const queue = initQueue();
assert.strictEqual(size(queue), 0);
enqueue(queue, 'first');
assert.strictEqual(size(queue), 1);
enqueue(queue, 'second');
enqueue(queue, 'third');
assert.strictEqual(size(queue), 3);
dequeue(queue);
assert.strictEqual(size(queue), 2);
```

Tu dois avoir une fonction `isEmpty`.

```js
assert.isFunction(isEmpty);
```

La fonction `isEmpty` doit renvoyer `true` pour une file vide.

```js
const queue = initQueue();
assert.isTrue(isEmpty(queue));
```

La fonction `isEmpty` doit renvoyer `false` pour une file non vide.

```js
const queue = initQueue();
enqueue(queue, 'element');
assert.isFalse(isEmpty(queue));
```

La fonction `isEmpty` doit renvoyer `true` après que tous les éléments ont été retirés de la file.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
dequeue(queue);
dequeue(queue);
assert.isTrue(isEmpty(queue));
```

# --seed--

## --seed-contents--

```js
function initQueue() {
  return {
    collection: []
  };
}

function print(queue) {
  console.log(queue.collection);
}

function enqueue(queue, element) {

}

function dequeue(queue) {

}

function front(queue) {

}

function size(queue) {

}

function isEmpty(queue) {

}
```

# --solutions--

```js
function initQueue() {
  return {
    collection: []
  };
}

function print(queue) {
  console.log(queue.collection);
}

function enqueue(queue, element) {
  queue.collection.push(element);
}

function dequeue(queue) {
  return queue.collection.shift();
}

function front(queue) {
  return queue.collection[0];
}

function size(queue) {
  return queue.collection.length;
}

function isEmpty(queue) {
  return queue.collection.length === 0;
}
```
