---
id: 6823f9df49cc206af5471a30
title: Déboguer la page de profil de Camperbot
challengeType: 25
dashedName: lab-debug-camperbots-profile-page
# Le demoType est volontairement retiré pour ne pas afficher d'exemple dans ce projet de débogage.
---

# --description--

Camperbot essaie de créer une page de profil. Il a demandé à un ami de relire son code, et celui-ci lui a dit qu'il contenait quelques erreurs.

Ton objectif est de corriger toutes les erreurs de Camperbot pour qu'il puisse continuer à construire sa page de profil. Termine les éléments des user stories ci-dessous, puis clique sur « Vérifier votre code » pour voir si tu as corrigé toutes les erreurs.

**Histoires utilisateur :**

1. Camperbot essaie d'utiliser un élément `heading2`, mais cet élément n'existe pas. Corrige ces balises pour utiliser un vrai élément de titre de deuxième niveau.
2. Camperbot essaie d'ajouter deux paragraphes avec `pp`, mais ces balises n'existent pas non plus. Corrige-les pour utiliser les bonnes balises de paragraphe.
3. Camperbot utilise un élément `h3` pour le sous-titre `Background and Interests`, mais il contient une erreur de syntaxe. Repère le problème et corrige-le.

# --hints--

Tu dois avoir exactement un élément `h1` sur la page.

```js
assert.lengthOf(document.querySelectorAll("h1"), 1);
```

Tu ne dois pas supprimer l'élément `<h1>Hello from Camperbot!</h1>` de la page.

```js
assert.match(code, /\<h1\s*\>\s*Hello\s*from\s*Camperbot\s*!\s*\<\/h1\s*\>/i);
```

Tu dois avoir exactement un élément `h2` sur la page.

```js
assert.lengthOf(document.querySelectorAll("h2"), 1);
```

Tu dois avoir un élément `h2` avec le texte `About`. Voici un exemple : `<h2>Text here</h2>`

```js
assert.match(code, /\<h2\s*\>\s*About\s*\<\/h2\s*\>/i);
```

Tu dois avoir un élément paragraphe avec le texte `My name is Camperbot and I love learning new things.` Voici un exemple : `<p>Text here</p>`

```js
assert.match(code, /\<p\s*\>\s*My\s*name\s*is\s*camperbot\s*and\s*I\s*love\s*learning\s*new\s*things\s*\.\s*\<\/p\s*\>/i);
```

Tu dois avoir un élément paragraphe avec le texte `I enjoy solving puzzles.` Voici un exemple : `<p>Text here</p>`

```js
assert.match(code, /\<p\s*\>\s*I\s*enjoy\s*solving\s*puzzles\s*\.\s*\<\/p\s*\>/i);
```

Tu dois avoir exactement deux éléments paragraphe sur la page.

```js
assert.lengthOf(document.querySelectorAll("p"), 2);
```

Tu dois avoir exactement un élément `h3` sur la page.

```js
assert.lengthOf(document.querySelectorAll("h3"), 1);
```

Tu dois avoir un élément `h3` avec le texte `Background and Interests`. Voici un exemple : `<h3>Text here</h3>`

```js
assert.match(code, /\<h3\s*\>\s*Background\s*and\s*Interests\s*\<\/h3\s*\>/i);
```

Tu dois avoir exactement cinq éléments au total sur la page.

```js
assert.lengthOf(document.querySelectorAll("body *:not(script)"), 5);
```

# --seed--

## --seed-contents--

```html
<h1>Hello from Camperbot!</h1>

<heading2>About</heading2>

<pp>My name is Camperbot and I love learning new things.</pp>

<h3>Background and Interests<h3/>
<pp>I enjoy solving puzzles.</pp>
```

# --solutions--

```html
<h1>Hello from Camperbot!</h1>

<h2>About</h2>

<p>My name is Camperbot and I love learning new things.</p>

<h3>Background and Interests</h3>
<p>I enjoy solving puzzles.</p>
```
