---
id: 66eaddd04a9e533fba689001
title: Concevoir une carte d'article de blog
challengeType: 25
dashedName: lab-blog-post-card
demoType: onClick
---

# --description--

Dans ce lab, tu vas pratiquer comment styliser les arrière-plans et les bordures en créant une carte d'article de blog.

**Objectif :** Remplir les user stories ci-dessous et faire passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu devrais avoir un élément `div` avec une classe `blog-post-card` pour contenir tous les éléments de ta carte.
2. Dans l'élément `.blog-post-card`, tu devrais avoir un élément image avec un attribut `alt` valide et du texte, et une classe `post-img`. Tu peux utiliser `https://cdn.freecodecamp.org/curriculum/labs/cover-photo.jpg` pour l'attribut `src` de l'image.
3. Tu devrais avoir un `div` avec une classe `post-content` dans l'élément `.blog-post-card` avec ce qui suit :
    - Un élément `h2` avec une classe `post-title` et du texte pour le titre de ton article de blog.
    - Un élément `p` avec une classe `post-excerpt` et du texte pour résumer ton article de blog.
    - Un élément `a` avec une classe `read-more` avec le texte `Read More`.
4. Tu devrais définir la couleur d'arrière-plan de l'élément `body` à une valeur autre que blanc.
5. Tu devrais appliquer les styles suivants à l'élément `.blog-post-card` :
    - Un arrière-plan blanc.
    - Des coins arrondis.
    - Une largeur de ton choix.
    - L'alignement de texte de ton choix.
6. L'élément `.post-img` devrait être stylisé pour que l'image remplisse toute la largeur de la carte et ait une bordure inférieure.
7. L'élément `.post-content` devrait être stylisé pour qu'il y ait du padding à l'intérieur de la carte.
8. Les éléments `.post-title` et `.post-excerpt` devraient avoir une couleur de texte autre que celle par défaut et des margins sur tous les côtés.
9. L'élément `.read-more` devrait être stylisé comme un bouton et avoir :
    - Une couleur de texte autre que celle par défaut.
    - Une couleur d'arrière-plan.
    - Des margins sur tous les côtés.
    - Du padding sur tous les côtés.
    - Des coins arrondis.
    - Une propriété `display` définie à `inline-block`.
    - Un effet hover qui change sa couleur d'arrière-plan.

**Note :** Assure-toi de lier ta feuille de style dans ton HTML et d'appliquer ton CSS.

# --hints--

Tu devrais avoir un `div` avec une classe `blog-post-card`.

```js
assert.exists(document.querySelector("div.blog-post-card"));
```

Tu devrais avoir un élément `img` avec une `class` `post-img`. Assure-toi que ton image a un attribut `alt` avec du texte et un attribut `src` avec une valeur.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const img = card.querySelector("img.post-img");
assert.exists(img);
assert.isString(img.alt);
assert.isNotEmpty(img.alt);
assert.isTrue(img.hasAttribute("src"));
assert.isNotEmpty(img.getAttribute("src"));
```

Tu devrais avoir un `div` avec une classe `post-content`.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);
const content = card.querySelector("div.post-content");
assert.exists(content);
```

Tu devrais avoir un élément `h2` avec une classe `post-title`. Assure-toi qu'il a du contenu texte pour le titre de ton article de blog.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const postTitle = content.querySelector("h2.post-title");
assert.exists(postTitle);
assert.isString(postTitle.textContent);
assert.isNotEmpty(postTitle.textContent);
```

Tu devrais avoir un `p` avec une classe `post-excerpt`. Assure-toi qu'il a du contenu texte pour résumer ton article de blog.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const postExcerpt = content.querySelector("p.post-excerpt");
assert.exists(postExcerpt);
```

Tu devrais avoir un élément `a` avec une `class` `read-more`.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const readMore = content.querySelector("a.read-more");
assert.exists(readMore);
```

Ton élément `.read-more` devrait avoir le texte `Read More`.

```js
const el = document.querySelector("div.post-content a.read-more");
assert.exists(el);
assert.strictEqual(el.textContent.trim(), "Read More");
```

Ton élément `.blog-post-card` devrait avoir une propriété `border-radius` avec une valeur (qui ne devrait pas être `0` ou une valeur négative).

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card);

const borderRadius = getComputedStyle(card).borderRadius;
assert.isAbove(parseInt(borderRadius), 0);
```

Ton élément `.blog-post-card` devrait avoir un arrière-plan blanc.

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card);

const backgroundColor = getComputedStyle(card).backgroundColor;
const whiteColors = ['rgb(255, 255, 255)', 'rgba(255, 255, 255, 1)', '#ffffff', '#fff', 'white'];

assert.oneOf(backgroundColor.toLowerCase(), whiteColors);
```

Ton élément `body` devrait avoir une couleur d'arrière-plan autre que blanc.

```js
const bodyBg = getComputedStyle(document.body).backgroundColor;
const whiteColors = ['rgb(255, 255, 255)', 'rgba(255, 255, 255, 1)', '#ffffff', '#fff', 'white'];

assert.isFalse(whiteColors.includes(bodyBg.toLowerCase()));
```

Tu devrais cibler `.blog-post-card` et définir sa propriété `width`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("width"));
```

Tu devrais cibler `.blog-post-card` et définir sa propriété `text-align`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("text-align"));
```

Tu devrais cibler `.post-content` et définir sa propriété `padding`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".post-content")?.getPropertyValue("padding"));
```

Ton élément `.read-more` devrait avoir un effet hover.

```js
const link = document.querySelector('.read-more');
assert.exists(link);

const d = new __helpers.CSSHelp(document).getStyle('.read-more:hover');
assert.isNotEmpty(d.backgroundColor);
```

Tu devrais cibler `.read-more` et définir sa propriété `color`.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreTextColor = getComputedStyle(readMore).color;
assert.notStrictEqual(readMoreTextColor, 'rgba(0, 0, 0, 0)');
```

Tu devrais cibler `.read-more` et définir sa propriété `background-color`.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreBackgroundColor = getComputedStyle(readMore).backgroundColor;
assert.notStrictEqual(readMoreBackgroundColor, 'rgba(0, 0, 0, 0)');
```

Tu devrais cibler `.read-more` et définir sa propriété `margin`. 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreMargin = getComputedStyle(readMore).marginTop;
assert.notStrictEqual(readMoreMargin, '0px');
```

Tu devrais cibler `.read-more` et définir sa propriété `display`. 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreDisplay = getComputedStyle(readMore).display;
assert.strictEqual(readMoreDisplay, 'inline-block');
```

Tu devrais cibler `.read-more` et définir sa propriété `border-radius`. 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreBorderRadius = getComputedStyle(readMore).borderRadius;
assert.notStrictEqual(readMoreBorderRadius, '0px');
```

Tu devrais cibler `.read-more` et définir sa propriété `padding`. 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMorePadding = getComputedStyle(readMore).padding;
assert.notStrictEqual(readMorePadding, '0px');
```

Ton élément `.post-img` devrait remplir la `width` de la carte.

```js
const postImg = document.querySelector('.post-img');
assert.exists(postImg);

const imgWidth = getComputedStyle(postImg).width;
const cardWidth = getComputedStyle(postImg.parentElement).width;
assert.closeTo(Number(imgWidth.replace(/px/ig, "")), Number(cardWidth.replace(/px/ig, "")), 15);
```

Ton élément `.post-img` devrait avoir une valeur `border-bottom`.

```js
const postImg = document.querySelector('.post-img');
assert.exists(postImg);

const imgBorderBottom = getComputedStyle(postImg).borderBottomWidth;
assert.notStrictEqual(imgBorderBottom, '0px');
```

Tes éléments `.post-title` et `.post-excerpt` devraient avoir des margins et des couleurs de texte autres que par défaut.

```js
const title = document.querySelector('.post-title');
const excerpt = document.querySelector('.post-excerpt');

function isColorApplied(element) {
  const color = getComputedStyle(element).getPropertyValue('color');
  return color && color !== 'rgba(0, 0, 0, 0)' && color !== 'transparent'; 
}

function isMarginApplied(element) {
  const margin = getComputedStyle(element).getPropertyValue('margin');
  return margin && margin !== '0px'; 
}

assert.isTrue(isColorApplied(title));
assert.isTrue(isMarginApplied(title));

assert.isTrue(isColorApplied(excerpt));
assert.isTrue(isMarginApplied(excerpt));

const styles = new __helpers.CSSHelp(document);
// Are CSS selectors used to style the elements:
const postTitle = [".post-title", ".post-excerpt, .post-title", ".post-title, .post-excerpt"].some(selector => {
  return !!styles.getStyle(selector)?.getPropVal('color');
});
const postExcerpt = [".post-excerpt", ".post-excerpt, .post-title", ".post-title, .post-excerpt"].some(selector => {
  return !!styles.getStyle(selector)?.getPropVal('color');
});
assert.isTrue(postTitle);
assert.isTrue(postExcerpt);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="blog-post-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/cover-photo.jpg" alt="Blog Post Image" class="post-img">
        <div class="post-content">
            <h2 class="post-title">Learn Web Development in 2024</h2>
            <p class="post-excerpt">Stay ahead of the curve with the latest trends in web development. Discover what's new and exciting in 2024</p>
            <a href="#" class="read-more">Read More</a>
        </div>
    </div>
</body>
</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background-color: #f0f0f0;
    font-family: Arial, sans-serif;
}

.blog-post-card {
    background-color: white;
    border-radius: 15px;
    box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
    overflow: hidden;
    width: 350px;
    margin: 20px;
    text-align: center;
}

.post-img {
    width: 100%;
    height: auto;
    border-bottom: 5px solid #333;
}

.post-content {
    padding: 20px;
}

.post-title {
    color: #333;
    margin: 0 0 10px;
}

.post-excerpt {
    color: #667;
    margin: 10px 0;
}

.read-more {
    display: inline-block;
    margin-top: 10px;
    padding: 10px 15px;
    background-color: #333;
    color: white;
    text-decoration: none;
    border-radius: 5px;
}

.read-more:hover {
    background-color: #555;
}
```
