---
id: 698c5229e47d57c6dfaa91bd
title: Étape 6
challengeType: 0
dashedName: step-6
---

# --description--

L'attribut `controls` fournit des contrôles de lecture pour l'élément `audio`, notamment pause/reprise, déplacement dans la piste et contrôle du volume.

L'attribut `controls` est un attribut booléen et n'a pas besoin de valeur.

Ajoute l'attribut `controls` à l'élément `audio`.

Tu devrais maintenant voir l'élément `audio` s'afficher sur la page.

# --hints--

Ton élément `audio` devrait avoir l'attribut `controls`.

```js
const audio = document.querySelector('audio');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop></audio>
--fcc-editable-region--

</body>
</html>
```
