---
id: 698c55e7aaacfae48a227459
title: Étape 7
challengeType: 0
dashedName: step-7
---

# --description--

Ajoute un nouvel élément `h2` sous l'élément `audio` avec le texte `Cruising for a Musing`.

Sous le `h2`, ajoute un élément `p` avec le texte `Artist: Quincy Larson`.

# --hints--

Tu devrais avoir un deuxième élément `h2`.

```js
assert.exists(document.querySelector('h2:nth-of-type(2)'));
```

Ton deuxième élément `h2` devrait contenir le texte `Cruising for a Musing`.

```js
const h2 = document.querySelector('h2:nth-of-type(2)');
assert.strictEqual(h2?.textContent.trim(), 'Cruising for a Musing');
```

Tu devrais avoir un deuxième élément `p`.

```js
assert.exists(document.querySelector('p:nth-of-type(2)'));
```

Ton deuxième élément `p` devrait contenir le texte `Artist: Quincy Larson`.

```js
const p = document.querySelector('p:nth-of-type(2)');
assert.strictEqual(p?.textContent.trim(), 'Artist: Quincy Larson');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>
  
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
