---
id: 698c5710136539797f296098
title: Étape 8
challengeType: 0
dashedName: step-8
---

# --description--

Sous l'élément `p`, ajoute un élément `audio`, donne-lui un attribut `src` avec la valeur `https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3`, ainsi que les attributs `loop` et `controls`.

# --hints--

Tu devrais avoir un deuxième élément `audio`.

```js
assert.exists(document.querySelector('audio:nth-of-type(2)'));
```

Le nouvel élément `audio` devrait être sous l'élément `p`.

```js
assert.exists(document.querySelector('p:nth-of-type(2) + audio:nth-of-type(2)'));
```

Ton élément `audio` devrait avoir un attribut `src`.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('src'));
```

Ton élément `audio` devrait avoir un attribut `src` avec la valeur `https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3`.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.strictEqual(audio?.getAttribute('src'),
"https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3");
```

Ton élément `audio` devrait avoir l'attribut `loop`.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('loop'));
```

Ton élément `audio` devrait avoir l'attribut `controls`.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>
  
  <h2>Cruising for a Musing</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
