---
id: 698c513508507ec8bb820227
title: Étape 4
challengeType: 0
dashedName: step-4
---

# --description--

Pour indiquer la ressource média de l'audio, tu dois ajouter l'attribut `src` à l'élément `audio`.

Ajoute l'attribut `src` avec la valeur `https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3`.

# --hints--

Ton élément `audio` devrait avoir un attribut `src`.

```js
const audio = document.querySelector('audio');
assert.isTrue(audio?.hasAttribute('src'));
```

Ton élément `audio` devrait avoir un attribut `src` avec la valeur `https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3`.

```js
const audio = document.querySelector('audio');
assert.strictEqual(audio?.getAttribute('src'),
"https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  <audio></audio>
--fcc-editable-region--

</body>
</html>
```
