---
id: 698c50535b54b2d75fa985fb
title: Étape 5
challengeType: 0
dashedName: step-5
---

# --description--

L'attribut `loop` relancera l'audio lorsque la lecture sera terminée. Pense à un meme sur Internet qui se répète. Si tu omets l'attribut `loop`, l'audio ne sera lu qu'une seule fois.

L'attribut `loop` est un attribut booléen et n'a pas besoin de valeur.

Ajoute l'attribut `loop` à l'élément `audio`.

# --hints--

Ton élément `audio` devrait avoir l'attribut `loop`.

```js
const audio = document.querySelector('audio');
assert.isTrue(audio?.hasAttribute('loop'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3"></audio>
--fcc-editable-region--

</body>
</html>
```
