---
id: 698c4f4e965bd853ad5084e1
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

Sous le `h1`, ajoute un élément `h2` avec le texte `Can't Stay Down`. C'est le titre de la première chanson.

Sous l'élément `h2`, ajoute un élément `p` avec le texte `Artist: Quincy Larson`.

# --hints--

Tu devrais avoir un élément `h2`.

```js
assert.exists(document.querySelector('h2'));
```

Ton élément `h2` devrait être sous l'élément `h1`.

```js
assert.exists(document.querySelector('h1 + h2'));
```

Ton élément `h2` devrait contenir le texte `Can't Stay Down`.

```js
const h2 = document.querySelector('h2');
assert.strictEqual(h2?.textContent.trim(), "Can't Stay Down");
```

Tu devrais avoir un élément `p`.

```js
assert.exists(document.querySelector('p'));
```

Ton élément `p` devrait contenir le texte `Artist: Quincy Larson`.

```js
const p = document.querySelector('p');
assert.strictEqual(p?.textContent.trim(), 'Artist: Quincy Larson');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
