---
id: 698c4fe24f1f1a17be74a23e
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

Ensuite, crée un élément `audio` sous l'élément `p`. Dans les prochaines étapes, tu ajouteras les attributs nécessaires pour permettre à l'élément audio de jouer de la musique.

# --hints--

Tu devrais avoir un élément `audio`.

```js
assert.exists(document.querySelector('audio'));
```

Le nouvel élément `audio` devrait être sous l'élément `p`.

```js
assert.exists(document.querySelector('p + audio'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Travailler avec l'élément audio HTML</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
