---
id: 698c4e69abe14a9746c72312
title: Étape 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Dans cet atelier, tu vas créer un lecteur de musique HTML. Le boilerplate HTML t'a été fourni.

Crée un élément `h1` et donne-lui le texte `freeCodeCamp Tunes`.

# --hints--

Tu devrais avoir un élément `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Ton élément `h1` devrait contenir le texte `freeCodeCamp Tunes`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent.trim(), 'freeCodeCamp Tunes');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
