---
id: 7a97b79c096ecb000d03bd36
title: 'Crée un calculateur du n-ième nombre de Fibonacci'
challengeType: 26
dashedName: build-an-nth-fibonacci-number-calculator
---

# --description--

**Objectif :** Remplis les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu dois créer une fonction nommée `fibonacci`.
2. Tu dois définir un tableau nommé `sequence` dans la fonction `fibonacci`, et il doit être initialisé avec les valeurs `[0, 1]`.
3. La fonction `fibonacci` doit accepter un paramètre, un entier non négatif `n`.
4. Appeler `fibonacci(n)` doit utiliser une approche de programmation dynamique pour calculer et renvoyer le `n`-ième nombre de la suite de Fibonacci, où chaque nombre est la somme des deux nombres précédents.
5. Chaque nombre de Fibonacci calculé doit être ajouté au tableau `sequence`.

# --hints--

Tu dois avoir une fonction nommée `fibonacci`.

```js
assert.isFunction(fibonacci);
```

Ta fonction `fibonacci` doit prendre un seul paramètre.

```js
const params = __helpers.getFunctionParams(fibonacci.toString());
assert.equal(params.length, 1);
```

Tu dois avoir un tableau nommé `sequence` dans la fonction `fibonacci`, initialisé à `[0, 1]`.

```js
assert.match(__helpers.removeJSComments(fibonacci.toString()), /(const|let|var)\s+sequence\s*=\s*\[\s*0\s*,\s*1\s*\]/);
```

`fibonacci(0)` devrait renvoyer `0`.

```js
assert.equal(fibonacci(0), 0);
```

`fibonacci(1)` devrait renvoyer `1`.

```js
assert.equal(fibonacci(1), 1);
```

`fibonacci(2)` devrait renvoyer `1`.

```js
assert.equal(fibonacci(2), 1);
```

`fibonacci(3)` devrait renvoyer `2`.

```js
assert.equal(fibonacci(3), 2);
```

`fibonacci(5)` devrait renvoyer `5`.

```js
assert.equal(fibonacci(5), 5);
```

`fibonacci(10)` devrait renvoyer `55`.

```js
assert.equal(fibonacci(10), 55);
```

`fibonacci(15)` devrait renvoyer `610`.

```js
assert.equal(fibonacci(15), 610);
```

Tu ne dois pas utiliser la récursivité dans ton code.

```js
const funcStr = __helpers.removeJSComments(fibonacci.toString());
const bodyMatch = funcStr.match(/^[^{]*\{([\s\S]*)\}[^}]*$/) || funcStr.match(/=>\s*([\s\S]+)$/);
assert(bodyMatch && !bodyMatch[1].match(/\bfibonacci\s*\(/));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function fibonacci(n) {
  const sequence = [0, 1];

  if (n < 2) {
    return sequence[n];
  }

  for (let i = 2; i <= n; i++) {
    sequence.push(sequence[i - 1] + sequence[i - 2]);
  }

  return sequence[n];
}
```
