---
id: 686386e76197458b7dfd52ef
title: Étape 6
challengeType: 0
dashedName: step-6
---

# --description--

Maintenant, tu vas créer un bouton CTA. Un Call to Action (CTA) est un bouton ou un lien conçu pour inciter l’utilisateur à faire une action précise, ici commencer à écouter.

À l’extérieur du `div` le plus profond, crée un élément d’ancre avec un `href` égal à `#` et le texte `Start listening`.

Cela termine le HTML de la première carte de tarification.

# --hints--

Tu dois créer un élément d’ancre.

```js
assert.exists(document.querySelector("a"))
```

Ton élément d’ancre doit avoir un `href` égal à `#`.

```js
const link = document.querySelector("a")
assert.equal(link?.getAttribute("href"), "#")
```

Ton élément d’ancre doit avoir le texte `Start listening`.

```js
assert.equal(document.querySelector("a")?.textContent, "Start listening")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-supported streaming</li>
              <li><span>&#10003;</span>Curated playlists</li>
            </ul>
          </div>
          --fcc-editable-region--

          --fcc-editable-region--
        </div>
      </div>
    </main>
  </body>
</html>
```
