---
id: 686386e76197458b7dfd52eb
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

Crée un élément `div` sous le `h1`. À l’intérieur de ce `div`, crée un autre élément `div`, puis imbrique encore un élément `div` à l’intérieur.

Et à l’intérieur du `div` le plus profond, crée un élément `h2` avec le texte `Listener`. Cela représentera le premier plan tarifaire.

# --hints--

Tu dois créer un élément `div`.

```js
assert.exists(document.querySelector("div"));
```

Tu dois créer un autre élément `div` à l’intérieur du `div` extérieur.

```js
assert.exists(document.querySelector("div > div"));
```

Tu dois créer un autre élément `div` à l’intérieur du `div` intérieur.

```js
assert.exists(document.querySelector("div > div > div"));
```

Tu dois créer un `h2` avec le texte `Listener` à l’intérieur de l’élément `div` le plus profond.

```js
assert.equal(document.querySelector("div > div > div > h2")?.textContent, "Listener");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
    --fcc-editable-region--

    --fcc-editable-region--
    </main>
  </body>
</html>
```
