---
id: 686386e76197458b7dfd52ee
title: Étape 5
challengeType: 0
dashedName: step-5
---

# --description--

À l’intérieur du `ul`, crée deux éléments de liste. Dans le premier élément de liste, crée un élément `span` et ajoute l’entité HTML `&#10003;` comme texte. Juste après ce `span` mais dans l’élément de liste, ajoute le texte `Ad-supported streaming`.

Dans le second élément de liste, crée un élément `span` et ajoute l’entité `&#10003;` comme texte. Juste après ce `span`, mais dans l’élément de liste, ajoute le texte `Curated playlists`.

# --hints--

Tu dois avoir deux éléments de liste à l’intérieur de ton élément `ul`.

```js
assert.lengthOf(document.querySelectorAll("ul > li"), 2)
```

Ton premier élément `li` doit contenir un élément `span` avec l’entité de caractère HTML `&#10003;`.

```js
const listItems = document.querySelectorAll("ul > li")
const firstSpan = listItems[0].querySelector("span")

assert.equal(firstSpan?.textContent, "✓")
```

Ton premier élément `li` doit avoir le texte `Ad-supported streaming` juste après l’élément `span`.

```js
const listItems = document.querySelectorAll("ul > li");
const span = listItems[0].querySelector("span");

assert.equal(span.nextSibling?.textContent.trim(), "Ad-supported streaming");
```

Ton second élément `li` doit contenir un élément `span` avec l’entité de caractère HTML `&#10003;`.

```js
const listItems = document.querySelectorAll("ul > li")
const secondSpan = listItems[1].querySelector("span")

assert.equal(secondSpan?.textContent, "✓")
```

Ton second élément `li` doit avoir le texte `Curated playlists` juste après l’élément `span`.

```js
const listItems = document.querySelectorAll("ul > li");
const span = listItems[1].querySelector("span");

assert.equal(span.nextSibling?.textContent.trim(), "Curated playlists");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>Start exploring millions of songs with basic features and ads.</p>
            <ul>
            --fcc-editable-region--

            --fcc-editable-region--
            </ul>
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
