---
id: 686386e76197458b7dfd52ec
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

Sous le `h2`, crée un élément `p` avec le texte `$0`. Puis, à l’intérieur de cet élément `p` et après le texte, crée un élément `span` avec le texte `/month`.

# --hints--

Tu dois créer un élément `p` avec le texte `$0`.

```js
const p = document.querySelector("p");
const pTextBeforeSpan = p.childNodes[0].textContent.trim();

assert.equal(pTextBeforeSpan, "$0");
```

Tu dois créer un élément `span` avec le texte `/month` à l’intérieur de l’élément `p`.

```js
const spanText = document.querySelector("p > span")?.textContent;
assert.equal(spanText, "/month");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            --fcc-editable-region--

            --fcc-editable-region--
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
