---
id: 686383abac6c625a84458ea9
title: Étape 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Dans cet atelier, tu vas approfondir Tailwind CSS en utilisant ses classes utilitaires de grille pour construire un composant de tarification.

Le boilerplate HTML de base t’a été fourni. Commence par créer un élément `h1` avec le texte `Choose your listening plan`.

# --hints--

Tu dois créer un élément `h1`.

```js
assert.exists(document.querySelector("h1"));
```

Ton élément `h1` doit avoir le texte `Choose your listening plan`.

```js
assert.equal(document.querySelector("h1")?.textContent, "Choose your listening plan");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
    --fcc-editable-region--

    --fcc-editable-region--
    </main>
  </body>
</html>
```
