---
id: 613297a923965e0703b64796
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

Tu connais peut-être déjà l'élément `meta` ; il sert à préciser des informations sur la page, comme le titre, la description, les mots-clés et l'auteur.

Donne à ta page un élément `meta` avec une valeur de `charset` appropriée.

L'attribut `charset` précise l'encodage des caractères de la page et, aujourd'hui, `UTF-8` est le seul encodage pris en charge par la plupart des navigateurs.

# --hints--

Tu dois créer un élément `meta` dans l'élément `head`.

```js
assert.exists(document.querySelector('head > meta'));
```

La balise `meta` doit avoir une propriété `charset` de `UTF-8`.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
