---
id: 61329d52e5010e08d9b9d66b
title: Étape 4
challengeType: 0
dashedName: step-4
---

# --description--

Un autre élément `meta` important pour l'accessibilité et le SEO est la définition `description`. La valeur de l'attribut `content` est utilisée par les moteurs de recherche pour fournir une description de ta page.

Ajoute un élément `meta` avec l'attribut `name` défini sur `description`, et donne-lui un attribut `content` utile.

# --hints--

Tu dois ajouter un nouvel élément `meta` au `head`.

```js
assert.equal(document.querySelectorAll('meta').length, 3);
```

La balise `meta` doit avoir un attribut `name` avec la valeur `description`.

```js
assert.exists(document.querySelector('meta[name="description"]'));
```

La balise `meta` doit avoir un attribut `content`.

```js
assert.notEmpty(document.querySelector('meta[name="description"]')?.content);
```

La valeur de l'attribut `content` ne doit pas dépasser 165 caractères. _C'est la longueur maximale de description de Google._

```js
assert.isAtMost(document.querySelector('meta[name="description"]')?.content?.length, 165);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
