---
id: 61329b210dac0b08047fd6ab
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

Toujours avec les éléments `meta`, une définition `viewport` indique au navigateur comment afficher la page. En inclure une améliore l'accessibilité visuelle sur mobile et améliore le _SEO_ (optimisation pour les moteurs de recherche).

Ajoute une définition `viewport` avec un attribut `content` détaillant la `width` et l'`initial-scale` de la page.

# --hints--

Tu dois créer un autre élément `meta` dans le `head`.

```js
assert.equal(document.querySelectorAll('head > meta')?.length, 2);
```

La balise `meta` doit avoir un attribut `name` avec la valeur `viewport`.

```js
assert.equal(document.querySelectorAll('head > meta[name="viewport"]')?.length, 1);
```

La balise `meta` doit avoir un attribut `content` avec la valeur `width=device-width, initial-scale=1`.

```js
assert.equal(document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1.0"]')?.length || document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1"]')?.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
