---
id: 6133d11ef548f51f876149e3
title: Étape 6
challengeType: 0
dashedName: step-6
---

# --description--

La navigation est une partie essentielle de l'accessibilité, et les lecteurs d'écran dépendent de toi pour fournir la structure de ta page. Cela se fait avec des éléments HTML sémantiques.

Ajoute un élément `header` et un élément `main` à ta page.

L'élément `header` servira à introduire la page, ainsi qu'à fournir un menu de navigation.

L'élément `main` contiendra le contenu principal de ta page.

# --hints--

Tu dois ajouter un élément `header` dans l'élément `body`.

```js
assert.exists(document.querySelector('body > header'));
```

Tu dois ajouter un élément `main` dans l'élément `body`.

```js
assert.exists(document.querySelector('body > main'));
```

L'élément `header` doit venir avant l'élément `main`.

```js
assert.exists(document.querySelector('header + main'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
