---
id: 6133acc353338c0bba9cb553
title: Étape 5
challengeType: 0
dashedName: step-5
---

# --description--

Enfin dans le `head`, l'élément `title` aide les lecteurs d'écran à comprendre le contenu d'une page. De plus, c'est une partie importante du _SEO_.

Donne à ta page un `title` descriptif et concis.

# --hints--

Tu dois ajouter un élément `title` dans l'élément `head`.

```js
assert.exists(document.querySelector('head > title'));
```

Tu ne dois pas rendre le `title` plus long que 60 caractères.

```js
assert.isAtMost(document.querySelector('head > title')?.textContent?.length, 60);
```

Essaie d’être plus descriptif avec ton élément `title`. _Indice : au moins 15 caractères_

```js
assert.isAtLeast(document.querySelector('head > title')?.textContent?.length, 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
