---
id: 613e275749ebd008e74bb62e
title: Étape 8
challengeType: 0
dashedName: step-8
---

# --description--

Actuellement, le `img` prend sa taille par défaut, qui est trop grande. CSS possède une fonction `max` qui renvoie la plus grande valeur parmi un ensemble de valeurs séparées par des virgules. Par exemple :

```css
img {
  width: max(250px, 25vw);
}
```

Dans l'exemple ci-dessus, la largeur de l'image sera de 250px si la largeur du viewport est inférieure à 1000 pixels. Si la largeur du viewport est supérieure à 1000 pixels, la largeur de l'image sera de 25vw. Cela s'explique par le fait que 25vw est égal à 25 % de la largeur du viewport.

Redimensionne l'image en utilisant son `id` comme sélecteur, et en définissant la `width` comme le maximum entre `10rem` et `18vw`.

# --hints--

Tu dois utiliser le sélecteur `#logo` pour cibler l'élément `img`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('#logo'));
```

La balise `img` doit avoir une propriété `width` de `max(10rem, 18vw)`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(10rem, 18vw)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--

--fcc-editable-region--

```
