---
id: 613e2546d0594208229ada50
title: Étape 7
challengeType: 0
dashedName: step-7
---

# --description--

Dans le `header`, donne du contexte sur la page en y imbriquant un élément `img`, un élément `h1` et un élément `nav`.

Le `img` doit pointer vers `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg`, avoir un `id` de `logo` et avoir un texte `alt` de `freeCodeCamp`.

Le `h1` doit contenir le texte `HTML/CSS Quiz`.

# --hints--

Tu dois ajouter un élément `img` dans l'élément `header`.

```js
assert.exists(document.querySelector('header > img'));
```

Tu dois ajouter un élément `h1` dans l'élément `header`.

```js
assert.exists(document.querySelector('header > h1'));
```

Tu dois ajouter un élément `nav` dans l'élément `header`.

```js
assert.exists(document.querySelector('header > nav'));
```

Tu dois placer les éléments `img`, `h1` et `nav` dans cet ordre.

```js
assert.exists(document.querySelector('img + h1 + nav'));
```

L'élément `img` doit avoir un attribut `src` avec la valeur `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg`.

```js
assert.equal(document.querySelector('img')?.src, 'https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg');
```

L'élément `img` doit avoir un attribut `id` avec la valeur `logo`.

```js
assert.equal(document.querySelector('img')?.id, 'logo');
```

L'élément `img` doit avoir un attribut `alt` avec la valeur `freeCodeCamp`.

```js
assert.equal(document.querySelector('img')?.alt, 'freeCodeCamp');
```

L'élément `h1` doit avoir le texte `HTML/CSS Quiz`.

```js
assert.include(document.querySelector('h1')?.innerText.toLowerCase(), 'html/css quiz');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <header>

    </header>
--fcc-editable-region--
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
