---
id: 614090d5a22b6f0a5a6b464c
title: Étape 13
challengeType: 0
dashedName: step-13
---

<!-- TODO: I purposefully added the `nav` styles without Camper input -->

# --description--

Le sélecteur de combinateur enfant `>` est utilisé entre des sélecteurs pour cibler uniquement les éléments qui correspondent au second sélecteur et qui sont des enfants directs du premier sélecteur.

Cela peut être utile quand tu as des éléments profondément imbriqués et que tu veux contrôler la portée de ton style.

Utilise le sélecteur `>` pour cibler les éléments de liste non ordonnée dans les éléments `nav`, et utilise _Flexbox_ pour espacer les enfants de manière uniforme.

# --hints--

Tu dois utiliser le sélecteur `nav > ul`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('nav > ul'));
```

Les éléments `nav > ul` doivent avoir une propriété `display` de `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('nav > ul')?.display, 'flex');
```

Les éléments `nav > ul` doivent avoir une propriété `justify-content` de `space-evenly`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('nav > ul')?.justifyContent, 'space-evenly');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

--fcc-editable-region--

--fcc-editable-region--
```
