---
id: 6140827cff96e906bd38fc2b
title: Étape 9
challengeType: 0
dashedName: step-9
---

# --description--

Comme décrit dans le <a href="https://design-style-guide.freecodecamp.org/" target="_blank" rel="noopener noreferrer nofollow">guide de style freeCodeCamp</a>, le logo doit conserver un ratio d'aspect de `35 / 4` et avoir du padding autour du texte.

Commence par changer le `background-color` en `#0a0a23` pour voir le logo. Ensuite, utilise la propriété `aspect-ratio` pour définir le ratio d'aspect souhaité à `35 / 4`. Enfin, ajoute un `padding` de `0.4rem` tout autour.

# --hints--

`#logo` doit avoir une propriété `background-color` de `#0a0a23`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.backgroundColor, 'rgb(10, 10, 35)');
```

Tu dois utiliser la propriété `aspect-ratio`.

```js
assert.notEmpty(new __helpers.CSSHelp(document).getStyle('#logo')?.aspectRatio);
```

Tu ne dois pas utiliser la propriété `height`.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('#logo')?.height);
```

Tu ne dois pas modifier la propriété `width`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(10rem, 18vw)');
```

La balise `img` doit avoir une propriété `aspect-ratio` de `35 / 4`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.aspectRatio, '35 / 4');
```

La balise `img` doit avoir une propriété `padding` de `0.4rem`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.padding, '0.4rem');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--
#logo {
  width: max(10rem, 18vw);

}
--fcc-editable-region--

```
