---
id: 61408f155e798909b6908712
title: Étape 12
challengeType: 0
dashedName: step-12
---

# --description--

Pour activer la navigation sur la page, ajoute une liste non ordonnée avec les trois éléments de liste suivants :

- `INFO`
- `HTML`
- `CSS`

Le texte des éléments de liste doit être enveloppé dans des balises d'ancre.

# --hints--

Tu dois imbriquer un élément `ul` dans l'élément `nav`.

```js
assert.equal(document.querySelectorAll('nav > ul')?.length, 1);
```

Tu dois imbriquer trois éléments `li` dans l'élément `ul`.

```js
assert.equal(document.querySelectorAll('nav > ul > li')?.length, 3);
```

Tu dois imbriquer un élément `a` dans chaque élément `li`.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')?.length, 3);
```

Le premier élément `a` doit avoir le texte `INFO`.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')[0]?.textContent.trim(), 'INFO');
```

Le deuxième élément `a` doit avoir le texte `HTML`.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')[1]?.textContent.trim(), 'HTML');
```

Le troisième élément `a` doit avoir le texte `CSS`.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')[2]?.textContent.trim(), 'CSS');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
--fcc-editable-region--
      <nav>

      </nav>
--fcc-editable-region--
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}
```
