---
id: 6140883f74224508174794c0
title: Étape 10
challengeType: 0
dashedName: step-10
---

# --description--

Fais en sorte que le `header` occupe toute la largeur de son conteneur parent, définis sa `height` sur `50px` et son `background-color` sur `#1b1b32`. Ensuite, définis le `display` pour utiliser _Flexbox_.

# --hints--

Tu dois utiliser le sélecteur d'élément `header`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('header'));
```

La balise `header` doit avoir une propriété `width` de `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.width, '100%');
```

La balise `header` doit avoir une propriété `height` de `50px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.height, '50px');
```

La balise `header` doit avoir une propriété `background-color` de `#1b1b32`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.backgroundColor, 'rgb(27, 27, 50)');
```

La balise `header` doit avoir une propriété `display` de `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.display, 'flex');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--

--fcc-editable-region--

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

```
