---
id: 6141fabd6f75610664e908fd
title: Étape 14
challengeType: 0
dashedName: step-14
---

# --description--

Comme il s'agit d'un quiz, tu auras besoin d'un formulaire pour que les utilisateurs puissent envoyer leurs réponses. Tu peux séparer sémantiquement le contenu dans le formulaire avec des éléments `section`.

Dans l'élément `main`, crée un formulaire contenant trois éléments `section`. Ensuite, fais soumettre le formulaire vers `https://freecodecamp.org/practice-project/accessibility-quiz` avec la bonne méthode.

# --hints--

Tu dois imbriquer un élément `form` dans ton élément `main`.

```js
assert.exists(document.querySelector('main > form'));
```

Tu dois imbriquer trois éléments `section` dans l'élément `form`.

```js
assert.equal(document.querySelectorAll('main > form > section')?.length, 3);
```

L'élément `form` doit avoir un attribut `action`.

```js
assert.notEmpty(document.querySelector('main > form')?.action);
```

Tu dois donner à l'attribut `action` une valeur de `https://freecodecamp.org/practice-project/accessibility-quiz`.

```js
assert.equal(document.querySelector('main > form')?.action, 'https://freecodecamp.org/practice-project/accessibility-quiz');
```

L'élément `form` doit avoir un attribut `method`.

```js
assert.notEmpty(document.querySelector('main > form')?.method);
```

L'élément `form` doit avoir un attribut `method` avec la valeur `post`.

```js
assert.equal(document.querySelector('main > form')?.method?.toLowerCase(), 'post');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
--fcc-editable-region--
    <main>

    </main>
--fcc-editable-region--
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}
```
