---
id: 614202874ca576084fca625f
title: Étape 16
challengeType: 0
dashedName: step-16
---

# --description--

Chaque rôle `region` nécessite un libellé, ce qui aide les utilisateurs de lecteurs d'écran à comprendre le but de la région. Une méthode pour ajouter un libellé consiste à ajouter un élément de titre dans la région puis à le référencer avec l'attribut `aria-labelledby`.

Ajoute les attributs `aria-labelledby` suivants aux éléments `section` :

- `student-info`
- `html-questions`
- `css-questions`

Ensuite, dans chaque élément `section`, imbrique un élément `h2` avec un `id` correspondant à l'attribut `aria-labelledby` associé. Donne à chaque `h2` un contenu textuel approprié.

# --hints--

Le premier élément `section` doit avoir un attribut `aria-labelledby` avec la valeur `student-info`.

```js
assert.equal(document.querySelectorAll('section')?.[0]?.getAttribute('aria-labelledby'), 'student-info');
```

Le deuxième élément `section` doit avoir un attribut `aria-labelledby` avec la valeur `html-questions`.

```js
assert.equal(document.querySelectorAll('section')?.[1]?.getAttribute('aria-labelledby'), 'html-questions');
```

Le troisième élément `section` doit avoir un attribut `aria-labelledby` avec la valeur `css-questions`.

```js
assert.equal(document.querySelectorAll('section')?.[2]?.getAttribute('aria-labelledby'), 'css-questions');
```

Tu dois imbriquer un élément `h2` dans le premier élément `section`.

```js
assert.equal(document.querySelectorAll('section')?.[0]?.querySelectorAll('h2')?.length, 1);
```

Tu dois imbriquer un élément `h2` dans le deuxième élément `section`.

```js
assert.equal(document.querySelectorAll('section')?.[1]?.querySelectorAll('h2')?.length, 1);
```

Tu dois imbriquer un élément `h2` dans le troisième élément `section`.

```js
assert.equal(document.querySelectorAll('section')?.[2]?.querySelectorAll('h2')?.length, 1);
```

Le premier élément `h2` doit avoir un attribut `id` avec la valeur `student-info`.

```js
assert.equal(document.querySelectorAll('h2')?.[0]?.id, 'student-info');
```

Le deuxième élément `h2` doit avoir un attribut `id` avec la valeur `html-questions`.

```js
assert.equal(document.querySelectorAll('h2')?.[1]?.id, 'html-questions');
```

Le troisième élément `h2` doit avoir un attribut `id` avec la valeur `css-questions`.

```js
assert.equal(document.querySelectorAll('h2')?.[2]?.id, 'css-questions');
```

Le premier élément `h2` doit avoir un contenu textuel approprié. _Indice : j'aurais choisi `Student Info`_.

```js
assert.isNotEmpty(document.querySelectorAll('h2')[0]?.textContent.trim());
```

Le deuxième élément `h2` doit avoir un contenu textuel approprié. _Indice : j'aurais choisi `HTML`_.

```js
assert.isNotEmpty(document.querySelectorAll('h2')[1]?.textContent.trim());
```

Le troisième élément `h2` doit avoir un contenu textuel approprié. _Indice : j'aurais choisi `CSS`_.

```js
assert.isNotEmpty(document.querySelectorAll('h2')[2]?.textContent.trim());
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
--fcc-editable-region--
      <form method="post" action="https://freecodecamp.org/practice-project/accessibility-quiz">
        <section role="region"></section>
        <section role="region"></section>
        <section role="region"></section>
      </form>
--fcc-editable-region--
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}
```
