---
id: 6143639d5eddc7094161648c
title: Étape 20
challengeType: 0
dashedName: step-20
---

# --description--

Il est important de relier chaque `input` à l'élément `label` correspondant. Cela fournit aux utilisateurs de technologies d'assistance une référence visuelle vers l'entrée.

Cela se fait en donnant au `label` un attribut `for`, qui contient l'`id` de l'`input`.

Cette section demandera le nom, l'adresse e-mail et la date de naissance d'un étudiant. Donne aux éléments `label` des attributs `for` appropriés, ainsi qu'un contenu textuel. Ensuite, relie les éléments `input` aux éléments `label` correspondants.

# --hints--

Le premier élément `label` doit avoir un attribut `for` approprié.

```js
assert.isAtLeast(document.querySelectorAll('label')?.[0]?.htmlFor?.length, 1);
```

Le deuxième élément `label` doit avoir un attribut `for` approprié.

```js
assert.isAtLeast(document.querySelectorAll('label')?.[1]?.htmlFor?.length, 1);
```

Le troisième élément `label` doit avoir un attribut `for` approprié.

```js
assert.isAtLeast(document.querySelectorAll('label')?.[2]?.htmlFor?.length, 1);
```

Le premier élément `label` doit avoir un contenu textuel approprié.

```js
assert.isNotEmpty(document.querySelectorAll('label')[0]?.textContent.trim());
```

Le deuxième élément `label` doit avoir un contenu textuel approprié.

```js
assert.isNotEmpty(document.querySelectorAll('label')[1]?.textContent.trim());
```

Le troisième élément `label` doit avoir un contenu textuel approprié.

```js
assert.isNotEmpty(document.querySelectorAll('label')[2]?.textContent.trim());
```

Le premier élément `input` doit avoir un attribut `id` correspondant à l'attribut `for` du premier élément `label`.

```js
assert.equal(document.querySelectorAll('input')?.[0]?.id, document.querySelectorAll('label')?.[0]?.htmlFor);
```

Le deuxième élément `input` doit avoir un attribut `id` correspondant à l'attribut `for` du deuxième élément `label`.

```js
assert.equal(document.querySelectorAll('input')?.[1]?.id, document.querySelectorAll('label')?.[1]?.htmlFor);
```

Le troisième élément `input` doit avoir un attribut `id` correspondant à l'attribut `for` du troisième élément `label`.

```js
assert.equal(document.querySelectorAll('input')?.[2]?.id, document.querySelectorAll('label')?.[2]?.htmlFor);
```

Tu ne dois pas utiliser le même attribut `id` pour plus d'un élément `input`.

```js
const id = (n) => document.querySelectorAll('input')?.[n]?.id;
assert.notEqual(id(0), id(1));
assert.notEqual(id(0), id(2));
assert.notEqual(id(1), id(2));
```

Tu ne dois pas utiliser le même attribut `for` pour plus d'un élément `label`.

```js
const htmlFor = (n) => document.querySelectorAll('label')?.[n]?.htmlFor;
assert.notEqual(htmlFor(0), htmlFor(1));
assert.notEqual(htmlFor(0), htmlFor(2));
assert.notEqual(htmlFor(1), htmlFor(2));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a href="#student-info">INFO</a></li>
          <li><a href="#html-questions">HTML</a></li>
          <li><a href="#css-questions">CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <form method="post" action="https://freecodecamp.org/practice-project/accessibility-quiz">
--fcc-editable-region--
        <section role="region" aria-labelledby="student-info">
          <h2 id="student-info">Student Info</h2>
          <div class="info">
            <label></label>
            <input />
          </div>
          <div class="info">
            <label></label>
            <input />
          </div>
          <div class="info">
            <label></label>
            <input />
          </div>
        </section>
--fcc-editable-region--
        <section role="region" aria-labelledby="html-questions">
          <h2 id="html-questions">HTML</h2>
        </section>
        <section role="region" aria-labelledby="css-questions">
          <h2 id="css-questions">CSS</h2>
        </section>
      </form>
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}

h1,
h2 {
  font-family: Verdana, Tahoma;
}

h2 {
  border-bottom: 4px solid #dfdfe2;
}

```
