---
id: 614ccc21ea91ef1736b9b578
title: Étape 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Pour ce workshop, un code de base t'a été fourni.

Commence ce parcours sur l'accessibilité en ajoutant un attribut `lang` à ton élément `html`. Cela aidera les lecteurs d'écran à identifier la langue de la page.

N'oublie pas non plus de lier ta feuille de style à la page.

# --hints--

L'élément `html` doit avoir un attribut `lang`. _Indice : tu peux utiliser la valeur `en` pour l'anglais._

```js
assert.match(code, /<html\s+lang=('|")[\w\-]+?\1\s*>/i);
// TODO: This should/could be fixed in the builder.js
// assert.notThrow(Intl.getCanonicalLocales(document.querySelector('html').lang));
```

Tu dois avoir un élément `link` dans ton élément `head`.

```js
assert.exists(document.querySelector("head>link"));
```

Ton élément `link` doit avoir `rel="stylesheet"` et `href="styles.css"`.

```js
const link = document.querySelector("head>link");
assert.equal(link?.getAttribute("rel"), "stylesheet");
const href = link?.getAttribute("data-href");
assert.oneOf(href, ["./styles.css", "styles.css"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
  <head>

  </head>
  <body>

  </body>
</html>
--fcc-editable-region--

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
