---
id: 66ec4c8e9878d8441956516f
title: Construire un tableau de catalogue de livres
challengeType: 25
dashedName: build-a-book-catalog-table
demoType: onClick
---

# --description--

**Objectif :** Respecte les récits utilisateur ci-dessous et fais passer tous les tests pour terminer le lab.

**Récits utilisateur :**

1. Tu dois créer un élément `table` qui liste des informations sur des livres.
1. Ton tableau doit avoir un élément d’en-tête de tableau avec une ligne dedans.
1. La ligne dans ton élément d’en-tête de tableau doit avoir quatre éléments d’en-tête de tableau, avec les textes `Title`, `Author`, `Genre` et `Publication Year`, dans cet ordre.
1. Ton tableau doit avoir un élément de corps de tableau avec au moins cinq lignes dedans.
1. Chaque ligne du corps de ton tableau doit avoir quatre éléments de données de tableau qui affichent le titre, l’auteur, le genre et l’année de publication du livre.
1. Ton tableau doit avoir un élément de pied de tableau avec une ligne dedans.
1. La ligne dans ton élément de pied de tableau doit avoir un élément de données de tableau qui s’étend sur quatre colonnes et contient le texte `Total Books: N`, où `N` doit être remplacé par le nombre de livres dans ton tableau.

# --hints--

Tu dois avoir un seul élément `table`.

```js
assert.lengthOf(document.querySelectorAll('table'), 1);
```

Tu dois avoir un seul élément `thead` dans ton élément `table`.

```js
assert.lengthOf(document.querySelectorAll('table thead'), 1);
```

Tu dois avoir un seul élément `tr` dans ton élément `thead`.

```js
assert.lengthOf(document.querySelectorAll('thead tr'), 1);
```

Tu dois avoir quatre éléments `th` dans la ligne de ton élément `thead`.

```js
assert.lengthOf(document.querySelectorAll('thead tr th'), 4);
```

Tes quatre éléments `th` doivent avoir les textes `Title`, `Author`, `Genre` et `Publication Year`, dans cet ordre.

```js
const ths = document.querySelectorAll('thead tr th');
assert.equal(ths[0]?.textContent.trim(), 'Title');
assert.equal(ths[1]?.textContent.trim(), 'Author');
assert.equal(ths[2]?.textContent.trim(), 'Genre');
assert.equal(ths[3]?.textContent.trim(), 'Publication Year');
```

Tu dois avoir un seul élément `tbody` dans ton élément `table`.

```js
assert.lengthOf(document.querySelectorAll('table tbody'), 1);
```

Ton élément `tbody` doit avoir au moins cinq lignes.

```js
assert.isAtLeast(document.querySelectorAll('tbody tr').length, 5);
```

Chaque ligne dans ton élément `tbody` doit avoir exactement quatre éléments `td` comme enfants.

```js
const rows = document.querySelectorAll('tbody tr');
assert.isAtLeast(rows.length, 1);

rows.forEach(row => {
  assert.lengthOf(row.querySelectorAll('td'), 4);
  assert.lengthOf(row.children, 4);
});
```

Chaque élément `td` dans le corps de ton tableau doit avoir du texte avec des informations sur un livre.

```js
const tds = document.querySelectorAll('tbody tr td');
assert.isAtLeast(tds.length, 1);

tds.forEach(td => {
  assert.isAtLeast(td.textContent.trim().length, 1);
});
```

Tu dois avoir un seul élément `tfoot` dans ton élément `table`.

```js
assert.lengthOf(document.querySelectorAll('table tfoot'), 1);
```

Tu dois avoir exactement un élément `tr` dans ton élément `tfoot`.

```js
assert.lengthOf(document.querySelectorAll('tfoot tr'), 1);
```

L’élément `td` dans la ligne de ton élément `tfoot` doit avoir son attribut `colspan` défini sur `4`.

```js
assert.equal(document.querySelector('tfoot tr td')?.colSpan, 4);
```

L’élément `td` dans la ligne de ton élément `tfoot` doit avoir le texte `Total Books: N`, où `N` est le nombre de livres dans ton tableau.

```js
const numberOfBooks = document.querySelectorAll('tbody tr')?.length;
assert.equal(document.querySelector('tfoot tr td').textContent.trim(), `Total Books: ${numberOfBooks}`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Book Catalog</title>
</head>

<body>
</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Book Catalog</title>
</head>
<body>
    <h1>Book Catalog</h1>
  <table>
    <thead>
      <tr>
        <th>Title</th>
        <th>Author</th>
        <th>Genre</th>
        <th>Publication Year</th>
      </tr>
    </thead>
    <tbody>
      <tr>
        <td>How to Contribute to Open-Source Projects – A Handbook for Beginners</td>
        <td>Hillary Nyakundi</td>
        <td>Open Source</td>
        <td>2023</td>
      </tr>
      <tr>
        <td>Learn Linux for Beginners: From Basics to Advanced Techniques</td>
        <td>Zaira Hira</td>
        <td>Linux</td>
        <td>2024</td>
      </tr>
      <tr>
        <td>How to Learn to Code and Get a Developer Job</td>
        <td>Quincy Larson</td>
        <td>Learn To Code</td>
        <td>2024</td>
      </tr>
      <tr>
        <td>The Regular Expressions Book – RegEx for JavaScript Developers</td>
        <td>Kolade Chris</td>
        <td>Regular Expressions</td>
        <td>2023</td>
      </tr>
      <tr>
        <td>The Python Code Example Handbook</td>
        <td>Farhan Hasin Chowdhury</td>
        <td>Python</td>
        <td>2023</td>
      </tr>
    </tbody>
    <tfoot>
      <tr>
        <td colspan="4">Total Books: 5</td>
      </tr>
    </tfoot>
  </table>
</body>
</html>
```
