---
id: 66da326c02141df538f29ba5
title: 🧪 Construire une page de paiement (TEST)
challengeType: 25
dashedName: build-a-checkout-page
demoType: onClick
---

# --description--

**Objectif :** Réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu devrais avoir un élément `h1` avec le texte `Checkout`.
1. Tu devrais avoir deux éléments `section` immédiatement après l'élément `h1`.
1. Tu devrais avoir un élément `h2` avec le texte `Your Cart` à l'intérieur de la première section.
1. Tu devrais avoir une image d'un article dans la première section avec un texte alternatif approprié. Tu peux utiliser cette image si tu veux : `https://cdn.freecodecamp.org/curriculum/labs/cube.jpg`
1. Tu devrais avoir un élément `h2` avec le texte `Payment Information` à l'intérieur de la deuxième section.
1. Tu devrais avoir un élément `form` à l'intérieur de la deuxième section.
1. Tu devrais avoir un champ de saisie avec un `id` et un `name` de `card-name`, et un `type` de `text` à l'intérieur de ton formulaire, avec un `label` qui lui est associé.
1. Tu devrais avoir un champ de saisie avec un `id` et un `name` de `card-number`, et un `type` de `text` à l'intérieur de ton formulaire, avec un `label` qui lui est associé.
1. Tu devrais avoir au moins deux éléments `input` avec l'attribut `required`.
1. Tu devrais inclure un élément `span` avec le texte `*` et `aria-hidden` mis à `true` à l'intérieur de l'élément `label` pour chaque champ obligatoire, afin que les champs obligatoires soient visuellement indiqués.
1. Tu devrais avoir un élément `p` avec un texte d'aide qui explique le format requis pour le numéro de carte, placé immédiatement après le champ de saisie du numéro de carte. Le `p` doit avoir un `id` de `card-number-help` et être référencé par le champ de saisie du numéro de carte en utilisant `aria-describedby`.

# --hints--

Tu devrais avoir un élément `h1` avec le texte `Checkout`.

```js
assert.equal(document.querySelector('h1')?.innerText, 'Checkout');
```

Tu ne devrais avoir qu'un seul élément `h1` sur ta page.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Tu devrais avoir au moins deux éléments `section` après l'élément `h1`.

```js
const h1 = document.querySelector('h1');
const firstSibling = h1?.nextElementSibling;
const secondSibling = firstSibling?.nextElementSibling;
assert.equal(firstSibling?.tagName, 'SECTION');
assert.equal(secondSibling?.tagName, 'SECTION');
```

Tu devrais avoir un élément `h2` avec le texte `Your Cart` à l'intérieur de la première section.

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.equal(firstSection?.querySelector('h2')?.innerText, 'Your Cart');
```

Tu devrais avoir au moins une image avec un texte alternatif dans ta première section.

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.isAtLeast(firstSection?.querySelector('img')?.alt?.length, 1);
```

Tu devrais avoir un élément `h2` avec le texte `Payment Information` à l'intérieur de la deuxième section.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.equal(secondSection?.querySelector('h2')?.innerText, 'Payment Information');
```

Tu devrais avoir un élément `form` à l'intérieur du deuxième élément `section`.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form'));
```

Tu devrais avoir un élément `input` avec un `id` et un `name` de `card-name` à l'intérieur de ton formulaire.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-name'));
assert.exists(secondSection?.querySelector('form input[name="card-name"]'));
```

Ton champ de saisie pour le nom sur la carte devrait avoir un `type` de `text`.

```js
const cardNameInput = document.getElementById('card-name');
assert.exists(cardNameInput);
assert.equal(cardNameInput.getAttribute('type'), 'text');
```

Tu devrais avoir un élément `input` avec un `id` et un `name` de `card-number` à l'intérieur de ton formulaire.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-number'));
assert.exists(secondSection?.querySelector('form input[name="card-number"]'));
```

Ton champ de saisie pour le numéro de carte devrait avoir un `type` de `text`.

```js
const cardNumberInput = document.getElementById('card-number');
assert.exists(cardNumberInput);
assert.equal(cardNumberInput.getAttribute('type'), 'text');
```

Tous tes éléments `input` qui ne sont pas de `type` `submit` devraient avoir un élément `label` qui leur est associé.

```js
const inputs = document.querySelectorAll('input:not([type="submit"])');
assert.isAtLeast(inputs.length, 1);
inputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
});
```

Tu devrais avoir au moins deux éléments `input` avec l'attribut `required`.

```js
const inputs = document.querySelectorAll('input');
const requiredInputs = Array.from(inputs).filter(input => input.required);
assert.isAtLeast(requiredInputs?.length, 2)
```

Tu devrais inclure un élément `span` avec le texte `*` à l'intérieur de l'élément `label` pour chaque champ obligatoire.

```js
const requiredInputs = Array.from(document.querySelectorAll('input'))
  .filter(input => input.required);

assert.isNotEmpty(requiredInputs);

requiredInputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
  const span = label?.querySelector('span');
  assert.equal(span?.textContent.trim(), '*');
});
```

Tes éléments `span` devraient avoir `aria-hidden` mis à `true`.

```js
const requiredInputs = Array.from(document.querySelectorAll('input'))
  .filter(input => input.required);

assert.isNotEmpty(requiredInputs);

requiredInputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
  const span = label?.querySelector('span');
  assert.equal(span?.getAttribute('aria-hidden'), 'true');
});
```

Tu devrais avoir un élément `p` avec un `id` de `card-number-help` immédiatement après le champ de saisie du numéro de carte.

```js
const cardNumberInput = document.getElementById('card-number');
const cardNumberHelp = document.getElementById('card-number-help');

assert.exists(cardNumberHelp);
const nextElement = cardNumberInput.parentElement?.tagName === 'LABEL'
  ? cardNumberInput.parentElement.nextElementSibling
  : cardNumberInput.nextElementSibling;
assert.equal(nextElement, cardNumberHelp);
```

Ton texte d'aide pour le numéro de carte ne devrait pas être vide.

```js
const cardNumberHelp = document.getElementById('card-number-help');
assert.isString(cardNumberHelp?.textContent);
assert.isAbove(cardNumberHelp?.textContent.trim().length, 0);
```

Ton champ de saisie du numéro de carte devrait avoir `aria-describedby` mis à `card-number-help`.

```js
const cardNumberInput = document.getElementById('card-number');
assert.equal(cardNumberInput.getAttribute('aria-describedby'), 'card-number-help');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>Checkout Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Checkout Page</title>
</head>
<body>
    <div class="container">
        <h1>Checkout</h1>

        <section>
            <h2>Your Cart</h2>
            <img src="https://cdn.freecodecamp.org/curriculum/labs/spoon.jpg" alt="Left-handed spoon">
            <p>Left-handed spoon</p>
            <p>$29.99</p>
        </section>

        <section>
            <h2>Payment Information</h2>
            <form action="/submit-payment" method="POST">
                <div>
                    <label for="card-name">Cardholder Name <span aria-hidden="true">*</span></label>
                    <input type="text" id="card-name" name="card-name" required>
                </div>
                <div>
                    <label for="card-number">Card Number <span aria-hidden="true">*</span></label>
                    <input type="text" id="card-number" name="card-number" aria-describedby="card-number-help" required>
                    <p id="card-number-help">Please enter your 16-digit card number without spaces or dashes.</p>
                </div>
                <div>
                    <label for="expiry-date">Expiry Date <span aria-hidden="true">*</span></label>
                    <input type="text" id="expiry-date" name="expiry-date" placeholder="MM/YY" required>
                </div>
                <div>
                    <label for="cvv">CVV <span aria-hidden="true">*</span></label>
                    <input type="text" id="cvv" name="cvv" required aria-label="Card Verification Value">
                </div>
                <input type="submit" value="Place Order">
            </form>
        </section>
    </div>
</body>
</html>
```
