---
id: a3bfc1673c0526e06d3ac698
title: Crée un calculateur de somme de nombres premiers
challengeType: 26
dashedName: build-a-prime-number-sum-calculator
---

# --description--

Dans ce lab, tu vas créer un calculateur qui prend un nombre et renvoie la somme de tous les nombres premiers inférieurs ou égaux à ce nombre.

**Objectif** : réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois avoir une fonction `sumPrimes` qui accepte un nombre comme argument.
1. La fonction `sumPrimes` doit renvoyer la somme de tous les nombres premiers inférieurs ou égaux au nombre fourni.
1. Si le nombre en entrée est inférieur à `2`, la fonction doit renvoyer `0`.

# --hints--

Tu dois avoir une fonction `sumPrimes`.

```js
assert.isFunction(sumPrimes);
```

`sumPrimes(10)` doit renvoyer `17`.

```js
assert.strictEqual(sumPrimes(10), 17);
```

`sumPrimes(5)` doit renvoyer `10`.

```js
assert.strictEqual(sumPrimes(5), 10);
```

`sumPrimes(2)` doit renvoyer `2`.

```js
assert.strictEqual(sumPrimes(2), 2);
```

`sumPrimes(0)` doit renvoyer `0`.

```js
assert.strictEqual(sumPrimes(0), 0);
```

`sumPrimes(977)` doit renvoyer `73156`.

```js
assert.strictEqual(sumPrimes(977), 73156);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sumPrimes(num) {
  function isPrime(n) {
    if (n < 2) return false;
    for (let i = 2; i <= Math.sqrt(n); i++) {
      if (n % i === 0) {
        return false;
      }
    }
    return true;
  }

  let sum = 0;

  for (let i = 2; i <= num; i++) {
    if (isPrime(i)) {
      sum += i;
    }
  }
  return sum;
}
```
