---
id: 66b205e6eacba4c4e54ea434
title: 'Crée un masqueur d''adresse e-mail'
challengeType: 26
dashedName: build-an-email-masker
---

# --description--

Dans ce lab, tu vas masquer la partie nom d'utilisateur d'une adresse e-mail avec des astérisques. Le masquage est un terme utilisé pour cacher ou remplacer des informations sensibles par des astérisques ou d'autres caractères.

Par exemple, si l'adresse e-mail était `myEmail@email.com`, alors l'adresse e-mail masquée sera `m*****l@email.com`.

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Crée une fonction nommée `maskEmail` qui prend `email` comme argument.
2. À l'intérieur de la fonction, tu dois masquer l'`email` et lui ajouter le nom de domaine. Souviens-toi que tu peux utiliser des méthodes comme `slice`, `repeat`, `indexOf` ou même `replace` pour t'aider.
3. En dehors de la fonction, déclare une variable nommée `email` pour stocker l'adresse e-mail que tu veux masquer.
4. Appelle la fonction `maskEmail` avec la variable `email` et affiche le résultat dans la console.
5. `maskEmail("apple.pie@example.com")` doit retourner `"a*******e@example.com"`.
6. `maskEmail("freecodecamp@example.com")` doit retourner `"f**********p@example.com"`.
7. `maskEmail("info@test.dev")` doit retourner `"i**o@test.dev"`.
8. `maskEmail("user@domain.org")` doit retourner `"u**r@domain.org"`.

# --hints--

Tu dois définir une fonction nommée `maskEmail`.

```js
assert.isFunction(maskEmail);
```

La fonction `maskEmail` doit prendre une chaîne, `email`, comme argument.

```js
assert.match(maskEmail.toString(), /\s*maskEmail\(\s*\w+\s*\)/);
```

En dehors de la fonction, tu dois avoir une variable `email`.

```js
assert.isDefined(email);
```

Tu dois affecter une adresse e-mail valide à ta variable `email`.

```js
assert.match(email, /^[a-zA-Z0-9._%-]+\+?[a-zA-Z0-9._%-]+@[a-zA-Z0-9.-]+\.[a-zA-Z]{2,}$/);
```

`maskEmail("apple.pie@example.com")` doit retourner `"a*******e@example.com"`.

```js
assert.strictEqual(maskEmail("apple.pie@example.com"),"a*******e@example.com");
```

`maskEmail("freecodecamp@example.com")` doit retourner `"f**********p@example.com"`.

```js
assert.strictEqual(maskEmail("freecodecamp@example.com"), "f**********p@example.com");
```

`maskEmail("info@test.dev")` doit retourner `"i**o@test.dev"`.

```js
assert.strictEqual(maskEmail("info@test.dev"), "i**o@test.dev");
```

`maskEmail("user@domain.org")` doit retourner `"u**r@domain.org"`.

```js
assert.strictEqual(maskEmail("user@domain.org"), "u**r@domain.org");
```

Ta fonction `maskEmail` doit produire le résultat correct.

```js
assert.strictEqual(maskEmail("johnappleseed@email.com"), "j***********d@email.com");
assert.strictEqual(maskEmail("john.doe@example.com"), "j******e@example.com");
assert.strictEqual(maskEmail("jane.smith@sample.org"), "j********h@sample.org");
assert.strictEqual(maskEmail("michael.jordan@testmail.net"), "m************n@testmail.net");
assert.strictEqual(maskEmail("mike+nospam@testmail.net"), "m*********m@testmail.net");
```

Tu dois afficher dans la console le résultat de l'appel de `maskEmail` avec `email` comme argument.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(maskEmail\(email\)\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function maskEmail(email) {
  const atIndex = email.indexOf("@");
  const userName = email.slice(0, atIndex);
  const domain = email.slice(atIndex);

  const maskedName =
    userName[0] + "*".repeat(userName.length - 2) + userName[userName.length - 1] + domain;

  return maskedName;
}

const email = "apple.pie@example.com";
console.log(maskEmail(email));
```
