---
id: 68e707257ea91c2e12d489fb
title: Étape 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Dans cet atelier, tu vas t'entraîner à ajouter des styles personnalisés aux boutons radio en construisant un formulaire de réunion parent-professeur. La structure HTML de base t'a été fournie.

Commence par ajouter un élément `main` avec une classe appelée `container`.

# --hints--

Tu devrais avoir un élément `main`.

```js
assert.exists(document.querySelector("main"));
```

Ton élément `main` devrait avoir une classe appelée `container`.

```js
assert.exists(document.querySelector("main.container"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
--fcc-editable-region--

--fcc-editable-region--
</body>

</html>
```
