---
id: 68e7e729f6d14f69a23786ef
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

Ensuite, dans ton élément `main`, ajoute un élément `h1` avec le texte `Parent Teacher Conference Form` et les classes `title` et `center`.

# --hints--

Tu devrais avoir un élément `h1`.

```js
assert.exists(document.querySelector("h1"));
```

Ton élément `h1` devrait avoir le texte `Parent Teacher Conference Form`.

```js
assert.strictEqual(document.querySelector("h1")?.innerText, "Parent Teacher Conference Form");
```

Ton élément `h1` devrait avoir les classes `title` et `center`.

```js
assert.exists(document.querySelector("h1.title.center"));
```

Ton élément `h1` devrait être dans l'élément `main`.

```js
assert.exists(document.querySelector("main h1.title.center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
--fcc-editable-region--

--fcc-editable-region--
  </main>
</body>

</html>
```
