---
id: 68ea8fdf2c9b167307e7f671
title: Étape 9
challengeType: 0
dashedName: step-9
---

# --description--

Il est maintenant temps d'ajouter les éléments de formulaire pour collecter les informations du parent.

Commence par ajouter un élément `label` avec le texte `Parent/Guardian Name: ` et l'attribut `for` défini sur `"parent-name"`.

Ensuite, sous ton élément `label`, ajoute un élément `input` avec l'attribut `type` défini sur `"text"`. Les attributs `name` et `id` devraient être définis sur `"parent-name"`. L'attribut `placeholder` devrait être défini sur `"E.g., Nancy Doe"`. Enfin, ton `input` devrait être requis.

# --hints--

Tu devrais avoir un total de trois éléments `label` sur la page.

```js
assert.lengthOf(document.querySelectorAll("label"), 3);
```

Ton troisième élément `label` devrait avoir le texte `Parent/Guardian Name: `. N'oublie pas l'espace après les deux-points.

```js
assert.match(code, /Parent\/Guardian\s+Name:\s+<\/label>/);
```

Ton troisième élément `label` devrait avoir un attribut `for`.

```js
assert.isTrue(document.querySelector("label:last-of-type")?.hasAttribute("for"));
```

L'attribut `for` de ton troisième élément `label` devrait avoir la valeur `"parent-name"`.

```js
assert.strictEqual(document.querySelector("fieldset:last-of-type label")?.getAttribute("for"), "parent-name");
```

Tu devrais avoir trois éléments `input` sur la page.

```js
assert.lengthOf(document.querySelectorAll("input"), 3);
```

Ton troisième élément `input` devrait être sous ton troisième élément `label`.

```js
const fieldsetChildren = [...document.querySelector("fieldset:last-of-type")?.children];
const inputEl = document.querySelector("fieldset:last-of-type input");
const labelEl = document.querySelector("fieldset:last-of-type label");

assert.isBelow(fieldsetChildren.indexOf(labelEl), fieldsetChildren.indexOf(inputEl));
```

Ton troisième élément `input` devrait avoir un attribut `type` défini sur `"text"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("type"), "text");
```

Ton troisième élément `input` devrait avoir un attribut `name` défini sur `"parent-name"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("name"), "parent-name");
```

Ton troisième élément `input` devrait avoir un attribut `id` défini sur `"parent-name"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("id"), "parent-name");
```

Ton troisième élément `input` devrait avoir un attribut `placeholder` défini sur `"E.g., Nancy Doe"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., Nancy Doe");
```

Ton troisième élément `input` devrait avoir un attribut `required`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
      --fcc-editable-region--

      --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
