---
id: 68ea8c330ed07b60879f3d38
title: Étape 7
challengeType: 0
dashedName: step-7
---

# --description--

Il est maintenant temps d'ajouter les éléments de formulaire pour collecter les informations sur la classe de l'élève.

Commence par ajouter un autre élément `label` avec le texte `Student Grade: ` et l'attribut `for` défini sur `"grade"`.

Ensuite, sous ton élément `label`, ajoute un élément `input` avec l'attribut `type` défini sur `"number"`. Les attributs `name` et `id` devraient être définis sur `"grade"`. L'attribut `placeholder` devrait être défini sur `"E.g., 4"`. Enfin, ton `input` de nombre devrait être requis.

# --hints--

Tu devrais avoir un total de deux éléments `label` sur la page.

```js
assert.lengthOf(document.querySelectorAll("label"), 2);
```

Ton deuxième élément `label` devrait avoir le texte `Student Grade: `. N'oublie pas l'espace après les deux-points.

```js
assert.match(code, /Student\s+Grade:\s+<\/label>/);
```

Ton deuxième élément `label` devrait avoir un attribut `for`.

```js
assert.isTrue(document.querySelector("label:last-of-type")?.hasAttribute("for"));
```

L'attribut `for` de ton deuxième élément `label` devrait avoir la valeur `"grade"`.

```js
assert.strictEqual(document.querySelector("label:last-of-type")?.getAttribute("for"), "grade");
```

Tu devrais avoir deux éléments `input` sur la page.

```js
assert.lengthOf(document.querySelectorAll("input"), 2);
```

Ton deuxième élément `input` devrait être sous ton deuxième élément `label`.

```js
const fieldsetChildren = [...document.querySelector("fieldset")?.children];
const inputEl = document.querySelector("fieldset input:last-of-type");
const labelEl = document.querySelector("label:last-of-type");

assert.isBelow(fieldsetChildren.indexOf(labelEl), fieldsetChildren.indexOf(inputEl));
```

Ton deuxième élément `input` devrait avoir un attribut `type` défini sur `"number"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("type"), "number");
```

Ton deuxième élément `input` devrait avoir un attribut `name` défini sur `"grade"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("name"), "grade");
```

Ton deuxième élément `input` devrait avoir un attribut `id` défini sur `"grade"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("id"), "grade");
```

Ton deuxième élément `input` devrait avoir un attribut `placeholder` défini sur `"E.g., 4"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., 4");
```

Ton deuxième élément `input` devrait avoir un attribut `required`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />
        --fcc-editable-region--

        --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
