---
id: 68ea81da9f687732368a5322
title: Étape 4
challengeType: 0
dashedName: step-4
---

# --description--

Il est maintenant temps d'ajouter les éléments `form` et `input`, qui représenteront les informations du parent et de l'élève.

Commence par ajouter un élément `form` sous l'élément `p`.

# --hints--

Tu devrais avoir un élément `form`.

```js
assert.exists(document.querySelector(".container form"));
```

L'élément `form` devrait être sous l'élément `p`.

```js
assert.exists(document.querySelector(".container > p+form"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>
    --fcc-editable-region--

    --fcc-editable-region--
  </main>
</body>

</html>
```
