---
id: 68eaa668d4d9fbe529a0e16b
title: Étape 14
challengeType: 0
dashedName: step-14
---

# --description--

Maintenant, ajoute un élément `label` avec le texte `Any specific concerns or topics you'd like to discuss?` et un attribut `for` défini sur `notes`.

Sous ce `label`, ajoute un élément `textarea`. Définis son `id` sur `notes`, son `name` sur `notes`, ses `rows` sur `4` et ses `cols` sur `50`.

# --hints--

Ton dernier `fieldset` devrait contenir un élément `label`.

```js
assert.exists(document.querySelector("fieldset:last-of-type label"));
```

L'élément `label` devrait avoir le texte `Any specific concerns or topics you'd like to discuss?`.

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.innerText, "Any specific concerns or topics you'd like to discuss?");
```

L'attribut `for` de l'élément `label` devrait être défini sur `"notes"`.

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.getAttribute("for"), "notes");
```

Ton dernier `fieldset` devrait contenir un élément `textarea`.

```js
assert.exists(document.querySelector("fieldset:last-of-type textarea"));
```

Ton élément `textarea` devrait être sous l'élément `label`.

```js
const label = document.querySelector("fieldset:last-of-type label");
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(label?.nextElementSibling, textarea);
```

L'élément `textarea` devrait avoir un attribut `id` défini sur `"notes"`.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("id"), "notes");
```

L'élément `textarea` devrait avoir un attribut `name` défini sur `"notes"`.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("name"), "notes");
```

L'élément `textarea` devrait avoir un attribut `rows` défini sur `4`.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("rows"), "4");
```

L'élément `textarea` devrait avoir un attribut `cols` défini sur `50`.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("cols"), "50");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        --fcc-editable-region--

        --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
