---
id: 68ea8666b8b473455f5b33ef
title: Étape 5
challengeType: 0
dashedName: step-5
---

# --description--

La première section du formulaire se concentrera sur les informations de l'élève, comme son nom et sa classe.

Dans l'élément `form`, ajoute un élément `fieldset`. Dans cet élément `fieldset`, ajoute un élément `legend` avec le texte `Student Information`.

Sous l'élément `legend`, ajoute un élément `label` avec le texte `Full Name: ` et un attribut `for` avec la valeur `student-name`.


# --hints--

Tu devrais avoir un élément `fieldset`.

```js
assert.exists(document.querySelector("fieldset"));
```

Ton élément `fieldset` devrait être dans l'élément `form`.

```js
assert.exists(document.querySelector("form fieldset"));
```

Tu devrais avoir un élément `legend`.

```js
assert.exists(document.querySelector("legend"));
```

Ton élément `legend` devrait être dans l'élément `fieldset`.

```js
assert.exists(document.querySelector("fieldset legend"));
```

Ton élément `legend` devrait avoir le texte `Student Information`.

```js
assert.strictEqual(document.querySelector("fieldset legend")?.innerText, "Student Information");
```

Tu devrais avoir un élément `label`.

```js
assert.exists(document.querySelector("label"));
```

Ton élément `label` devrait avoir le texte `Full Name: `. N'oublie pas l'espace après les deux-points.

```js
assert.match(code, /Full\s+Name:\s+<\/label>/);
```

Ton élément `label` devrait avoir un attribut `for`.

```js
assert.isTrue(document.querySelector("label")?.hasAttribute("for"));
```

L'attribut `for` de ton élément `label` devrait avoir la valeur `student-name`.

```js
assert.strictEqual(document.querySelector("label")?.getAttribute("for"), "student-name");
```

Ton élément `label` devrait être sous l'élément `legend`.

```js
const fieldsetChildren = [...document.querySelector("fieldset")?.children];
const legendEl = document.querySelector("fieldset legend");
const labelEl = document.querySelector("label");

assert.isBelow(fieldsetChildren.indexOf(legendEl), fieldsetChildren.indexOf(labelEl));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
    --fcc-editable-region--

    --fcc-editable-region--
    </form>
  </main>
</body>

</html>
```
