---
id: 68ea975c7cbbb894a386e402
title: Étape 11
challengeType: 0
dashedName: step-11
---

# --description--

L'étape suivante consiste à ajouter les éléments `label` et `input` pour la méthode de contact par e-mail.

Commence par ajouter un élément `label` avec une classe de `"contact-method"` et un attribut `for` défini sur `"email"`. Le texte du `label` devrait être `Email: `.

Sous ton élément `label`, ajoute un bouton radio avec les attributs `id` et `value` définis sur `"email"`. L'attribut `name` devrait être défini sur `"contact-method"` et la classe devrait être définie sur `"contact-method-radio-btn"`.

Enfin, assure-toi que ce bouton radio est coché par défaut.

# --hints--


Ton dernier `fieldset` devrait contenir un élément `label`.

```js
assert.exists(document.querySelector("fieldset:last-of-type label"));
```

Ce nouvel élément `label` (pour l'e-mail) devrait avoir une classe de `"contact-method"`.

```js
assert.exists(document.querySelector("fieldset:last-of-type label.contact-method"));
```

L'attribut `for` du nouvel élément `label` devrait être défini sur `"email"`.

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.getAttribute("for"), "email");
```

L'élément `label` devrait avoir le texte `Email: `. Assure-toi de garder l'espace après les deux-points.

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.match(label?.innerText, /Email:\s?/);
```

Tu devrais avoir un élément `input` dans le dernier `fieldset`.

```js
assert.exists(document.querySelector("fieldset:last-of-type input"));
```

Ton élément `input` devrait avoir un attribut `type` défini sur `radio`.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("type"), "radio");
```

L'élément `input` devrait avoir un attribut `id` défini sur `"email"`.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("id"), "email");
```

L'élément `input` devrait avoir un attribut `value` défini sur `"email"`.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("value"), "email");
```

L'élément `input` devrait avoir un attribut `name` défini sur `"contact-method"`.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("name"), "contact-method");
```

L'élément `input` devrait avoir une classe de `"contact-method-radio-btn"`.

```js
assert.exists(document.querySelector("fieldset:last-of-type input.contact-method-radio-btn"));
```

Ton bouton `radio` devrait être coché par défaut.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.isTrue(input?.hasAttribute("checked"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
      --fcc-editable-region--

      --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
