---
id: 68ea95c833f0518ae6e366cd
title: Étape 10
challengeType: 0
dashedName: step-10
---

# --description--

Dans les prochaines étapes, tu ajouteras les éléments de formulaire chargés de collecter la méthode de contact préférée de l'utilisateur.

Commence par ajouter un autre élément `fieldset` avec un élément `legend` imbriqué dedans. Ton élément `legend` devrait avoir le texte `Preferred Contact Method`.

# --hints--

Tu as un total de trois éléments `fieldset` sur la page.

```js
assert.lengthOf(document.querySelectorAll("fieldset"), 3);
```

Tu devrais avoir un élément `legend` dans ton troisième élément `fieldset`.

```js
assert.exists(document.querySelector("fieldset:last-of-type legend"));
```

Ton élément `legend` devrait avoir le texte `Preferred Contact Method`.

```js
assert.strictEqual(document.querySelector("fieldset:last-of-type legend")?.innerText, "Preferred Contact Method");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>
      --fcc-editable-region--

      --fcc-editable-region--
    </form>
  </main>
</body>

</html>
```
