---
id: 68ea7c41f2555e16174e7cfa
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

Ensuite, ajoute un élément paragraphe avec le texte `Please fill out the form below to help schedule your parent-teacher conference.`. Ton élément paragraphe devrait aussi avoir les classes `description` et `center`.

# --hints--

Tu devrais avoir un élément `p`.

```js
assert.exists(document.querySelector("p"));
```

Ton élément `p` devrait avoir le texte `Please fill out the form below to help schedule your parent-teacher conference.`

```js
assert.strictEqual(document.querySelector("p")?.innerText, "Please fill out the form below to help schedule your parent-teacher conference.");
```

Ton élément `p` devrait avoir les classes `description` et `center`.

```js
assert.exists(document.querySelector("p.description.center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    --fcc-editable-region--

    --fcc-editable-region--
  </main>
</body>

</html>
```
