---
id: 68ea8a5db73b155708992687
title: Étape 6
challengeType: 0
dashedName: step-6
---

# --description--

L'étape suivante consiste à ajouter l'élément `input` associé aux informations de l'élève.

Commence par ajouter un élément `input` avec un attribut `type` défini sur `"text"`. Ajoute ensuite un attribut `name` et un attribut `id`, tous les deux définis sur `student-name`.

Ensuite, ajoute un attribut `placeholder` défini sur `E.g., Jane Doe`. Enfin, ajoute un attribut `required`.

# --hints--

Tu devrais avoir un élément `input`.

```js
assert.exists(document.querySelector("fieldset input"));
```

Ton élément `input` devrait avoir un attribut `type` défini sur `"text"`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("type"), "text");
```

Ton élément `input` devrait avoir un attribut `name` défini sur `"student-name"`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("name"), "student-name");
```

Ton élément `input` devrait avoir un attribut `id` défini sur `"student-name"`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("id"), "student-name");
```

Ton élément `input` devrait avoir un attribut `placeholder` défini sur `E.g., Jane Doe`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., Jane Doe");
```

Ton élément `input` devrait avoir un attribut `required`.

```js
const inputEl = document.querySelector("fieldset input");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>

      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
      --fcc-editable-region--

      --fcc-editable-region--
      </fieldset>


    </form>
  </main>
</body>

</html>
```
