---
id: 68eaa310c153c0d0c5c9b78e
title: Étape 13
challengeType: 0
dashedName: step-13
---

# --description--

C'est une bonne pratique de permettre aux utilisateurs d'ajouter des notes ou préoccupations supplémentaires.

Sous ton troisième élément `fieldset`, ajoute un nouvel élément `fieldset`. À l'intérieur, ajoute un élément `legend` avec le texte `Additional Notes`.


# --hints--

Tu devrais avoir un total de 4 éléments `fieldset` sur cette page.

```js
const fieldsets = document.querySelectorAll('fieldset');
assert.lengthOf(fieldsets, 4);
```

Tu devrais ajouter un élément `fieldset` sous le troisième élément `fieldset`.

```js
const fieldsets = document.querySelectorAll('fieldset');
assert.exists(fieldsets[2]?.nextElementSibling?.matches('fieldset'));
```

Le nouvel élément `fieldset` devrait contenir un élément `legend`.

```js
const newFieldset = document.querySelectorAll('fieldset')[3];
assert.exists(newFieldset?.querySelector('legend'));
```

L'élément `legend` devrait avoir le texte `Additional Notes`.

```js
const newFieldset = document.querySelectorAll('fieldset')[3];
const legend = newFieldset?.querySelector('legend');
assert.equal(legend?.innerText?.trim(), 'Additional Notes');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>
      --fcc-editable-region--

      --fcc-editable-region--
    </form>
  </main>
</body>

</html>
```
