---
id: 68eb058081504963447bfa60
title: Étape 17
challengeType: 0
dashedName: step-17
---

# --description--

Ensuite, tu devrais styliser le conteneur principal. Sélectionne l'élément avec la classe `container` et définis sa propriété `background-color` sur le code hexadécimal `#ffffff1a`.

Un code hexadécimal est une combinaison de six chiffres et lettres utilisée en HTML et CSS pour représenter les couleurs. Il commence par un `#` suivi de trois paires :

- Les deux premiers caractères représentent le rouge
- Les deux suivants représentent le vert
- Les deux derniers représentent le bleu

Par exemple, `#ffffff` est un blanc pur parce qu'il utilise la valeur maximale pour le rouge, le vert et le bleu.

Tu peux aussi ajouter deux caractères supplémentaires à la fin pour contrôler l'opacité (appelée canal alpha). Dans `#ffffff1a`, `1a` rend la couleur blanche semi-transparente. Plus la valeur alpha est faible, plus la couleur paraît transparente. Tu en apprendras davantage sur les codes hexadécimaux plus tard !

# --hints--

Tu devrais cibler l'élément avec la classe `container`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container'));
```

Définis la propriété `background-color` sur `#ffffff1a`.

```js
const style = new __helpers.CSSHelp(document).getStyle('.container')
assert.strictEqual(style?.backgroundColor, 'rgba(255, 255, 255, 0.1)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        <label for="notes"
            >Any specific concerns or topics you'd like to discuss?</label>
        <textarea id="notes" name="notes" rows="4" cols="50"></textarea>
      </fieldset>

      <button class="submit-btn" type="submit">Submit Form</button>
    </form>
  </main>
</body>

</html>
```


```css
body {
  background-color: MidnightBlue;
  color: whitesmoke;
}

--fcc-editable-region--

--fcc-editable-region--
```
