---
id: 68eb372f9b470238bf60577d
title: Étape 34
challengeType: 0
dashedName: step-34
---

# --description--

Maintenant, il est temps de définir ce qui se passe quand le bouton radio est coché.

Tu styles le cercle intérieur uniquement quand le bouton radio est sélectionné, donc ton sélecteur devrait être `.contact-method-radio-btn:checked::before`.

Dans la règle, définis `transform` sur `translate(3px, 3px) scale(1)` et `background-color` sur `lightgreen`.

La propriété `transform` déplace légèrement le cercle intérieur et l'agrandit à sa taille complète avec `scale(1)`, tandis qu'un arrière-plan `lightgreen` remplit le cercle de couleur pour montrer que l'option est sélectionnée, ce qui le rend visible.

# --hints--

Tu devrais cibler le pseudo-élément `.contact-method-radio-btn:checked::before`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.contact-method-radio-btn:checked::before'));
```

Définis la propriété `transform` sur `translate(3px, 3px) scale(1)`.

```js
const style = new __helpers.CSSHelp(document).getStyle('.contact-method-radio-btn:checked::before');
assert.strictEqual(style?.transform, 'translate(3px, 3px) scale(1)');
```

Définis la propriété `background-color` sur `lightgreen`.

```js
const style = new __helpers.CSSHelp(document).getStyle('.contact-method-radio-btn:checked::before');
assert.strictEqual(style?.backgroundColor, 'lightgreen');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        <label for="notes"
            >Any specific concerns or topics you'd like to discuss?</label>
        <textarea id="notes" name="notes" rows="4" cols="50"></textarea>
      </fieldset>

      <button class="submit-btn" type="submit">Submit Form</button>
    </form>
  </main>
</body>

</html>
```

```css
body {
  background-color: MidnightBlue;
  color: whitesmoke;
}

.container {
    background-color: #ffffff1a;
    width: 80%;
    max-width: 600px;
    border-radius: 10px;
    margin: 20px auto;
    padding: 10px 20px;
    box-shadow: 0 5px 15px black;
}

.center {
  text-align: center;
}

.description {
  font-size: 1.2rem;
}

fieldset {
  border: 1px solid gray;
  border-radius: 5px;
  margin: 20px 0;
  padding: 20px;
}

fieldset legend {
  font-size: 1.3rem;
  font-weight: 600;
}

label {
  font-size: 1.2rem;
}

label:not(.contact-method) {
  display: block;
  margin: 10px 0;
}

input:not(.contact-method-radio-btn),
textarea {
  background-color: #ffffff1a;
  width: 95%;
  border: 1px solid gray;
  border-radius: 5px;
  padding: 10px;
}

input,
input::placeholder,
textarea {
  color: whitesmoke;
}


.contact-method-radio-btn {
  appearance: none;
  width: 20px;
  height: 20px;
  border-radius: 50%;
  border: 2px solid gray;
  vertical-align: bottom;
}

.contact-method-radio-btn::before {
  display: block;
  content: " ";
  width: 10px;
  height: 10px;
  border-radius: 50%;
  transform: translate(3px, 3px) scale(0);
  transition: all 0.3s ease-in;
}

--fcc-editable-region--

--fcc-editable-region--
```
