---
id: 6733b5632fb854bc6c58dc0d
title: Étape 2
challengeType: 1
dashedName: step-2
---

# --description--

Crée maintenant un autre objet, avec les mêmes trois propriétés, pour représenter un autre cultivar de lavande espagnole appelé `Pretty Polly`, et assigne-le à une variable nommée `prettyPolly`.

# --hints--

Tu dois avoir une variable nommée `prettyPolly`.

```js
assert.exists(prettyPolly);
```

Tu dois assigner un objet à ta variable `prettyPolly`.

```js
assert.isObject(prettyPolly);
```

Ton objet `prettyPolly` doit avoir une propriété `commonName` définie à `"Spanish lavender"`.

```js
assert.strictEqual(prettyPolly.commonName, "Spanish lavender");
```

Ton objet `prettyPolly` doit avoir une propriété `scientificName` définie à `"Lavandula stoechas"`.

```js
assert.strictEqual(prettyPolly.scientificName, "Lavandula stoechas");
```

Ton objet `prettyPolly` doit avoir une propriété `cultivar` définie à `"Pretty Polly"`.

```js
assert.strictEqual(prettyPolly.cultivar, "Pretty Polly");
```

Ton objet `prettyPolly` doit avoir trois clés.

```js
assert.lengthOf(Object.keys(prettyPolly), 3);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}
--fcc-editable-region--

--fcc-editable-region--
```
