---
id: 6733b5e6f210c1bd1afdf8b3
title: Étape 3
challengeType: 1
dashedName: step-3
---

# --description--

Les objets plantes restants ont été ajoutés pour toi. Il est temps de créer ton catalogue.

Comme les objets JavaScript classiques, les objets `Map` contiennent une collection de paires clé-valeur. Pour créer un objet `Map`, tu peux utiliser le constructeur `Map` :

```js
const myMap = new Map();
```

Déclare une variable `catalog` et assigne-lui un objet `Map`.

# --hints--

Tu dois avoir une variable nommée `catalog`.

```js
assert.exists(catalog);
```

Tu dois assigner un objet `Map` à ta variable `catalog`.

```js
assert.instanceOf(catalog, Map);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--

--fcc-editable-region--
```
