---
id: 673339d96c56d36e963e888e
title: Étape 1
challengeType: 1
dashedName: step-1
---

# --description--

Dans ce workshop, tu vas construire un catalogue de pépinière tout en pratiquant les objets Map et Set de JavaScript.

Pour représenter chaque plante de ton catalogue, tu utiliseras des objets, chacun avec trois propriétés :

- `commonName`, le nom commun de la plante
- `scientificName`, le nom scientifique (ou nom latin) de la plante
- `cultivar`, la variété de la plante

Commence par créer un objet avec les trois propriétés ci-dessus pour le `Spanish lavender` (nom scientifique `Lavandula stoechas`) d'un cultivar appelé `Ballerina` et assigne-le à une variable nommée `ballerina`.

# --hints--

Tu dois avoir une variable nommée `ballerina`.

```js
assert.exists(ballerina);
```

Tu dois assigner un objet à ta variable `ballerina`.

```js
assert.isObject(ballerina);
```

Ton objet `ballerina` doit avoir une propriété `commonName` définie à `"Spanish lavender"`.

```js
assert.strictEqual(ballerina.commonName, "Spanish lavender");
```

Ton objet `ballerina` doit avoir une propriété `scientificName` définie à `"Lavandula stoechas"`.

```js
assert.strictEqual(ballerina.scientificName, "Lavandula stoechas");
```

Ton objet `ballerina` doit avoir une propriété `cultivar` définie à `"Ballerina"`.

```js
assert.strictEqual(ballerina.cultivar, "Ballerina");
```

Ton objet `ballerina` doit avoir trois clés.

```js
assert.lengthOf(Object.keys(ballerina), 3);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
