---
id: 67349ba230cff237c0f2bdf2
title: Étape 4
challengeType: 1
dashedName: step-4
---

# --description--

Les objets Map mémorisent l'ordre dans lequel les clés sont insérées dans la map. Pour définir une paire clé-valeur dans ton objet map, tu peux utiliser la méthode `set`. L'exemple suivant définit la clé `"name"` à la valeur `"Jerry"` dans l'objet map `myMap`.

```js
myMap.set("name", "Jerry");
```

Définis la clé `"ballerina"` à la valeur de ton objet `ballerina` dans ta map `catalog`. Ensuite, affiche ton catalogue dans la console.

Note que pour voir correctement ton objet map, tu devras ouvrir la console du navigateur.

# --hints--

Tu dois définir la clé `"ballerina"` à la valeur de ton objet `ballerina` dans ta map `catalog`.

```js
assert.deepEqual(catalog.get("ballerina"), ballerina)
```

Ta map `catalog` doit contenir une seule paire clé-valeur.

```js
assert.lengthOf(catalog, 1);
```

Tu dois afficher ta map `catalog` dans la console.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*catalog\s*\)/)
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();

--fcc-editable-region--
```
