---
id: 6734ab61e916486c80c490d1
title: Étape 11
challengeType: 1
dashedName: step-11
---

# --description--

Quand la valeur associée à une clé est un objet, la méthode `get` renvoie une référence à cet objet. Tu peux utiliser cette référence pour modifier l'objet à l'intérieur de la map.

Pour le tester, définis la propriété `small` de `getBallerina` à un entier différent de `20`. Ensuite, regarde la console du navigateur pour vérifier que la valeur a changé dans ton catalogue.

# --hints--

Tu dois définir la propriété `small` de `getBallerina` à un entier différent de `20`.

```js
const small = catalog.get(ballerina).small;
assert.match(__helpers.removeJSComments(code), /getBallerina\s*(?:\.\s*small|\[\s*("|')small\1\s*\])\s*=\s*\d+/);
assert.notEqual(small, 20);
assert.isTrue(Number.isInteger(small));
assert.isAtLeast(small, 0);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });

const getBallerina = catalog.get(ballerina);
console.log(getBallerina);
--fcc-editable-region--

--fcc-editable-region--
console.log(catalog);
```
