---
id: 6734a2385c52c54d916a8b86
title: Étape 7
challengeType: 1
dashedName: step-7
---

# --description--

Ajoute une autre paire clé-valeur à ton catalogue. Utilise ta variable `prettyPolly` comme clé et fais que la valeur soit un objet avec les propriétés `small`, `medium` et `large` dont les valeurs sont à nouveau des entiers.

# --hints--

Tu dois définir ton objet `prettyPolly` comme clé de ta map `catalog`.

```js
assert.isTrue(catalog.has(prettyPolly));
```

La valeur de la clé `prettyPolly` dans la map `catalog` doit être un objet avec les propriétés `small`, `medium` et `large`.

```js
const b = catalog.get(prettyPolly)
assert.property(b, "small");
assert.property(b, "medium");
assert.property(b, "large");
```

La propriété `small` de l'objet associé à la clé `prettyPolly` dans la map `catalog` doit être un entier positif.

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
```

La propriété `medium` de l'objet associé à la clé `prettyPolly` dans la map `catalog` doit être un entier positif.

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
```

La propriété `large` de l'objet associé à la clé `prettyPolly` dans la map `catalog` doit être un entier positif.

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Tu dois toujours avoir la clé `ballerina` dans ta map `catalog`.

```js
assert.isTrue(catalog.has(ballerina));
```

Ta map `catalog` doit contenir deux éléments.

```js
assert.lengthOf(catalog, 2);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });

--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
