---
id: 6734c10a327760a665b7d5b9
title: Étape 13
challengeType: 1
dashedName: step-13
---

# --description--

Dans les prochaines étapes, tu vas travailler sur une fonction qui te permettra de mettre à jour le catalogue lorsque tu vends des plantes.

Crée une fonction nommée `sellPlants` qui prend trois arguments : un objet plante, la taille du pot et le nombre de pots à vendre, dans cet ordre.

Quand le nombre de pots spécifié est supérieur aux pots disponibles, fais renvoyer à ta fonction `Not enough <pot-size> size pots for <item-name>. Only <pot-number> left.`. Remplace `<pot-size>` par la taille du pot, `<pot-number>` par les pots restants pour cette taille, et `<item-name>` par le nom scientifique de la plante suivi d'un espace et du nom du cultivar entouré de guillemets simples.

Pour vérifier que tout fonctionne, affiche dans la console le résultat de l'appel de ta fonction `sellPlants` avec l'objet `ballerina`, `"small"` et `25`.

# --hints--

Tu ne dois pas modifier ta map `catalog`.

```js
const testCatalog = new Map();
testCatalog.set(ballerina, { small: 20, medium: 15, large: 12 });
testCatalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
testCatalog.set(willowVale, { small: 3, medium: 5, large: 0 });
testCatalog.set(hidcote, { small: 33, medium: 13, large: 18 });
testCatalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
testCatalog.set(royalCrown, { small: 40, medium: 22, large: 9 });
assert.deepEqual(catalog, testCatalog);
```

Tu dois avoir une fonction nommée `sellPlants`.

```js
assert.isFunction(sellPlants);
```

Ta fonction `sellPlants` doit prendre trois arguments.

```js
assert.lengthOf(sellPlants, 3);
```

`sellPlants(ballerina, "small", 25)` devrait renvoyer `Not enough small size pots for Lavandula stoechas 'Ballerina'. Only 20 left.`

```js
assert.equal(sellPlants(ballerina, "small", 25), "Not enough small size pots for Lavandula stoechas 'Ballerina'. Only 20 left.")

// prevent hardcoding
assert.equal(sellPlants(hidcote, "large", 95), "Not enough large size pots for Lavandula angustifolia 'Hidcote'. Only 18 left.")

assert.equal(sellPlants(royalCrown, "medium", 95), "Not enough medium size pots for Lavandula dentata 'Royal Crown'. Only 22 left.")
```

Tu dois afficher `sellPlants(ballerina, "small", 25)` dans la console.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*sellPlants\s*\(\s*ballerina\s*,\s*("|')small\1\s*,\s*25\s*\)\s*\)/);
```

`sellPlants(ballerina, "small", 10)` devrait renvoyer `undefined`.

```js
assert.isUndefined(sellPlants(ballerina, "small", 10));
// prevent hardcoding
assert.isUndefined(sellPlants(hidcote, "large", 5));
assert.isUndefined(sellPlants(royalCrown, "medium", 5));
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });
--fcc-editable-region--

--fcc-editable-region--
console.log(catalog)
```
