---
id: 6734a606e5968a5a19b77171
title: Étape 8
challengeType: 1
dashedName: step-8
---

# --description--

Continue d'ajouter les plantes restantes à ton catalogue comme tu l'as fait dans les étapes précédentes.

# --hints--

Tu dois définir ton objet `willowVale` comme clé de ta map `catalog`.

```js
assert.isTrue(catalog.has(willowVale));
```

La valeur de la clé `willowVale` dans la map `catalog` doit être un objet avec les propriétés `small`, `medium` et `large`, chacune avec la valeur d'un entier positif.

```js
const b = catalog.get(willowVale)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Tu dois définir ton objet `hidcote` comme clé de ta map `catalog`.

```js
assert.isTrue(catalog.has(hidcote));
```

La valeur de la clé `hidcote` dans la map `catalog` doit être un objet avec les propriétés `small`, `medium` et `large`, chacune avec la valeur d'un entier positif.

```js
const b = catalog.get(hidcote)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Tu dois définir ton objet `imperialGem` comme clé de ta map `catalog`.

```js
assert.isTrue(catalog.has(imperialGem));
```

La valeur de la clé `imperialGem` dans la map `catalog` doit être un objet avec les propriétés `small`, `medium` et `large`, chacune avec la valeur d'un entier positif.

```js
const b = catalog.get(imperialGem)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Tu dois définir ton objet `royalCrown` comme clé de ta map `catalog`.

```js
assert.isTrue(catalog.has(royalCrown));
```

La valeur de la clé `royalCrown` dans la map `catalog` doit être un objet avec les propriétés `small`, `medium` et `large`, chacune avec la valeur d'un entier positif.

```js
const b = catalog.get(royalCrown)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Tu dois toujours avoir la clé `ballerina` dans ta map `catalog`.

```js
assert.isTrue(catalog.has(ballerina));
```

Tu dois toujours avoir la clé `prettyPolly` dans ta map `catalog`.

```js
assert.isTrue(catalog.has(prettyPolly));
```

Ta map `catalog` doit contenir six éléments.

```js
assert.lengthOf(catalog, 6);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });

--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
