---
id: 6734a13a29acdf47f2b8e161
title: Étape 6
challengeType: 1
dashedName: step-6
---

# --description--

Contrairement aux objets classiques, pour lesquels les clés doivent toujours être des chaînes, les objets map te permettent d'utiliser n'importe quel type de données comme clés. Dans ton catalogue, tu vas utiliser des paires clé-valeur objet-objet, en associant chaque objet plante à un objet contenant le nombre de pots disponibles pour différentes tailles.

Modifie ton appel à `set` pour utiliser ton objet `ballerina` comme clé. Pour la valeur, utilise un objet contenant trois propriétés, `small`, `medium` et `large`, chacune avec la valeur d'un entier positif.

# --hints--

Tu dois définir ton objet `ballerina` comme clé de ta map `catalog`.

```js
assert.isTrue(catalog.has(ballerina));
```

La valeur de la clé `ballerina` dans la map `catalog` doit être un objet avec les propriétés `small`, `medium` et `large`.

```js
const b = catalog.get(ballerina)
assert.property(b, "small");
assert.property(b, "medium");
assert.property(b, "large");
```

La propriété `small` de l'objet associé à la clé `ballerina` dans la map `catalog` doit être un entier positif.

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
```

La propriété `medium` de l'objet associé à la clé `ballerina` dans la map `catalog` doit être un entier positif.

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
```

La propriété `large` de l'objet associé à la clé `ballerina` dans la map `catalog` doit être un entier positif.

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Ta map `catalog` doit contenir une seule paire clé-valeur.

```js
assert.lengthOf(catalog, 1);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set("ballerina", ballerina);
--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
