---
id: 6734a772196245601e4b80e0
title: Étape 10
challengeType: 1
dashedName: step-10
---

# --description--

Quand tu as besoin de récupérer un élément d'une map, tu peux utiliser la méthode `get`. Elle prend une clé comme argument et renvoie la valeur associée.

Utilise la méthode `get` pour récupérer la valeur de la clé `ballerina` depuis ton catalogue et assigne-la à une variable nommée `getBallerina`. Ensuite, affiche ta nouvelle variable dans la console.

# --hints--

Tu dois avoir une variable nommée `getBallerina`.

```js
assert.exists(getBallerina);
```

Ta variable `getBallerina` doit avoir la valeur de `catalog.get(ballerina)`.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+getBallerina\s*=\s*catalog\s*\.\s*get\s*\(\s*ballerina\s*\)/);
```

Tu dois afficher ta variable `getBallerina` dans la console.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*getBallerina\s*\)/);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });
--fcc-editable-region--

--fcc-editable-region--
console.log(catalog);
```
