---
id: 673521668be7905059b2d555
title: Étape 22
challengeType: 1
dashedName: step-22
---

# --description--

Enfin, tu peux utiliser la méthode `forEach` pour exécuter une fonction sur chaque paire clé-valeur de ta map.

```js
myMap.forEach((value, key) => {console.log(key, value)});
```

Note que le callback passé à `forEach` accède à la valeur via son premier argument et à la clé via son deuxième argument pour chaque entrée.

Modifie ta fonction pour qu'elle renvoie une chaîne contenant tout ton catalogue et affiche le résultat dans la console. Construis ta chaîne en utilisant la structure suivante pour chaque entrée de `catalog` :

```md
<scientific-name> '<cultivar>': <small-pots> S, <medium-pots> M, <large-pots> L

```

Note que la représentation de chaque entrée doit se terminer par un caractère de nouvelle ligne.

# --hints--

Tu dois afficher le résultat de l'appel à `displayCatalog` dans la console.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*displayCatalog\s*\(\s*\)\s*\)/)
```

Ta fonction `displayCatalog` doit utiliser la méthode `forEach` sur `catalog`.

```js
let flag = false;
const temp = catalog.forEach;
try {
    catalog.forEach = obj => {flag = true};
    displayCatalog();
    assert.isTrue(flag);
} finally {
    catalog.forEach = temp;
}
```

`displayCatalog` devrait renvoyer `--fcc-expected--`, mais `--fcc-actual--` a été trouvé.

```js
const displayCatalogTest = () => {
    let catalogString = "";
    catalog.forEach((val, key) => {
        catalogString += `${key.scientificName} '${key.cultivar}': ${val.small} S, ${val.medium} M, ${val.large} L
`   })
    return catalogString
}
assert.equal(displayCatalog(), displayCatalogTest())
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });

const removePlant = plant => catalog.delete(plant);

const sellPlants = (plant, size, potsNo) => {
    if (!catalog.has(plant)) return "Item not found.";
    const name = `${plant.scientificName} '${plant.cultivar}'`
    const pots = catalog.get(plant);
    if (pots[size] - potsNo < 0) {
        return `Not enough ${size} size pots for ${name}. Only ${pots[size]} left.`
    }
    pots[size] -= potsNo;
    return `Catalog successfully updated.`
}
--fcc-editable-region--
const displayCatalog = () => {
    for (const entry of catalog.entries()) {
        console.log(entry);
    }
}

displayCatalog()
--fcc-editable-region--
```
