---
id: 680a431692a5f863f726c438
title: Étape 3
challengeType: 20
dashedName: step-3
---

# --description--

Dans ce workshop, tu vas construire un *chiffrement de César*. C'est l'une des techniques les plus simples pour chiffrer du texte : elle consiste à remplacer chaque lettre du texte en clair par la lettre située un nombre fixe de positions plus loin dans l'alphabet. Par exemple, avec un décalage de 5, `a` serait remplacé par `f`, `b` par `g`, et ainsi de suite.

Pour implémenter ce chiffrement, tu vas créer une nouvelle version de ton alphabet qui commence à la position indiquée par le décalage. Comme tu l'as appris dans une leçon précédente, tu peux extraire une partie d'une chaîne avec le découpage de chaîne :

```py
fcc = 'freeCodeCamp'
print(fcc[8:]) # Camp
```

Crée une variable nommée `shifted_alphabet` et utilise la syntaxe de découpage pour lui assigner la portion de `alphabet` qui commence à l'indice de `shift`. Puis appelle la fonction intégrée `print()` pour afficher `shifted_alphabet` dans le terminal et observe le résultat.

# --hints--

Tu devrais déclarer une nouvelle variable nommée `shifted_alphabet`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shifted_alphabet")`)) })

```

Tu devrais assigner `alphabet[shift:]` à `shifted_alphabet`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shifted_alphabet").is_equivalent("shifted_alphabet = alphabet[shift:]")`)) })
```

Tu devrais appeler la fonction `print()` en lui passant `shifted_alphabet` comme argument.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
--fcc-editable-region--

--fcc-editable-region--
```
