---
id: 680a549f60f5a3fedfc0edd9
title: Étape 4
challengeType: 20
dashedName: step-4
---

# --description--

Comme tu peux le voir dans la sortie, l'alphabet décalé commence à la lettre `f` parce que `shift` a la valeur `5`. Mais les cinq premières lettres de l'alphabet – `a`, `b`, `c`, `d` et `e` – manquent maintenant dans l'alphabet décalé, tu vas donc devoir les ajouter à la fin. Tu peux utiliser l'opérateur `+` pour concaténer la portion manquante de l'alphabet.

Modifie l'affectation existante de la variable `shifted_alphabet` : utilise la syntaxe de découpage pour extraire la première portion manquante de `alphabet` et concatène-la à `alphabet[shift:]`.

Pour rappel, `sentence[start:stop]` renvoie les caractères de `sentence` de la position `start` (incluse) à `stop` (exclue).

# --hints--

Tu ne devrais pas modifier directement la variable `alphabet`.

```js
({ test: () => runPython(`
  assert alphabet == "abcdefghijklmnopqrstuvwxyz"
`) })
```

Tu devrais assigner la concaténation de `alphabet[shift:]` et de la première portion manquante de `alphabet` à la variable `shifted_alphabet` sur la même ligne où elle a été déclarée. Utilise `shift` pour indiquer où arrêter le découpage.

```js
({ test: () => runPython(`
_shifted_alpha = _Node(_code).find_variable("shifted_alphabet")
_first = _shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[:shift]")
_second =_shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[0:shift]")
assert _first or _second
`) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
--fcc-editable-region--
shifted_alphabet = alphabet[shift:]
--fcc-editable-region--
print(shifted_alphabet)
```
