---
id: 680a41ef5bfdc059cfd9d4b5
title: Étape 2
challengeType: 20
dashedName: step-2
---

# --description--

Déclare une autre variable appelée `alphabet` et assigne-lui la chaîne `abcdefghijklmnopqrstuvwxyz`.

# --hints--

Tu devrais déclarer une variable appelée `alphabet`. Place bien le nom de la variable au début de la ligne.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("alphabet")`)) })
```

Tu devrais assigner la chaîne `abcdefghijklmnopqrstuvwxyz` à ta variable `alphabet`. N'oublie pas d'entourer la chaîne de quotes simples ou doubles.

```js
({ test: () => assert(runPython(`
_Node(_code).find_variable("alphabet").is_equivalent("alphabet = 'abcdefghijklmnopqrstuvwxyz'")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
shift = 5
```
