---
id: 680a7e200400089967b8b1f4
title: Étape 6
challengeType: 20
dashedName: step-6
---

# --description--

La méthode `str.maketrans()` prend deux chaînes de même longueur et renvoie une table de traduction qui associe chaque caractère de la première chaîne au caractère correspondant de la seconde. Chaque caractère de la table de traduction est stocké comme un ordinal Unicode, un nombre qui identifie le caractère de façon unique.

```py
lower_chars = 'abc'
upper_chars = 'ABC'

table = str.maketrans(lower_chars, upper_chars)
print(table) # {97: 65, 98: 66, 99: 67}
```

Plus tard dans le parcours, tu en apprendras davantage sur le type de structure renvoyé par `str.maketrans()`. Pour l'instant, crée une table de traduction qui associe les caractères de `alphabet` aux caractères de `shifted_alphabet` et assigne-la à une variable nommée `translation_table`.

# --hints--

Tu devrais déclarer une variable nommée `translation_table`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("translation_table")`)) })
```

Tu devrais appeler `str.maketrans()` en lui passant `alphabet` et `shifted_alphabet` comme arguments, et assigner le résultat à `translation_table`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("translation_table").is_equivalent("translation_table = str.maketrans(alphabet, shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--

--fcc-editable-region--
```
