---
id: 680a836235461fce1ff53b3f
title: Étape 8
challengeType: 20
dashedName: step-8
---

# --description--

La méthode `translate()` prend comme argument la table de traduction générée par `str.maketrans()`. Elle s'appelle sur une chaîne et renvoie une copie de la chaîne d'origine où les caractères ont été remplacés d'après la table de traduction :

```py
t = str.maketrans('lk', 'br')
sentence = 'The tent gave in to the leaks.'

print(sentence.translate(t))
# Output: The tent gave in to the bears.
```

Appelle la méthode `translate()` sur `text` en lui passant `translation_table` comme argument et assigne le résultat à une variable nommée `encrypted_text`.

# --hints--

Tu devrais avoir une variable nommée `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("encrypted_text")`)) })
```

Tu devrais assigner `text.translate(translation_table)` à `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = text.translate(translation_table)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
text = 'hello world'
--fcc-editable-region--

--fcc-editable-region--
```
