---
id: 680a82806c4c6bc3567f2708
title: Étape 7
challengeType: 20
dashedName: step-7
---

# --description--

Déclare une nouvelle variable nommée `text` et assigne-lui la chaîne `hello world`. Ce sera le message à chiffrer.

# --hints--

Tu devrais déclarer une variable nommée `text`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("text")`)) })
```

Tu devrais assigner la chaîne `hello world` à `text`. N'oublie pas d'entourer la chaîne de quotes simples ou doubles et fais attention à la casse.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("text").is_equivalent("text = 'hello world'")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--

--fcc-editable-region--
```
