---
id: 680a3d4c065557260de2ff33
title: Étape 1
challengeType: 20
dashedName: step-1
---

# --description--

Dans ce workshop, tu vas créer un programme de chiffrement de César, une méthode de chiffrement qui décale les lettres de l'alphabet pour encoder des messages.

Commence par créer une variable appelée `shift` et assigne la valeur `5` à ta nouvelle variable.

# --hints--

Tu devrais déclarer une variable appelée `shift`. Place bien le nom de la variable au début de la ligne.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shift")`)) })
```

Tu devrais assigner la valeur `5` à ta variable `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shift").is_equivalent("shift = 5")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
