---
id: 6810869355fa2468f9ebaf9c
title: Étape 9
challengeType: 20
dashedName: step-9
---

# --description--

Maintenant, affiche `encrypted_text` pour voir le résultat dans le terminal.

# --hints--

Tu devrais afficher `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(encrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
text = 'hello world'
encrypted_text = text.translate(translation_table)
--fcc-editable-region--

--fcc-editable-region--
```
