---
id: 6810891bd6de8a87a833df42
title: Étape 10
challengeType: 20
dashedName: step-10
---

# --description--

Comme tu peux le voir dans la sortie, le message a été chiffré. La prochaine étape sera de rendre ton code réutilisable, au cas où tu voudrais chiffrer d'autres messages.

Pour ça, tu vas créer une fonction. Pour rappel, voici comment créer une fonction nommée `spam` qui affiche `Spam!` dans le terminal :

```py
def spam():
    print('Spam!')
```

Crée une fonction nommée `caesar`. Place tout ton code existant dans le corps de la fonction. Veille à garder le même niveau d'indentation pour toutes les lignes du corps de la fonction.

# --hints--

Tu devrais avoir une fonction nommée `caesar`.

```js
({ test: () => assert(runPython(`_Node(_code).has_function("caesar")`)) })
```

Tu devrais déplacer tout le code que tu as écrit jusqu'ici dans le corps de la fonction `caesar`. Assure-toi de garder le même niveau d'indentation pour toutes les lignes du corps de la fonction...

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_body().is_equivalent("alphabet = 'abcdefghijklmnopqrstuvwxyz'\\nshift = 5\\nshifted_alphabet = alphabet[shift:] + alphabet[:shift]\\ntranslation_table = str.maketrans(alphabet, shifted_alphabet)\\ntext = 'hello world'\\nencrypted_text = text.translate(translation_table)\\nprint(encrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
text = 'hello world'
encrypted_text = text.translate(translation_table)
print(encrypted_text)
--fcc-editable-region--
```
