---
id: 6818ed1adc249bb91d373729
title: Étape 14
challengeType: 20
dashedName: step-14
---

# --description--

Dans les leçons précédentes, tu as appris que l'instruction `return` sert à renvoyer une valeur depuis une fonction, pour que tu puisses l'utiliser ailleurs dans ton code :

```py
def spam():
    return 'Spam!'
```

Supprime l'appel `print(encrypted_text)` de ta fonction. Puis, supprime la déclaration de la variable `encrypted_text` et renvoie directement `text.translate(translation_table)` depuis ta fonction.

# --hints--

Tu ne devrais pas avoir `print(encrypted_text)` dans la fonction `caesar`.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").has_call("print(encrypted_text)")`)) })
```

Tu ne devrais pas avoir de variable `encrypted_text` dans la fonction `caesar`.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").has_variable("encrypted_text")`)) })
```

Ta fonction `caesar` devrait renvoyer `text.translate(translation_table)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_return("text.translate(translation_table)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)
--fcc-editable-region--

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
