---
id: 6818e9c7c8dc5694e8f64fc3
title: Étape 12
challengeType: 20
dashedName: step-12
---

# --description--

Maintenant, teste la fonction `caesar` en l'appelant avec la chaîne `freeCodeCamp` et le nombre `3` comme arguments. Assigne l'appel de fonction à une variable nommée `encrypted_text`.

# --hints--

Tu devrais avoir une variable nommée `encrypted_text` dans la portée globale.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("encrypted_text")`)) })
```

Tu devrais assigner `caesar('freeCodeCamp', 3)` à la variable `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = caesar('freeCodeCamp', 3)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
