---
id: 6818fcb250f34e62cab32c39
title: Étape 15
challengeType: 20
dashedName: step-15
---

# --description--

Tu as peut-être remarqué que, même si le message est chiffré, les caractères en majuscules n'ont pas été touchés. Cela arrive parce que la table de traduction n'inclut pas les lettres majuscules.

Pour corriger cela, tu vas modifier les chaînes passées à ton appel `str.maketrans()` avec la méthode `upper()` :

```py
greet = "Hello"
print(greet.upper()) # HELLO
```

Mets à jour ton appel `str.maketrans()` en concaténant à chaque argument la version en majuscules de l'argument lui-même.

# --hints--

Tu devrais modifier les arguments de ton appel `str.maketrans()` en concaténant à chaque argument la version en majuscules de l'argument lui-même.

```js
({ test: () => assert(runPython(`caesar('This Is A TEST', 5) == 'Ymnx Nx F YJXY'`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--
    translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
