---
id: 6818b80519e6c6c47a8bb6e8
title: Étape 11
challengeType: 20
dashedName: step-11
---

# --description--

Comme tu l'as appris dans une leçon précédente, les fonctions peuvent avoir des paramètres, des variables que tu peux utiliser à l'intérieur de la fonction. Voici une fonction `sum` avec deux paramètres, `num1` et `num2`.

```py
def sum(num1, num2):
    print(num1 + num2)
```

Le message à chiffrer et le décalage sont encore écrits en dur dans ta fonction. Donne à ta fonction deux paramètres nommés `text` et `shift`, et supprime les déclarations des variables `text` et `shift` du corps de la fonction `caesar`.

# --hints--

Ta fonction `caesar` devrait avoir deux paramètres, `text` et `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_args("text, shift")`)) })
```

Tu ne devrais pas avoir de variable `text` déclarée dans le corps de la fonction `caesar`.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").find_body().has_variable("text")`)) })
```

Tu ne devrais pas avoir de variable `shift` déclarée dans le corps de la fonction `caesar`.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").find_body().has_variable("shift")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shift = 5
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    text = 'hello world'
--fcc-editable-region--
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)
```
