---
id: 6818eb3f98c0bca5af7c8d03
title: Étape 13
challengeType: 20
dashedName: step-13
---

# --description--

Ton code est maintenant réutilisable. Pourtant, la fonction `caesar` affiche un message dans le terminal et renvoie `None` par défaut. Pour le vérifier, affiche `encrypted_text` à la fin de ton code.

# --hints--

Tu devrais afficher `encrypted_text` en dehors du corps de la fonction.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(encrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)

encrypted_text = caesar('freeCodeCamp', 3)
--fcc-editable-region--

--fcc-editable-region--
```
