---
id: 6819c5c9faf2548e556163d5
title: Étape 18
challengeType: 20
dashedName: step-18
---

# --description--

Comme tu t'en souviens des leçons précédentes, l'opérateur `not` est un opérateur logique unaire qui nie une expression :

```py
print(not True) # False
print(not False) # True
```

Utilise l'opérateur `not` pour corriger la condition de l'instruction `if` afin que ta fonction renvoie `Shift must be an integer value.` quand `shift` n'est *pas* un entier.

# --hints--

Tu devrais utiliser l'opérateur `not` pour nier `isinstance(shift, int)` dans ton instruction `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(shift, int)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

--fcc-editable-region--
    if isinstance(shift, int):
--fcc-editable-region--
        return 'Shift must be an integer value.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
