---
id: 6819be9055828d65c813eec5
title: Étape 17
challengeType: 20
dashedName: step-17
---

# --description--

La fonction `isinstance()` renvoie `True` si son premier argument est une instance du second, et `False` sinon :

```py
print(isinstance('Hello World', str)) # True
print(isinstance(42, int)) # True
```

Remplace la condition actuelle de ton instruction `if` par un appel à `isinstance()`. Passe `shift` comme premier argument, et `int` comme second argument.

# --hints--

Ton instruction `if` devrait utiliser `isinstance(shift, int)` comme condition.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("isinstance(shift, int)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

--fcc-editable-region--
    if True:
--fcc-editable-region--
        return 'Shift must be an integer value.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
