---
id: 6819ca464db219ab86f37dc4
title: Étape 20
challengeType: 20
dashedName: step-20
---

# --description--

Comme tu l'as déjà vérifié, le décalage passé pour chiffrer le texte doit être un entier positif. Il ne peut toutefois pas dépasser `25` (le dernier indice de `alphabet`).

Ajoute une seconde condition à l'instruction `if` qui vérifie si `shift` est supérieur à `25`. Souviens-toi que l'opération logique OU en Python s'écrit avec l'opérateur `or`.

Mets aussi à jour le message renvoyé pour qu'il soit `Shift must be an integer between 1 and 25.`

# --hints--

Tu devrais utiliser l'opérateur `or` pour ajouter une seconde condition après `shift < 1` qui vérifie si `shift` est supérieur à `25`. Tu ne devrais pas modifier la condition existante.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("caesar").find_ifs()[1].find_conditions()[0]
_options = ["shift > 25", "25 < shift", "shift >= 26", "26 =< shift"]
assert any(_cond.is_equivalent(f"shift < 1 or {option}") for option in _options)
`) })
```

Ta seconde instruction `if` devrait renvoyer `Shift must be an integer between 1 and 25.`

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[1].find_bodies()[0].has_return("'Shift must be an integer between 1 and 25.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

--fcc-editable-region--
    if shift < 1:
        return 'Shift must be a positive integer.'
--fcc-editable-region--

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
