---
id: 6819cf38880d8cc9b9f6af7a
title: Étape 21
challengeType: 20
dashedName: step-21
---

# --description--

Python te permet de spécifier une valeur par défaut pour les paramètres d'une fonction, ce qui crée une fonction que tu peux appeler avec moins d'arguments, ou aucun. Voici comment créer une fonction avec un paramètre `name` qui a une valeur par défaut :

```py
def greet(name='Polly'):
    return 'Hello ' + name

print(greet()) # Hello Polly
```

Ajoute un troisième paramètre nommé `encrypt` à ta fonction et donne-lui la valeur par défaut `True`.

# --hints--

Tu devrais ajouter un troisième paramètre nommé `encrypt` avec la valeur par défaut `True` à ta fonction. Ne modifie pas les paramètres existants.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_args("text, shift, encrypt=True")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar(text, shift):
--fcc-editable-region--

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
