---
id: 6819bc0637b80256a33adccc
title: Étape 16
challengeType: 20
dashedName: step-16
---

# --description--

Maintenant que tu as implémenté les fonctionnalités de base du chiffrement, il est temps d'ajouter un peu de validation. Pour ça, tu vas avoir besoin d'une instruction `if`.

Au début du corps de ta fonction, crée une instruction `if`. Pour l'instant, utilise `True` comme condition, et dans le corps de l'instruction `if` renvoie la chaîne `Shift must be an integer value.`

# --hints--

Tu devrais avoir une instruction `if` à l'intérieur de ta fonction `caesar`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0]`)) })
```

Ton instruction `if` devrait utiliser `True` comme condition.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("True")`)) })
```

Tu devrais renvoyer la chaîne `Shift must be an integer value.` depuis ton instruction `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_bodies()[0].has_return("'Shift must be an integer value.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

--fcc-editable-region--

--fcc-editable-region--

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
