---
id: 681a44f4c3235f7d8f428545
title: Étape 25
challengeType: 20
dashedName: step-25
---

# --description--

Maintenant, tu vas tester la fonction `decrypt`. Remplace la valeur assignée à `encrypted_text` par la chaîne suivante, qui représente un message à déchiffrer : `Pbhentr vf sbhaq va hayvxryl cynprf.`

Puis, déclare une variable nommée `decrypted_text` et assigne-lui un appel à `decrypt` avec `encrypted_text` comme premier argument et un décalage de `13` comme second argument.

Enfin, affiche `decrypted_text` dans le terminal. Avec ça, le chiffrement de César est terminé.

# --hints--

Tu devrais assigner `Pbhentr vf sbhaq va hayvxryl cynprf.` à `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = 'Pbhentr vf sbhaq va hayvxryl cynprf.'")`)) })
```

Tu devrais avoir une variable nommée `decrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("decrypted_text")`)) })
```

Tu devrais assigner `decrypt(encrypted_text, 13)` à `decrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("decrypted_text").is_equivalent("decrypted_text = decrypt(encrypted_text, 13)")`)) })
```

Tu devrais afficher `decrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(decrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift

    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

def encrypt(text, shift):
    return caesar(text, shift)

def decrypt(text, shift):
    return caesar(text, shift, encrypt=False)

--fcc-editable-region--
encrypted_text = encrypt('freeCodeCamp', 3)
print(encrypted_text)
--fcc-editable-region--
```

# --solutions--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift

    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

def encrypt(text, shift):
    return caesar(text, shift)

def decrypt(text, shift):
    return caesar(text, shift, encrypt=False)

encrypted_text = 'Pbhentr vf sbhaq va hayvxryl cynprf.'
decrypted_text = decrypt(encrypted_text, 13)
print(decrypted_text)
```
