---
id: 681a2a6cf95e82ca30866364
title: Étape 23
challengeType: 20
dashedName: step-23
---

# --description--

Déclare deux fonctions nommées `encrypt` et `decrypt`, toutes les deux avec les paramètres `text` et `shift`.

Tu utiliseras `encrypt` pour le chiffrement, et `decrypt` pour le déchiffrement, en donnant à chaque action un nom parlant.

Renvoie un appel à `caesar` en lui passant `text` et `shift` depuis tes deux nouvelles fonctions, mais n'oublie pas de passer aussi `False` comme troisième argument pour la fonction `decrypt`.

# --hints--

Tu devrais avoir une fonction nommée `encrypt` avec deux paramètres, `text` et `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("encrypt").has_args("text, shift")`)) })
```

Ta fonction `encrypt` devrait renvoyer `caesar(text, shift)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("encrypt").has_return("caesar(text, shift)")`)) })
```

Tu devrais avoir une fonction nommée `decrypt` avec deux paramètres, `text` et `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("decrypt").has_args("text, shift")`)) })
```

Ta fonction `decrypt` devrait renvoyer `caesar(text, shift, False)`.

```js
({ test: () => runPython(`
_n = _Node(_code).find_function("decrypt")
assert _n.has_return("caesar(text, shift, encrypt=False)") or _n.has_return("caesar(text, shift, False)")
`) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift

    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

--fcc-editable-region--

--fcc-editable-region--

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
