---
id: 681a2e35c406ade2f0a5f9b2
title: Étape 24
challengeType: 20
dashedName: step-24
---

# --description--

Il est temps de tester la fonction `encrypt`. Avec les mêmes arguments, remplace ton appel à `caesar` par un appel à `encrypt`. Tu verras la même sortie dans le terminal.

# --hints--

Ta variable `encrypted_text` devrait avoir la valeur de `encrypt('freeCodeCamp', 3)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = encrypt('freeCodeCamp', 3)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift

    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

def encrypt(text, shift):
    return caesar(text, shift)

def decrypt(text, shift):
    return caesar(text, shift, encrypt=False)

--fcc-editable-region--
encrypted_text = caesar('freeCodeCamp', 3)
--fcc-editable-region--
print(encrypted_text)
```
