---
id: 681a1d60e037127e7c8e0230
title: Étape 22
challengeType: 20
dashedName: step-22
---

# --description--

Tu vas utiliser le paramètre ajouté à l'étape précédente pour déterminer si la fonction doit chiffrer le texte qu'on lui passe (comportement par défaut, `encrypt=True`), ou si elle doit déchiffrer un message chiffré.

Crée une instruction `if` qui vérifie si `encrypt` n'est pas truthy. Dans la nouvelle instruction `if`, définis `shift` à `- shift`. C'est nécessaire pour que le décalage se fasse dans le sens opposé au chiffrement.

# --hints--

Tu devrais définir `shift` à `- shift` quand `encrypt` n'est pas truthy.

```js
({ test: () => assert(runPython(`caesar('Ymnx Nx F Yjxy', 5, False) == 'This Is A Test'`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

--fcc-editable-region--

--fcc-editable-region--

    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
