---
id: 66ed41f912d0bb1dc62da5dd
title: Crée un bot d'anecdotes JavaScript
challengeType: 26
dashedName: lab-javascript-trivia-bot
---

# --description--

**Objectif :** Réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu devrais afficher `"Hello! I'm your coding fun fact guide!"` dans la console comme message de bienvenue pour l'utilisateur.
1. Tu devrais créer trois variables : `botName`, `botLocation` et `favoriteLanguage`, qui stockent respectivement le nom du bot, son origine et son langage de programmation préféré.
1. Tu devrais afficher `"My name is (botName) and I live on (botLocation)."` dans la console.
1. Tu devrais afficher `"My favorite programming language is (favoriteLanguage)."` dans la console.
1. Tu devrais utiliser `let` pour créer une variable `codingFact` et lui affecter une chaîne qui est une anecdote sur le langage de programmation préféré de ton bot, en incluant l'utilisation de la variable `favoriteLanguage`.
1. Tu devrais afficher `codingFact` dans la console.
1. Tu devrais réaffecter la variable `codingFact` avec une nouvelle anecdote sur le langage préféré du bot, en réutilisant la variable `favoriteLanguage`.
1. Tu devrais afficher de nouveau `codingFact` dans la console.
1. Tu devrais réaffecter encore la variable `codingFact` avec une autre nouvelle anecdote sur le langage préféré du bot, en utilisant la variable `favoriteLanguage`.
1. Tu devrais afficher `codingFact` dans la console une troisième fois.
1. Tu devrais afficher `"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."` dans la console comme message d'adieu du bot.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Tu devrais afficher `"Hello! I'm your coding fun fact guide!"` dans la console.

```js
assert.equal(getLogs()[0], "Hello! I'm your coding fun fact guide!")
```

Tu devrais déclarer une variable `botName`. Vérifie bien l'orthographe et la casse.

```js
assert.exists(botName);
```

Ta variable `botName` devrait être une chaîne.

```js
assert.isString(botName);
```

Tu devrais déclarer une variable `botLocation`. Vérifie bien l'orthographe et la casse.

```js
assert.exists(botLocation);
```

Ta variable `botLocation` devrait être une chaîne.

```js
assert.isString(botLocation);
```

Tu devrais déclarer une variable `favoriteLanguage`. Vérifie bien l'orthographe et la casse.

```js
assert.exists(favoriteLanguage);
```

Ta variable `favoriteLanguage` devrait être une chaîne.

```js
assert.isString(favoriteLanguage);
```

Tu devrais afficher dans la console `"My name is (botName) and I live on (botLocation)."` et ajouter les variables à la chaîne.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.equal(getLogs()[1], `My name is ${botName} and I live on ${botLocation}.`)
```

Tu devrais afficher dans la console `"My favorite programming language is (favoriteLanguage)."` et ajouter la variable à la chaîne.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.equal(getLogs()[2], `My favorite programming language is ${favoriteLanguage}.`)
```

Tu devrais utiliser `let` pour déclarer une nouvelle variable `codingFact`.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.match(code, /let\s+codingFact/)
```

Tu devrais donner à `codingFact` une valeur qui inclut `favoriteLanguage`.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.match(code, /let\s+codingFact/);
assert.match(code, /codingFact\s*=\s*(("|'|`)?.+?\1?\s*\+?\s*|favoriteLanguage\s*\+\s*(("|'|`)?.+?\3?))/);
assert.exists(codingFact);
assert.isNotEmpty(codingFact);
assert.include(String(codingFact), favoriteLanguage);
assert.include(getLogs()[3], favoriteLanguage);
```

Tu devrais afficher `codingFact` dans la console.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.include(getLogs()[3], favoriteLanguage);
assert.isAtLeast(loggingCodingFacts.length, 1);
```

Tu devrais affecter une nouvelle valeur à `codingFact` qui contient aussi `favoriteLanguage`, et l'afficher dans la console.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
const [first, second, third] = codeWithoutComments.match(/(let )?\s*codingFact\s*=\s*(("|')?.+?\2?\s*\+?\s*|favoriteLanguage\s*\+\s*(("|')?.+?\2?))/g);
assert.include(getLogs()[4], favoriteLanguage);
assert.notEqual(getLogs()[4], getLogs()[3]);
assert.isAtLeast(loggingCodingFacts.length, 2);
assert.exists(second);
assert.isNotEmpty(codingFact);
```

Tu devrais affecter une valeur à `codingFact` pour la troisième fois, qui contient aussi `favoriteLanguage`, et l'afficher dans la console.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.include(getLogs()[5], favoriteLanguage);
assert.notEqual(getLogs()[5], getLogs()[4]);
assert.equal(getLogs()[5], codingFact);
assert.lengthOf(loggingCodingFacts, 3);
assert.isNotEmpty(codingFact);
```

Tu devrais afficher dans la console `"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."` et ajouter les valeurs des variables.

```js
assert.equal(getLogs()[6], `It was fun sharing these facts with you. Goodbye! - ${botName} from ${botLocation}.`);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
console.log("Hello! I'm your coding fun fact guide!");

const botName = "JsBot";
const botLocation = "planet Eris";
const favoriteLanguage = "JavaScript";

console.log("My name is " + botName + " and I live on " + botLocation + ".");
console.log("My favorite programming language is " + favoriteLanguage + ".");

let codingFact = "Did you know that " + favoriteLanguage + " was created in just 10 days?";

console.log(codingFact);

codingFact = "Another fun fact: " + favoriteLanguage + "was originally called Mocha!";

console.log(codingFact);

codingFact = "Also, " + favoriteLanguage + " is the most popular programming language in the world.";

console.log(codingFact);

console.log("It was fun sharing these facts with you. Goodbye! - " + botName + " from " + botLocation + ".");

```
