---
id: 68e6698b29e562d21b1d260c
title: Étape 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Dans cet atelier, tu vas construire une liste des principaux navigateurs web. Le boilerplate HTML t'a été fourni.

Commence par ajouter à ta page un titre qui contient `List of Major Web Browsers` en utilisant un élément `h1` à l'intérieur de l'élément `body`.

# --hints--

Tu devrais avoir un élément `h1` à l'intérieur de l'élément `body`.

```js
const h1 = document.querySelector('body h1');
assert.exists(h1);
```

Ton élément `h1` devrait contenir le texte `List of Major Web Browsers`.

```js
const h1 = document.querySelector('body h1');
assert.strictEqual(h1?.textContent.trim(), 'List of Major Web Browsers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
    <head>
        <meta charset="UTF-8">
        <meta name="viewport" content="width=device-width, initial-scale=1.0">
        <title>List of Browsers and Descriptions</title>
    </head>
    <body>
    --fcc-editable-region--
        
    --fcc-editable-region--
    </body>
</html>
```
