---
id: 68ecfa0f03d88990bb96a33c
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

Comme tu te souviens d'une leçon précédente, les listes de description sont utilisées pour présenter des termes et des définitions dans un format organisé et facile à lire. 

Voici un exemple :

```html
<dl>
  <dt>HTML</dt>
  <dd>HyperText Markup Language</dd>
  <dt>CSS</dt>
  <dd>Cascading Style Sheets</dd>
</dl>
```

Sous l'élément `h1`, crée un élément `dl`. Il contiendra la liste des navigateurs.

# --hints--

Tu devrais avoir un élément `dl`.

```js
const dl = document.querySelector('dl');
assert.exists(dl);
```

Ton élément `dl` devrait être sous l'élément `h1`.

```js
assert.exists(document.querySelector('h1 + dl'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        --fcc-editable-region--
        
        --fcc-editable-region--
    </body> 
</html>
```

