---
id: 68ecfa1003d88990bb96a33d
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

Les listes de description sont constituées d'une liste de termes et de détails. L'élément `dt` représente le terme de description, alors que l'élément `dd` représente les détails de la description. 

Le premier navigateur que tu vas ajouter à ta liste de description sera Google Chrome.

À l'intérieur de ton élément `dl`, ajoute un élément `dt` contenant le texte `Google Chrome`.

Sous ton élément `dt`, ajoute un élément `dd` avec le texte suivant :

```md
This is a free web browser developed by Google and first released in 2008.
```

# --hints--

Tu devrais avoir un élément `dt` à l'intérieur de ton élément `dl`.

```js
const dlElement = document.querySelector('dl dt');
assert.exists(dlElement);
```

Ton élément `dt` devrait avoir le texte `Google Chrome`.

```js
const dt = document.querySelector('dt');
assert.strictEqual(dt?.textContent.trim(), "Google Chrome");
```

Tu devrais avoir un élément `dd` sous ton élément `dt`.

```js
const ddElement = document.querySelector('dl dd');
assert.exists(ddElement);
```

Ton élément `dd` devrait contenir le texte `This is a free web browser developed by Google and first released in 2008.`

```js
const ddElement = document.querySelector('dd');
assert.strictEqual(ddElement?.textContent.trim(), "This is a free web browser developed by Google and first released in 2008.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
        --fcc-editable-region--
            
        --fcc-editable-region--
        </dl>
    </body> 
</html>
```

