---
id: 69038dd51a41ae215fcc030a
title: Étape 5
challengeType: 0
dashedName: step-5
---

# --description--

Le troisième navigateur que tu vas ajouter à la liste sera le navigateur Safari.

Ajoute un autre élément `dt` contenant le texte `Safari`.

Sous ton élément `dt`, ajoute un élément `dd` avec le texte suivant :

```md
This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.
```

# --hints--

Tu devrais avoir trois éléments `dt` dans ton code.

```js
const dtElements = document.querySelectorAll('dt');
assert.lengthOf(dtElements, 3);
```

Ton troisième élément `dt` devrait avoir le texte `Safari`.

```js
const dt = document.querySelectorAll('dt')[2];
assert.strictEqual(dt?.textContent.trim(), "Safari");
```

Tu devrais avoir trois éléments `dd` dans ton code.

```js
const ddElements = document.querySelectorAll('dd');
assert.lengthOf(ddElements, 3);
```

Ton troisième élément `dd` devrait contenir le texte `This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.`

```js
const ddElement = document.querySelectorAll('dd')[2];
assert.strictEqual(ddElement?.textContent.trim(), "This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
            <dt>Google Chrome</dt>
            <dd>This is a free web browser developed by Google and first released in 2008.</dd>

            <dt>Firefox</dt>
            <dd>This is a free web browser developed by the Mozilla Corporation and first created in 2004.</dd>

            --fcc-editable-region--
            
            --fcc-editable-region--
        </dl>
    </body> 
</html>
```
