---
id: 66cc095dfe1682753d2ab030
title: Étape 2
challengeType: 1
dashedName: step-2
---

# --description--

L'un des concepts que tu as appris dans les leçons précédentes était comment retourner des valeurs depuis une fonction.

Voici un rappel de comment retourner une valeur depuis une fonction :

```js
function myFunction() {
  return "Hello World";
}
```

À l'intérieur de ta fonction `addTwoAndSeven`, retourne la somme de `2` et `7`.

# --hints--

Ta fonction `addTwoAndSeven` doit retourner la somme de `2` et `7`.

```js
assert.strictEqual(addTwoAndSeven(), 9);
```

Ta fonction `addTwoAndSeven` doit retourner la somme en utilisant l'opérateur `+`.

```js
assert.match(code, /return\s*(2\s*\+\s*7|7\s*\+\s*2)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
--fcc-editable-region--

--fcc-editable-region--
}
```
