---
id: 66cc21d23238dc8240a8a182
title: Étape 16
challengeType: 1
dashedName: step-16
---

# --description--

Ce serait bien que ta calculatrice puisse calculer le carré d'un nombre. Le carré d'un nombre est le nombre multiplié par lui-même.

Pour calculer le carré d'un nombre en JavaScript, tu peux utiliser la méthode <dfn>Math.pow()</dfn>. La méthode `Math.pow()` prend deux arguments : le nombre de base et l'exposant.

Tu peux aussi utiliser l'opérateur d'exponentiation (`**`) pour calculer le carré d'un nombre.

Voici un exemple :

```js
// base number is 5
// the exponent is 2
Math.pow(5, 2); // 25

// using the exponentiation operator
5 ** 2; // 25
```

Déclare une fonction appelée `calculateSquare` qui prend un paramètre `num` et retourne le carré de `num`.

# --hints--

Tu dois avoir une fonction appelée `calculateSquare`.

```js
assert.isFunction(calculateSquare);
```

Ta fonction `calculateSquare` doit prendre un paramètre `num`.

```js
const regex = __helpers.functionRegex('calculateSquare', ['num']);
assert.match(__helpers.removeJSComments(code), regex);
```

Ta fonction `calculateSquare` doit retourner le carré de `num`.

```js
assert.equal(calculateSquare(5), 25);
assert.equal(calculateSquare(10), 100);
assert.equal(calculateSquare(3), 9);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

--fcc-editable-region--

--fcc-editable-region--
```
