---
id: 66cc1a3a39aef47d6473cb2f
title: Étape 13
challengeType: 1
dashedName: step-13
---

# --description--

L'étape suivante consiste à ajouter la fonctionnalité de division à ta calculatrice.

Déclare une fonction appelée `calculateQuotient` qui prend deux paramètres, `num1` et `num2`.

À l'intérieur de la fonction, retourne le quotient des deux paramètres.

Puis, ajoute un `console.log` qui appelle la fonction `calculateQuotient` avec les arguments `7` et `11`.

# --hints--

Tu dois avoir une fonction appelée `calculateQuotient`.

```js
assert.isFunction(calculateQuotient);
```

Ta fonction `calculateQuotient` doit avoir deux paramètres appelés `num1` et `num2`.

```js
const regex = __helpers.functionRegex('calculateQuotient', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

Ta fonction `calculateQuotient` doit retourner le quotient des deux paramètres.

```js
assert.equal(calculateQuotient(10, 2), 5);
assert.equal(calculateQuotient(100, 10), 10);
```

Tu dois avoir un `console.log` qui appelle la fonction `calculateQuotient` avec les arguments `7` et `11`.

```js
assert.match(code, /console\.log\s*\(\s*calculateQuotient\s*\(\s*7\s*,\s*11\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));

--fcc-editable-region--

--fcc-editable-region--
```
