---
id: 66cc172af7159a7c67804544
title: Étape 12
challengeType: 1
dashedName: step-12
---

# --description--

Maintenant, il est temps d'ajouter la fonctionnalité de multiplication à ta calculatrice.

Déclare une fonction appelée `calculateProduct` qui prend deux paramètres, `num1` et `num2`.

À l'intérieur de la fonction, retourne le produit des deux paramètres.

Puis, ajoute un `console.log` qui appelle la fonction `calculateProduct` avec les arguments `13` et `5`.

# --hints--

Tu dois avoir une fonction appelée `calculateProduct`.

```js
assert.isFunction(calculateProduct);
```

Ta fonction `calculateProduct` doit avoir deux paramètres appelés `num1` et `num2`.

```js
const regex = __helpers.functionRegex('calculateProduct', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

Ta fonction `calculateProduct` doit retourner le produit des deux paramètres.

```js
assert.equal(calculateProduct(5, 2), 10);
assert.equal(calculateProduct(10, 5), 50);
```

Tu dois avoir un `console.log` qui appelle la fonction `calculateProduct` avec les arguments `13` et `5`.

```js
assert.match(code, /console\.log\s*\(\s*calculateProduct\s*\(\s*13\s*,\s*5\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

--fcc-editable-region--

--fcc-editable-region--
```
