---
id: 66cc0a9e06e00b75d6782be9
title: Étape 3
challengeType: 1
dashedName: step-3
---

# --description--

Dans les leçons précédentes, tu as appris à appeler (invoquer) une fonction. Appeler une fonction signifie exécuter le code à l'intérieur de la fonction.

Voici un rappel de comment appeler une fonction :

```js
function myFunction() {
  return "Hello World";
}

// function call
myFunction();
```

Ajoute une instruction `console.log`, et à l'intérieur, appelle la fonction `addTwoAndSeven`.

Tu devrais maintenant voir la somme de `2` et `7` affichée dans la console.

# --hints--

Tu dois avoir un `console.log` dans ton code.

```js
assert.match(code, /console\.log(.*)/);
```

Tu dois appeler la fonction `addTwoAndSeven` à l'intérieur du `console.log`.

```js
assert.match(code, /console\.log\(\s*addTwoAndSeven\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}
--fcc-editable-region--

--fcc-editable-region--
```
