---
id: 66cc0f1ae40802781b2ea972
title: Étape 6
challengeType: 1
dashedName: step-6
---

# --description--

Dans les leçons précédentes, tu as appris à travailler avec les paramètres de fonction.

Un paramètre de fonction est une variable définie dans la déclaration de la fonction et qui sert d'espace réservé.

Voici un exemple de fonction qui possède un paramètre :

```js
// The parameter is `name`
function greetUser(name) {
  return `Hello, ${name}!`;
}
```

Déclare une fonction appelée `calculateSum` qui prend deux paramètres, `num1` et `num2`.

# --hints--

Tu dois déclarer une fonction appelée `calculateSum`.

```js
assert.isFunction(calculateSum);
```

Ta fonction `calculateSum` doit prendre deux paramètres, `num1` et `num2`.

```js
const regex = __helpers.functionRegex('calculateSum', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
