---
id: 66cc10958d72e57907f898b2
title: Étape 7
challengeType: 1
dashedName: step-7
---

# --description--

À l'intérieur de ta fonction `calculateSum`, tu devras retourner la somme des deux paramètres.

# --hints--

Ta fonction `calculateSum` doit retourner la somme des deux paramètres.

```js
assert.equal(calculateSum(1, 2), 3);
assert.equal(calculateSum(5, 5), 10);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
--fcc-editable-region--

--fcc-editable-region--
}
```
