---
id: 66cc0de2920bec775f610424
title: Étape 5
challengeType: 1
dashedName: step-5
---

# --description--

Même si les fonctions fonctionnent comme prévu, il y a beaucoup de répétition dans ton code.

Puisque tu construis une calculatrice, tu ne veux pas avoir à créer une fonction pour chaque combinaison possible de nombres que tu veux additionner.

Il serait préférable de créer une seule fonction réutilisable capable d'additionner deux nombres quelconques.

Supprime tout le code que tu as écrit jusqu'ici.

À l'étape suivante, tu réviseras comment travailler avec les paramètres et les arguments dans les fonctions.

# --hints--

Tu ne dois pas avoir de fonction appelée `addTwoAndSeven`.

```js
assert.notMatch(code, /function\s+addTwoAndSeven\s*\(\s*\)\s*{/);
```

Tu ne dois pas avoir d'appel à la fonction `addTwoAndSeven`.

```js
assert.notMatch(code, /console\.log\(\s*addTwoAndSeven\(\s*\)\s*\)/);
```

Tu ne dois pas avoir de fonction appelée `addThreeAndFour`.

```js
assert.notMatch(code, /function\s+addThreeAndFour\s*\(\s*\)\s*{/);
```

Tu ne dois pas avoir d'appel à la fonction `addThreeAndFour`.

```js
assert.notMatch(code, /console\.log\(\s*addThreeAndFour\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

function addThreeAndFour() {
  return 3 + 4;
}

console.log(addThreeAndFour());
--fcc-editable-region--
```
