---
id: 66cc12fa504b0479dac479a0
title: Étape 8
challengeType: 1
dashedName: step-8
---

# --description--

L'avantage des fonctions est qu'elles peuvent être appelées avec différents arguments, ce qui te permet de réutiliser le même code avec diverses valeurs.

Voici quelques exemples d'appels de fonction avec différents arguments de type chaîne :

```js
// function definition
function greetUser(name) {
  console.log(`Hello ${name}!`);
}

// function calls
greetUser("John"); // "Hello John!"
greetUser("Jane"); // "Hello Jane!"
```

Ajoute un `console.log` qui appelle la fonction `calculateSum` avec les arguments `2` et `5`.

# --hints--

Tu dois avoir un `console.log` qui appelle la fonction `calculateSum` avec les arguments `2` et `5`.

```js
assert.match(code, /console\.log\(calculateSum\(\s*2\s*,\s*5\)\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
    return num1 + num2;
}

--fcc-editable-region--

--fcc-editable-region--
```
