---
id: 66cc06012e46aa6bc9b8c001
title: Étape 1
challengeType: 1
dashedName: step-1
---

# --description--

Dans ce workshop, tu vas réviser comment travailler avec les fonctions en construisant une application de calculatrice.

Dans les leçons précédentes, tu as appris à déclarer des fonctions comme ceci :

```js
// regular function
function myFunction() {

}

// arrow function
const myArrowFunction = () => {

}
```

Commence par créer une fonction appelée `addTwoAndSeven`.

Tu peux choisir d'utiliser la syntaxe de fonction classique ou la syntaxe de fonction fléchée.

# --hints--

Tu dois déclarer une fonction appelée `addTwoAndSeven`.

```js
assert.isFunction(addTwoAndSeven);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
