---
id: 66cc16abc30b367ba1d1f9ef
title: Étape 11
challengeType: 1
dashedName: step-11
---

# --description--

Maintenant, il est temps de tester ta fonction `calculateDifference`.

Commence par ajouter un `console.log` qui appelle la fonction `calculateDifference` avec les arguments `22` et `5`.

Puis, ajoute un deuxième `console.log` qui appelle la fonction `calculateDifference` avec les arguments `12` et `1`.

Enfin, ajoute un troisième `console.log` qui appelle la fonction `calculateDifference` avec les arguments `17` et `9`.

# --hints--

Tu dois avoir un `console.log` qui appelle la fonction `calculateDifference` avec les arguments `22` et `5`.

```js
assert.match(code, /console\.log\s*\(\s*calculateDifference\s*\(\s*22\s*,\s*5\s*\)\s*\)\s*;?/);
```

Tu dois avoir un `console.log` qui appelle la fonction `calculateDifference` avec les arguments `12` et `1`.

```js
assert.match(code, /console\.log\s*\(\s*calculateDifference\s*\(\s*12\s*,\s*1\s*\)\s*\)\s*;?/);
```

Tu dois avoir un `console.log` qui appelle la fonction `calculateDifference` avec les arguments `17` et `9`.

```js
assert.match(code, /console\.log\s*\(\s*calculateDifference\s*\(\s*17\s*,\s*9\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

--fcc-editable-region--

--fcc-editable-region--
```
