---
id: 66cc0ba5881acb7692cfc4de
title: Étape 4
challengeType: 1
dashedName: step-4
---

# --description--

Maintenant, il est temps d'ajouter une autre fonction.

Déclare une fonction appelée `addThreeAndFour` qui retourne la somme de `3` et `4`.

Puis appelle la fonction `addThreeAndFour` à l'intérieur d'un `console.log` pour voir le résultat.

# --hints--

Tu dois déclarer une fonction appelée `addThreeAndFour`.

```js
assert.isFunction(addThreeAndFour);
```

Ta fonction `addThreeAndFour` doit retourner la somme de `3` et `4`.

```js
assert.equal(addThreeAndFour(), 7);
```

Ta fonction `addThreeAndFour` doit retourner la somme en utilisant l'opérateur `+`.

```js
const declarationStyle = /return\s*\(?\s*(3\s*\+\s*4|4\s*\+\s*3)\s*\)?\s*;?/;
const arrowStyle = /=>\s*\(?\s*(3\s*\+\s*4|4\s*\+\s*3)\s*\)?\s*;?/;
const combinedRegex = new RegExp(`(?:${declarationStyle.source})|(?:${arrowStyle.source})`);
assert.match(code, combinedRegex);
```

Tu dois appeler la fonction `addThreeAndFour` à l'intérieur d'un `console.log`.

```js
assert.match(code, /console\.log\(\s*addThreeAndFour\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

--fcc-editable-region--

--fcc-editable-region--
```
