---
id: 66cc14ca14e65e7a9ebecb80
title: Étape 10
challengeType: 1
dashedName: step-10
---

# --description--

Maintenant que ta calculatrice peut additionner deux nombres, il est temps de créer une fonction qui soustraira deux nombres.

Déclare une fonction appelée `calculateDifference` qui prend deux paramètres, `num1` et `num2`.

À l'intérieur de la fonction, retourne la différence entre les deux paramètres.

# --hints--

Tu dois avoir une fonction appelée `calculateDifference`.

```js
assert.isFunction(calculateDifference);
```

Ta fonction `calculateDifference` doit avoir deux paramètres appelés `num1` et `num2`.

```js
const regex = __helpers.functionRegex('calculateDifference', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

Ta fonction `calculateDifference` doit retourner la différence entre les deux paramètres.

```js
assert.equal(calculateDifference(5, 2), 3);
assert.equal(calculateDifference(10, 5), 5);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

--fcc-editable-region--

--fcc-editable-region--
```
