---
id: 66cc2739f687e484d50bb6f1
title: Étape 18
challengeType: 1
dashedName: step-18
---

# --description--

La dernière fonctionnalité de ta calculatrice sera la fonctionnalité de racine carrée.

Un carré est un nombre multiplié par lui-même. Ainsi, la racine carrée serait le nombre qui, multiplié par lui-même, donne le nombre d'origine.

Par exemple, la racine carrée de `9` est `3` car `3 * 3 = 9`.

Pour obtenir la racine carrée d'un nombre en JavaScript, tu peux utiliser la méthode <dfn>Math.sqrt()</dfn>.

Voici un exemple :

```js
// result: 3
Math.sqrt(9);
```

Déclare une fonction appelée `calculateSquareRoot` qui prend un paramètre `num` et retourne la racine carrée de `num`.

# --hints--

Tu dois avoir une fonction appelée `calculateSquareRoot`.

```js
assert.isFunction(calculateSquareRoot);
```

Ta fonction `calculateSquareRoot` doit prendre un paramètre `num`.

```js
const regex = __helpers.functionRegex('calculateSquareRoot', ['num']);
assert.match(__helpers.removeJSComments(code), regex);
```

Ta fonction `calculateSquareRoot` doit retourner la racine carrée de `num`.

```js
assert.equal(calculateSquareRoot(9), 3);
assert.equal(calculateSquareRoot(16), 4);
assert.equal(calculateSquareRoot(25), 5);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

function calculateSquare(num) {
  return num ** 2;
}

console.log(calculateSquare(2));
console.log(calculateSquare(9));

--fcc-editable-region--

--fcc-editable-region--
```
