---
id: 66cdf76685e4cb5a8726e27b
title: Étape 9
challengeType: 1
dashedName: step-9
---

# --description--

Pour tester davantage ta fonction `calculateSum`, tu devras l'appeler avec différents arguments.

Ajoute un autre `console.log` qui appelle la fonction `calculateSum` avec les arguments `10` et `10`.

En dessous de ce `console.log`, ajoute un autre `console.log` qui appelle la fonction `calculateSum` avec les arguments `5` et `5`.

# --hints--

Tu dois avoir un `console.log` qui appelle la fonction `calculateSum` avec les arguments `10` et `10`.

```js
assert.match(code, /console\.log\(calculateSum\(10\s*,\s*10\)\)/);
```

Tu dois avoir un `console.log` qui appelle la fonction `calculateSum` avec les arguments `5` et `5`.

```js
assert.match(code, /console\.log\(calculateSum\(5\s*,\s*5\)\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
    return num1 + num2;
}

console.log(calculateSum(2, 5));

--fcc-editable-region--

--fcc-editable-region--
```
