---
id: 66ebd4ae2812430bb883c787
title: Construire un hub d'événements
challengeType: 25
dashedName: lab-event-hub
demoType: onClick
---

# --description--

Dans ce labo, tu vas utiliser les éléments HTML sémantiques pour créer la structure d'une page web. Tu ajouteras du contenu et des images pour qu'elle ressemble à un vrai hub d'événements.

**Objectif :** réponds aux user stories ci-dessous et fais passer tous les tests pour terminer le labo.

**User Stories :**

1. Tu dois avoir un élément `header`.
1. Dans l'élément `header`, tu dois avoir un élément `h1` qui contient le texte `Event Hub`, et un élément `nav`.
1. Dans l'élément `nav`, tu dois avoir une liste non ordonnée de deux éléments contenant des liens vers différentes sections de la page. Le premier élément doit avoir le texte `Upcoming Events`, et le deuxième doit avoir le texte `Past Events`.
1. Chaque lien doit être représenté par un élément `a` avec un attribut `href` qui pointe vers la section correspondante de la page, respectivement `#upcoming-events` et `#past-events`.
1. Tu dois avoir un élément `main` qui contient les différentes sections de la page.
1. Dans l'élément `main`, tu dois avoir deux éléments `section`.
1. Le premier élément `section` doit avoir un attribut `id` avec la valeur `upcoming-events`.
1. Dans la section `#upcoming-events`, tu dois avoir :

   - Un élément `h2` avec le texte `Upcoming Events`.
   - Deux éléments `article`. Chaque article doit représenter un événement, et il doit avoir :
      - Un élément `h3` pour le titre de l'événement.
      - Un élément `p` pour la description de l'événement. Tu peux ajouter une date en bas si tu veux.

1. Le deuxième élément `section` doit avoir un attribut `id` avec la valeur `past-events`.
1. Dans la section `#past-events`, tu dois avoir :

   - Un élément `h2` avec le texte `Past Events`.
   - Deux éléments `article`. Chaque élément article doit représenter un événement passé, et il doit avoir :
     - Un élément `h3` pour le titre de l'événement.
     - Un élément `p` pour la description de l'événement. Tu peux ajouter une date en bas si tu veux.
     - Un élément image avec l'attribut `src` qui pointe vers un fichier image et l'attribut `alt` avec une description de l'image.

**Note :** tu peux utiliser le texte que tu veux pour les descriptions et les dates des événements. Tu peux utiliser les URLs d'images suivantes pour les images si tu veux :

- `https://cdn.freecodecamp.org/curriculum/labs/past-event1.jpg`.
- `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg`.

# --hints--

Tu dois avoir un élément `header`.

```js
assert.isNotNull(document.querySelector("header"));
```

Ton élément `header` doit venir après la balise ouvrante `body`.

```js
assert.equal(document.querySelector("body")?.firstElementChild?.tagName, "HEADER");
```

Dans l'élément `header`, tu dois avoir un élément `h1` qui contient le texte `Event Hub`.

```js
const h1Element = document.querySelector('header h1');
assert.strictEqual(h1Element?.innerText.trim(), "Event Hub");
```

Dans l'élément `header`, après l'élément `h1`, tu dois avoir un élément `nav`.

```js
assert.isNotNull(document.querySelector("header>h1+nav"));
```

Ton élément `nav` doit contenir une liste non ordonnée de deux éléments.

```js
const liElements = document.querySelectorAll('header nav>ul>li');

assert.isNotNull('header nav>ul');
assert.strictEqual(liElements.length, 2);
```

Le premier élément de la liste non ordonnée doit être un lien.

```js
const firstLink = document.querySelectorAll('header nav ul li a')[0];
assert.exists(firstLink);
```

Le deuxième élément de la liste non ordonnée doit être un lien.

```js
const secondLink = document.querySelectorAll('header nav ul li a')[1];
assert.exists(secondLink);
```

Le texte du premier élément de la liste non ordonnée doit être `Upcoming Events`.

```js
const firstLink = document.querySelectorAll('header nav>ul>li>a')[0];
assert.strictEqual(firstLink.innerText.trim(), "Upcoming Events");
```

Le premier élément de la liste non ordonnée doit avoir l'attribut `href` défini sur `#upcoming-events`.

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[0];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#upcoming-events");
```

Le texte du deuxième élément de la liste non ordonnée doit être `Past Events`.

```js
const secondLink = document.querySelectorAll('header nav>ul>li>a')[1];
assert.strictEqual(secondLink.innerText.trim(), "Past Events");
```

Le deuxième élément de la liste non ordonnée doit avoir l'attribut `href` défini sur `#past-events`.

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[1];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#past-events");
```

Tu dois avoir un élément `main` après la balise fermante de l'élément `header`.

```js
const mainElement = document.querySelector("body>header+main");
assert.isNotNull(mainElement);
```

Dans l'élément `main`, tu dois avoir deux éléments `section`.

```js
const sectionElements = document.querySelectorAll('body>header+main>section');
assert.strictEqual(sectionElements.length, 2);
```

Ton premier élément `section` doit avoir un attribut `id` avec la valeur `upcoming-events`.

```js
const firstSection = document.querySelectorAll('body>header+main>section')[0];
const idAttribute = firstSection?.getAttribute("id");
assert.strictEqual(idAttribute, "upcoming-events");
```

Ton deuxième élément `section` doit avoir un attribut `id` avec la valeur `past-events`.

```js
const secondSection = document.querySelectorAll('body>header+main>section')[1];
const idAttribute = secondSection?.getAttribute("id");
assert.strictEqual(idAttribute, "past-events");
```

Dans la section `#upcoming-events`, tu dois avoir un élément `h2` avec le texte `Upcoming Events`.

```js
const h2Element = document.querySelector('#upcoming-events h2');
assert.strictEqual(h2Element?.innerText.trim(), "Upcoming Events");
```

Dans la section `#upcoming-events`, tu dois avoir deux éléments `article` sous l'élément `h2`.

```js
const articleElements = document.querySelectorAll('#upcoming-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

Les deux éléments `article` dans la section `#upcoming-events` doivent avoir un élément `h3` pour le titre de l'événement.

```js
const h3Elements = document.querySelectorAll('#upcoming-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

Les deux éléments `article` dans la section `#upcoming-events` doivent avoir un élément paragraphe pour la description de l'événement.

```js
const articles = document.querySelectorAll('#upcoming-events article');
assert.isNotEmpty(articles);
articles.forEach(article => {
    assert.isAtLeast(article.querySelectorAll('h3 ~ p').length, 1);
});
```

Dans la section `#past-events`, tu dois avoir un élément `h2` avec le texte `Past Events`.

```js
const h2Element = document.querySelector('#past-events h2');
assert.strictEqual(h2Element?.innerText.trim(), "Past Events");
```

Dans la section `#past-events`, tu dois avoir deux éléments `article` sous l'élément `h2`.

```js
const articleElements = document.querySelectorAll('#past-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

Les deux éléments `article` dans la section `#past-events` doivent avoir un élément `h3` pour le titre de l'événement.

```js
const h3Elements = document.querySelectorAll('#past-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

Les deux éléments `article` dans la section `#past-events` doivent avoir un élément paragraphe pour la description de l'événement.

```js
const articles = document.querySelectorAll('#past-events article');
assert.isNotEmpty(articles);
articles.forEach(article => {
    assert.isAtLeast(article.querySelectorAll('h3 ~ p').length, 1);
});
```

Les deux éléments `article` dans la section `#past-events` doivent avoir un élément image.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);
```

Les deux éléments image dans la section `#past-events` doivent avoir l'attribut `src` qui pointe vers un fichier image.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
    assert.exists(img.getAttribute("src"));
}
```

Les deux éléments image dans la section `#past-events` doivent avoir l'attribut `alt` avec une description de l'image.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
    assert.exists(img.getAttribute("alt"));
}
```

Chaque élément `h3` doit avoir le titre de l'événement.

```js
const eventTitles = document.querySelectorAll('h3');
assert.isNotEmpty(eventTitles);
eventTitles.forEach((eventTitle => assert.isNotEmpty(eventTitle.innerText.trim())));
```

Chaque élément `p` doit avoir la description de l'événement.

```js
const eventDescriptions = document.querySelectorAll('p');
assert.isNotEmpty(eventDescriptions);
eventDescriptions.forEach((eventDescription => assert.isNotEmpty(eventDescription.innerText.trim())));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Event Hub</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Event Hub</title>
</head>
<body>
    <header>
        <h1>Event Hub</h1>
        <nav>
            <ul>
                <li><a href="#upcoming-events">Upcoming Events</a></li>
                <li><a href="#past-events">Past Events</a></li>
            </ul>
        </nav>
    </header>
    <main>
        <section id="upcoming-events">
            <h2>Upcoming Events</h2>
            <article>
                <h3>Conférence IA et apprentissage automatique 2024</h3>
                <p>Rejoins-nous pour explorer les dernières avancées en intelligence artificielle et en apprentissage automatique. Des leaders du secteur partageront des idées et des cas concrets sur la façon dont l'IA transforme différents domaines.</p>
                <p>Date : 10 août 2024</p>
            </article>
            <article>
                <h3>Bootcamp de développement web</h3>
                <p>Un atelier pratique conçu pour les développeurs qui veulent améliorer leurs compétences avec les technologies web modernes, notamment React, Node.js et GraphQL. Parfait pour les débutants comme pour les développeurs expérimentés.</p>
                <p>Date : 5 septembre 2024</p>
            </article>
        </section>
        <section id="past-events">
            <h2>Past Events</h2>
            <article>
                <h3>Sommet cybersécurité 2024</h3>
                <p>Un événement consacré aux dernières tendances et menaces en cybersécurité. Des experts ont présenté des stratégies pour protéger les données et garantir la confidentialité dans un monde de plus en plus numérique.</p>
                <p>Date : 15 juin 2024</p>
                <img src="https://cdn.freecodecamp.org/curriculum/labs/past-event1.jpg" alt="Image du Sommet cybersécurité 2024">
            </article>
            <article>
                <h3>Expo blockchain 2024</h3>
                <p>Un événement complet sur l'avenir de la technologie blockchain. Les sujets abordés incluaient la finance décentralisée (DeFi), les contrats intelligents et l'impact de la blockchain sur différents secteurs.</p>
                <p>Date : 20 juillet 2024</p>
                <img src="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg" alt="Image de l'Expo blockchain 2024">
            </article>
        </section>
    </main>
</body>
</html>
```
