---
id: 686daa7ed79ceacd0b264e80
title: Étape 5
challengeType: 0
dashedName: step-5
---

# --description--

Tu te rapproches du résultat. Regarde maintenant cet exemple :

```html
<svg viewBox="0 0 50 50">
</svg>
```

L'attribut `viewBox` contrôle la partie de l'image visible dans le SVG.

- Les deux premiers nombres (`0 0`) définissent la position de départ du `viewBox`, c'est-à-dire le coin supérieur gauche (x et y).
- Les deux nombres suivants (`50 50`) définissent la largeur et la hauteur du `viewBox`.

Définis l'attribut `viewBox` sur `0 0 24 24`.

# --hints--

Tu dois avoir un attribut `viewBox`.

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
assert.exists(viewBox);
```

Tu dois définir la position x du `viewBox` sur 0.

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
const x = viewBox?.trim().split(' ')[0];
assert.strictEqual(x, '0');
```

Tu dois définir la position y du `viewBox` sur 0.

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
const y = viewBox?.trim().split(' ')[1];
assert.strictEqual(y, '0');
```

Tu dois définir la largeur du `viewBox` sur 24.

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
const width = viewBox?.trim().split(' ')[2];
assert.strictEqual(width, '24');
```

Tu dois définir la hauteur du `viewBox` sur 24.

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
const height = viewBox?.trim().split(' ')[3];
assert.strictEqual(height, '24');
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Icône de coeur</title>
  </head>
  <body>
    --fcc-editable-region--
    <svg width="24" height="24">
    --fcc-editable-region--
      <path
        d="M12 21s-6-4.35-9.33-8.22C-.5 7.39 3.24 1 8.4 4.28 10.08 5.32 12 7.5 12 7.5s1.92-2.18 3.6-3.22C20.76 1 24.5 7.39 21.33 12.78 18 16.65 12 21 12 21z"
      ></path>
    </svg>
  </body>
</html>
```
