---
id: 686d749d6f32f708fdf2b7ad
title: Étape 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Dans une leçon précédente, tu as découvert les éléments `svg` et la façon dont ils sont souvent utilisés comme icônes dans les projets. Dans de vraies bases de code, tu utiliserais généralement des bibliothèques d'icônes, donc tu n'aurais pas besoin de créer des éléments `svg` à partir de zéro. Cependant, dans cet atelier, tu vas construire une icône de coeur pour découvrir les attributs principaux utilisés dans un élément `svg`.

Commence par créer un élément `svg` sur la page.

# --hints--

Tu dois avoir un élément `svg` sur la page.

```js
const svg = document.querySelector('svg');
assert.exists(svg);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Icône de coeur</title>
  </head>
  <body>
    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>
```
