---
id: 6995d0325fd7da40854fac40
title: Étape 1
challengeType: 1
dashedName: step-1
---

# --description--

Dans une leçon précédente, tu as appris la recherche binaire et en quoi elle diffère de la recherche linéaire.

Dans cet atelier, tu vas mettre cela en pratique en construisant un algorithme de recherche binaire.

Pour commencer, crée une fonction `binarySearch` avec les paramètres `searchList` et `value`.

`searchList` est le tableau d'éléments dans lequel la fonction cherchera, et `value` est la cible que la fonction tentera de trouver dans `searchList`.

# --hints--

Tu dois créer une fonction `binarySearch`.

```js
assert.isFunction(binarySearch);
```

Ta fonction `binarySearch` doit avoir les paramètres `searchList` et `value`.

```js
const binSearchParams = __helpers.getFunctionParams(binarySearch.toString());

const binarySearchParamsNames = binSearchParams.map(param => param.name)
assert.include(binarySearchParamsNames, "searchList");
assert.include(binarySearchParamsNames, "value");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
