---
id: 6996d0a2c2ac55041145a13f
title: Étape 14
challengeType: 1
dashedName: step-14
---

# --description--

As-tu remarqué que le deuxième appel de fonction a maintenant les valeurs `3, 5, 4` dans la liste ? Cela indique que la recherche binaire fonctionne comme prévu.

Voici comment cela s'est passé : l'algorithme a d'abord vérifié `3` comme milieu de la liste initiale. Comme `4` est supérieur à `3`, la recherche s'est déplacée vers la moitié droite. Elle a ensuite examiné `5` comme nouveau milieu. Comme `4` est inférieur à `5`, la recherche s'est déplacée vers la gauche, identifiant finalement `4` comme le milieu de la plage finale.

Pour tester à nouveau la fonction, appelle-la avec `[1, 3, 5, 9, 14, 22], 10` et affiche-la dans la console. C'est une situation où la valeur ne sera pas trouvée.

# --hints--

Tu dois appeler la fonction avec `binarySearch([1, 3, 5, 9, 14, 22], 10)` et l'afficher dans la console.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*binarySearch\s*\(\s*\[\s*1\s*,\s*3\s*,\s*5\s*,\s*9\s*,\s*14\s*,\s*22\s*\]\s*,\s*10\s*\)\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);

    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
--fcc-editable-region--

--fcc-editable-region--
```
