---
id: 6996c7363a4a437839f1f15d
title: Étape 4
challengeType: 1
dashedName: step-4
---

# --description--

Maintenant que tu as défini les bornes de ta recherche, il est temps de créer la boucle qui effectuera l'algorithme de recherche binaire.

La recherche binaire fonctionne en réduisant de façon répétée l'espace de recherche. Ce processus continue tant qu'il y a une plage valide d'éléments à vérifier.

Tu peux exprimer cela avec une boucle `while` qui continuera tant que ton pointeur `low` est inférieur ou égal à ton pointeur `high`.

Si `low` devient jamais supérieur à `high`, cela signifie que l'espace de recherche est devenu vide, et que `value` n'est pas dans le tableau.

À l'intérieur de la fonction `binarySearch`, crée une boucle `while` avec une condition qui vérifie si `low` est inférieur ou égal à high.

# --hints--

Tu dois créer une boucle `while` qui s'exécute tant que `low` est inférieur ou égal à `high`.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /while\s*\(\s*low\s*<=\s*high\s*\)/);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
--fcc-editable-region--

--fcc-editable-region--
}
```
