---
id: 6996cfb6c01e9a802add38be
title: Étape 13
challengeType: 1
dashedName: step-13
---

# --description--

Enfin, si `value` n'est pas trouvée au milieu et n'est pas non plus supérieure à `valueAtMiddle`, alors `value` doit être inférieure à `valueAtMiddle`, et doit être à gauche. Le bloc `else` final gérera cela.

Donc, termine la boucle en retirant le mot-clé `break` et en ajoutant un bloc `else`. À l'intérieur, mets à jour la variable `high` en soustrayant `1` de `mid`.

# --hints--

Tu dois remplacer le mot-clé `break` par un bloc `else` qui met à jour la variable `high` à `mid - 1`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /else\s*\{\s*high\s*=\s*mid\s*-\s*1\s*;?\s*\}/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);

    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    }

--fcc-editable-region--
    break;
--fcc-editable-region--

  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
```
