---
id: 6996c60281feca6718cc15c1
title: Étape 2
challengeType: 1
dashedName: step-2
---

# --description--

À l'intérieur de ta fonction `binarySearch`, crée une variable `pathToTarget` et initialise-la à un tableau vide, pour commencer.

# --hints--

Tu dois avoir une variable `pathToTarget` définie à un tableau vide `[]` dans ta fonction `binarySearch`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+pathToTarget\s*=\s*\[\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
--fcc-editable-region--

--fcc-editable-region--
}
```
