---
id: 6996d1032feed88e2957dc8e
title: Étape 15
challengeType: 1
dashedName: step-15
---

# --description--

Tu peux voir que le dernier appel affiche un tableau vide. Pour préciser davantage que la valeur n'est pas dans le tableau recherché, mets à jour l'instruction `return` pour renvoyer un tableau contenant deux éléments, un tableau vide, et la chaîne `Value not found`.

# --hints--

Tu dois mettre à jour l'instruction `return` pour renvoyer `[], 'Value not found'`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /return\s*\[\s*\[\s*\]\s*,\s*('|"|`)Value not found\1\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);

    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
--fcc-editable-region--
    return [];
--fcc-editable-region--
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
console.log(binarySearch([1, 3, 5, 9, 14, 22], 10));
```
