---
id: 6996c8f94fe6b250f11a2a1e
title: Étape 7
challengeType: 1
dashedName: step-7
---

# --description--

Maintenant, inclus `valueAtMiddle` dans le tableau `pathToTarget` pour suivre les étapes prises pendant la recherche, qu'il s'agisse ou non de la cible. Tu peux le faire avec la méthode `push()`.

# --hints--

Tu dois utiliser la méthode `push()` pour ajouter `valueAtMiddle` au tableau `pathToTarget`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /pathToTarget\s*\.\s*push\(\s*valueAtMiddle\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
