---
id: 6996ca21a74104721f2c7e5d
title: Étape 9
challengeType: 1
dashedName: step-9
---

# --description--

Tu dois tester les choses pour comprendre le flux de l'algorithme à ce stade initial.

Pour cela, tu dois d'abord sortir de la boucle. C'est parce que l'implémentation actuelle ne permettra qu'une seule itération, donc si la condition n'est pas remplie, il y aura une boucle infinie.

Juste après l'instruction `if`, utilise le mot-clé `break` pour sortir de la boucle `while`. Ensuite, après la boucle `while`, renvoie un tableau vide pour indiquer que la valeur n'a pas été trouvée.

# --hints--

Tu dois utiliser le mot-clé `break` pour sortir de la boucle `while` après l'instruction `if`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?\}\s*break\s*;?/
);
```

Tu dois renvoyer un tableau vide après ta boucle `while`.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /return\s*\[\s*\]/);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    if (value === valueAtMiddle) {
      return pathToTarget;
    }
--fcc-editable-region--

  }

--fcc-editable-region--
}
```
