---
id: 6996c675b599baa7b103192d
title: Étape 3
challengeType: 1
dashedName: step-3
---

# --description--

Dans la recherche binaire, il y a l'indice le plus bas possible et l'indice le plus haut possible.

L'indice le plus bas possible représente la borne gauche de la portion actuelle des valeurs recherchées, ici `searchList`. L'indice le plus haut possible, quant à lui, représente la borne droite des valeurs recherchées.

Pour tenir compte de ces deux, tu dois considérer tout le tableau recherché. Donc, à l'intérieur de la fonction `binarySearch`, définis une variable `low` avec une valeur de `0`, et une variable `high` avec une valeur qui correspond au dernier indice du tableau recherché.

# --hints--

Tu dois avoir une variable nommée `low` définie à `0` dans ta fonction `binarySearch`.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /(const|let|var)\s+low\s*=\s*0/);
```

Tu dois avoir une variable nommée `high` définie à `searchList.length - 1` dans ta fonction `binarySearch`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+high\s*=\s*searchList\.length\s*-\s*1/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
--fcc-editable-region--

--fcc-editable-region--
}
```
