---
id: 6996c8375f4301004459ee85
title: Étape 6
challengeType: 1
dashedName: step-6
---

# --description--

Maintenant que tu as calculé l'indice `mid`, tu dois réellement récupérer la valeur de `searchList` à cet indice. Cette valeur est celle que tu compareras à ton paramètre `value`, qui est la cible que tu recherches.

Crée une variable `valueAtMiddle` et assigne-lui l'élément de `searchList` situé à l'indice `mid`.

# --hints--

Tu dois créer une variable `valueAtMiddle` définie à `searchList[mid]`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+valueAtMiddle\s*=\s*searchList\s*\[\s*mid\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
