---
id: 6996c796385b6ea4e222d5e8
title: Étape 5
challengeType: 1
dashedName: step-5
---

# --description--

Pour démarrer avec la boucle, tu dois trouver l'indice du milieu de l'espace de recherche actuel.

Tu peux obtenir l'indice du milieu en calculant la moyenne des variables `low` et `high`. Tu devras utiliser la division entière (floor) pour obtenir la moyenne après avoir ajouté les deux valeurs, afin que la réponse soit toujours un entier arrondi vers le bas, car les indices doivent être des entiers.

Dans ta boucle `while`, crée une variable `mid` définie à la moyenne de `low` et `high`.

# --hints--

Dans la boucle while, tu dois créer une variable `mid` définie à `Math.floor((low + high) / 2)`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+mid\s*=\s*Math\.floor\(\s*\(\s*low\s*\+\s*high\s*\)\s*\/\s*2\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
