---
id: 6996ccadf9048c79b857bb15
title: Étape 11
challengeType: 1
dashedName: step-11
---

# --description--

Tu peux voir que le premier appel renvoie `[3]`. C'est parce que `3` est la valeur du milieu dans `[1, 2, 3, 4, 5]`.

En revanche, le deuxième appel renvoie un tableau vide. Cela se produit parce qu'après avoir vérifié la valeur du milieu, la boucle manque actuellement la logique pour réduire la plage de recherche en mettant à jour les variables `low` ou `high`. Elle n'agit que comme une seule vérification du milieu, puis se termine sans chercher davantage.

Pour permettre à la recherche binaire de continuer à réduire sa recherche, ajoute une instruction `else if` qui vérifie si `value` est supérieure à `valueAtMiddle`.

# --hints--

Tu dois ajouter une instruction `else if` avec la condition `value > valueAtMiddle`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /else\s+if\s*\(\s*value\s*>\s*valueAtMiddle\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
--fcc-editable-region--
    if (value === valueAtMiddle) {
      return pathToTarget;
    }
--fcc-editable-region--
    break;
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
```
