---
id: 6996c95e5bede03bdf2937a2
title: Étape 8
challengeType: 1
dashedName: step-8
---

# --description--

La prochaine chose à faire est de créer une condition qui vérifiera si la `value` cible est au milieu.

À l'intérieur de la boucle `while`, crée une instruction `if` qui vérifie si la `value` cible est égale à `valueAtMiddle`. Si c'est le cas, renvoie la variable `pathToTarget`.

# --hints--

Tu dois créer une instruction `if` avec la condition `value === valueAtMiddle`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)/
);
```

Tu dois renvoyer la variable `pathToTarget` à l'intérieur de l'instruction `if`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?return\s+pathToTarget/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
