---
id: 587d825a367417b2b2512c89
title: Implémenter l'algorithme de tri rapide
challengeType: 26
dashedName: implement-the-quicksort-algorithm
---

# --description--

Dans ce lab, tu vas implémenter l'algorithme de tri : le tri rapide (quicksort).

Le tri rapide est une approche récursive efficace de type diviser pour régner pour trier un tableau. Dans cette méthode, une valeur pivot est choisie dans le tableau d'origine. Le tableau est ensuite partitionné en deux sous-tableaux de valeurs inférieures et de valeurs supérieures à la valeur pivot. Ensuite, le résultat des appels récursifs de l'algorithme de tri rapide sur les deux sous-tableaux est combiné. Cela continue jusqu'à atteindre le cas de base d'un tableau vide ou à un seul élément, qui est renvoyé. Le déroulement des appels récursifs redonne le tableau trié.

Bien que le choix de la valeur pivot soit important, n'importe quel pivot convient pour nos besoins ici. Pour simplifier, le premier ou le dernier élément peut être utilisé.

Le tri rapide est une méthode de tri très efficace, offrant en moyenne des performances *O(nlog(n))*. Il est aussi relativement facile à implémenter. Ces attributs en font une méthode de tri populaire et utile.

**Objectif :** Remplir les user stories ci-dessous et faire passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu dois avoir une fonction `quicksort` qui prend un tableau d'entiers en entrée.
1. La fonction `quicksort` doit renvoyer un tableau qui contient les mêmes entiers que l'entrée, mais dans l'ordre du plus bas au plus haut.
1. La fonction `quicksort` doit s'appeler récursivement pour trier le tableau.

# --hints--

Tu dois avoir une fonction `quicksort`.

```js
assert.isFunction(quicksort);
```

`quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` doit renvoyer un tableau inchangé sauf pour l'ordre.

```js
assert.sameMembers(
  quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`quicksort` doit renvoyer un tableau trié (du plus petit au plus grand).

```js
assert.sameOrderedMembers(
  quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
);
```

`quicksort` ne doit pas utiliser la méthode intégrée `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    quicksort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function quicksort(array) {
  if (array.length === 0) {
    return [];
  } else {
    const pivotValue = array[0];

    // Sort elements into three piles
    let lesser = [];
    let equal = [];
    let greater = [];
    for (let e of array) {
      if (e < pivotValue) {
        lesser.push(e);
      } else if (e > pivotValue) {
        greater.push(e);
      } else {
        equal.push(e);
      }
    }

    return [...quicksort(lesser), ...equal, ...quicksort(greater)];
  }
}
```
