---
id: 668f08ea07b99b1f4a91acab
title: Créer une page de recette
challengeType: 25
dashedName: build-a-recipe-page
demoType: onClick
---

# --description--

**Objectif :** respecte les user stories ci-dessous et fais passer tous les tests pour terminer ce lab.

**Histoires utilisateur :**

1. Tu dois avoir une déclaration `!DOCTYPE html`.
1. Tu dois avoir un élément `html` avec `lang` défini sur `en`.
1. Tu dois avoir un élément `head` contenant un élément `title` avec le nom de ta recette, et un élément `meta` avec un attribut `charset` défini sur `UTF-8`.
1. Tu dois avoir un élément `body`.
1. Tu dois avoir un élément `h1` avec le nom de ta recette.
1. Tu dois avoir un élément `p` qui présente la recette sous le `h1`.
1. Tu dois avoir un élément `h2` avec le texte `Ingredients` pour la section des ingrédients.
1. Tu dois avoir une liste non ordonnée (élément `ul`) avec au moins quatre éléments de liste (`li`) qui liste tes ingrédients sous le premier élément `h2`.
1. Tu dois avoir un deuxième élément `h2` avec le texte `Instructions` pour la section des instructions.
1. Tu dois avoir une liste ordonnée (élément `ol`) avec au moins quatre éléments de liste qui liste les étapes de la recette dans l'ordre, sous le deuxième `h2`.
1. Tu dois avoir un élément `img` avec un attribut `src` défini sur une image valide. Tu peux utiliser `https://cdn.freecodecamp.org/curriculum/labs/recipe.jpg` si tu veux, et un attribut `alt` qui décrit l'image.

# --hints--

Ta page de recette doit avoir une déclaration `!DOCTYPE html`.

```js
assert.match(code, /<!DOCTYPE html>/i);
```

Tu dois avoir un élément `html` avec `lang` défini sur `en`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[\s\S]*<\/\s*html\s*>/gi);
```

Tu dois avoir un élément `head` dans l'élément `html`.

```js
assert.match(code, /<html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\/\s*html\s*>/i);
```

Tu dois avoir un élément `title` dans ton élément `head`.

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*title\s*>[\s\S]*<\/\s*title\s*>[\s\S]*<\/\s*head\s*>/i);
```

Ton élément `title` doit contenir le titre de ta recette.

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

Tu dois avoir un élément `meta` dans ton élément `head`.

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*meta[\s\S]*>[\s\S]*<\/\s*head\s*>/i);
```

Ton élément `meta` doit avoir son attribut `charset` défini sur `UTF-8`.

```js
assert.match(code, /<\s*meta[\s\S]+?charset\s*=\s*('|")UTF-8\1/i);
```

Tu dois avoir un élément `body` dans ton élément `html`.

```js
assert.match(code, /<\s*html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\s*body\s*>[\s\S]*<\/\s*body\s*>[\s\S]*<\/\s*html\s*>/i);
```

Tu dois avoir un élément `h1` avec le nom de ta recette.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

Tu dois avoir un seul élément `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Tu dois avoir un élément `p` sous ton élément `h1`.

```js
assert.strictEqual(document.querySelector('h1')?.nextElementSibling, document.querySelector('p'));
```

Ton premier élément `p` doit décrire ta recette.

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

Ton premier élément `h2` doit avoir le texte `Ingredients`.

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText, 'Ingredients');
```

Tu dois avoir un élément de liste non ordonnée sous ton premier élément `h2`.

```js
assert.strictEqual(document.querySelector('ul')?.previousElementSibling.tagName, 'H2');
```

Tu dois avoir au moins quatre éléments de liste dans ta liste non ordonnée avec les ingrédients.

```js
const els = document.querySelectorAll('ul > li');
assert.isAbove(els.length, 3);
els.forEach(el => assert.isAbove(el.innerText.trim().length, 0))
```

Ton deuxième élément `h2` doit avoir le texte `Instructions`.

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText, 'Instructions');
```

Tu dois avoir un élément de liste ordonnée sous ton deuxième élément `h2`.

```js
assert.strictEqual(document.querySelectorAll('h2')?.[1]?.nextElementSibling?.tagName, "OL");
```

Tu dois avoir au moins quatre éléments de liste dans ta liste ordonnée avec les instructions.

```js
const els = document.querySelectorAll('ol > li');
assert.isAbove(els.length, 3);
els.forEach(el => assert.isAbove(el.innerText.trim().length, 0))
```

Tu dois avoir au moins un élément `img`.

```js
assert.exists(document.querySelector('img'));
```

Tous tes éléments `img` doivent avoir un attribut `src` et une valeur valides.

```js
const img = document.querySelector('img');
const rawSrc = img?.getAttribute('src');
const resolvedSrc = img?.src;
const re = new RegExp(window.location.href, "ig");

assert.isAbove(rawSrc?.trim().length, 0, "The 'src' attribute must be explicitly set.");
assert.notMatch(resolvedSrc, re, "The 'src' should not start with the current page URL");

img.onload = () => {
  console.log('Image loaded successfully.');
};

img.onerror = (error) => {
  console.error('Image failed to load:', error);
  assert.fail("Your image's URL should be valid."); // Make the test instafail
};

if (img.complete) {
  img.onload && img.onload();
};
```

Tous tes éléments `img` doivent avoir un attribut `alt` pour décrire l'image.

```js
assert.isAbove(document.querySelector('img')?.alt.length, 0);
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Chocolate chip cookies recipe</title>
</head>

<body>
  <h1>Chocolate Chip Cookies</h1>
  <p>Welcome to the ultimate guide for making mini chocolate chip cookies! These bite-sized treats are perfect for
    satisfying your sweet tooth without overindulging. Follow this simple recipe to create delicious,
    crispy-on-the-outside, chewy-on-the-inside mini chocolate chip cookies that everyone will love.</p>
  <img src="https://cdn.freecodecamp.org/curriculum/labs/recipe.jpg" alt="Ingredients for baking: three eggs, a bowl of flour, a glass of milk, and a whisk arranged on a wooden table.">
  <h2>Ingredients</h2>
  <ul>
    <li>1 cup all-purpose flour</li>
    <li>1/2 teaspoon baking soda</li>
    <li>1/4 cup unsalted butter, softened</li>
    <li>1/4 cup granulated sugar</li>
    <li>1/2 teaspoon vanilla extract</li>
    <li>1/2 cup mini chocolate chips</li>
  </ul>
  <h2>Instructions</h2>
  <ol>
    <li>Preheat your oven to 350°F (175°C) and line a baking sheet with parchment paper.</li>
    <li>In a bowl, whisk together the flour and baking soda.</li>
    <li>In another bowl, beat the butter, sugar, and vanilla extract until creamy.</li>
    <li>Gradually add the dry ingredients to the wet mixture, then fold in the mini chocolate chips.</li>
    <li>Drop small spoonfuls of dough onto the baking sheet.</li>
    <li>Bake for 8-10 minutes, then let cool before enjoying!</li>
  </ol>
</body>

</html>
```
