---
id: 67172b43f84bcd2dec238a3d
title: Crée un organisateur de livres
challengeType: 26
dashedName: build-a-book-organizer
---

# --description--

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois avoir un tableau d'objets nommé `books`, où chaque objet du tableau doit avoir une chaîne `title`, une autre chaîne `authorName` et un nombre `releaseYear`.

2. Ton tableau `books` doit contenir au minimum trois objets.

3. Tu dois avoir une fonction de rappel nommée `sortByYear` qui accepte deux livres comme paramètres pour trier le tableau.

4. La fonction `sortByYear` doit renvoyer `-1` si le `releaseYear` du premier livre est inférieur à celui du second livre.

5. La fonction `sortByYear` doit renvoyer `1` si le `releaseYear` du premier livre est supérieur à celui du second livre.

6. La fonction `sortByYear` doit renvoyer `0` si les deux valeurs `releaseYear` sont égales.

7. Tu dois filtrer les livres écrits après une certaine année, comme 1950, depuis le tableau `books` et enregistrer le tableau filtré dans un nouveau tableau nommé `filteredBooks`.

8. Tu dois trier les livres du tableau `filteredBooks` selon leur `releaseYear` dans l'ordre croissant. Tu as appris dans une leçon précédente que la méthode `sort()` trie le tableau sur place. Cela signifie que le tableau `filteredBooks` sera modifié.

# --hints--

Tu dois avoir une fonction `sortByYear` dans ton code.

```js
assert.isFunction(sortByYear);
```

Ta fonction `sortByYear` doit prendre deux paramètres.

```js
assert.lengthOf(sortByYear, 2);
```

Ta fonction `sortByYear` doit renvoyer `-1` si le `releaseYear` de l'objet `book1` est inférieur à celui de l'objet `book2`, `1` si le `releaseYear` de l'objet `book1` est supérieur à celui de l'objet `book2`, et `0` dans tous les autres cas.

```js
assert.equal(sortByYear({releaseYear: 1913}, {releaseYear: 1925}), -1);
assert.equal(sortByYear({releaseYear: 1925}, {releaseYear: 1913}), 1);
assert.equal(sortByYear({releaseYear: 1925}, {releaseYear: 1925}), 0);
```

Tu dois avoir un tableau `books` dans ton code.

```js
assert.isArray(books)
```

Ton tableau `books` doit contenir au moins trois objets de livre.

```js
assert.isAtLeast(books.length, 3)
```

Ton tableau `books` doit contenir des objets ayant chacun une chaîne `title`, une autre chaîne `authorName` et un nombre `releaseYear`.

```js

books.forEach((book) => {
  assert.isObject(book);
  assert.sameMembers(Object.keys(book), ["title", "authorName", "releaseYear"]);
  assert.isString(book.title);
  assert.isNotEmpty(book.title);
  assert.isString(book.authorName);
  assert.isNotEmpty(book.authorName);
  assert.isNumber(book.releaseYear);
})
```

Tu dois avoir un tableau `filteredBooks` dans ton code.

```js
assert.isArray(filteredBooks)
```

Le tableau `filteredBooks` doit contenir certains livres filtrés depuis le tableau `books` et ne doit pas être vide.

```js
assert.isBelow(filteredBooks.length, books.length)
assert.isNotEmpty(filteredBooks)
```

Le tableau `filteredBooks` doit contenir uniquement des livres provenant du tableau `books` d'origine.

```js
filteredBooks.forEach(book => {
  assert.isTrue(
    books.some(original =>
      original.title === book.title &&
      original.authorName === book.authorName &&
      original.releaseYear === book.releaseYear
    )
  );
});
```

Le tableau `filteredBooks` doit inclure uniquement des livres publiés pendant ou avant une année donnée.

```js
const maxYear = Math.max(...filteredBooks.map(book => book.releaseYear));
const removedYears = books.filter(book => !filteredBooks.includes(book)).map(book => book.releaseYear);
assert.isTrue(removedYears.every(year => year > maxYear));
```

Tu dois appeler la fonction d'ordre supérieur `sort` en passant la fonction de rappel `sortByYear` sur le tableau `filteredBooks`.

```js
assert.match(__helpers.removeJSComments(code), /\s*filteredBooks\s*\.\s*sort\s*\(\s*sortByYear\s*\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sortByYear(book1, book2) {
  if (book1.releaseYear < book2.releaseYear) {
    return -1;
  }

  if (book1.releaseYear > book2.releaseYear) {
    return 1;
  }

  return 0;
}

const books = [
  {
    title: "The Great Gatsby",
    authorName: "F. Scott Fitzgerald",
    releaseYear: 1925,
  },
  {
    title: "In Search of Lost Time",
    authorName: "Marcel Proust",
    releaseYear: 1913,
  },
  {
    title: "Ulysses",
    authorName: "James Joyce",
    releaseYear: 1922,
  },
  {
    title: "One Hundred Years of Solitude",
    authorName: "Gabriel García Márquez",
    releaseYear: 1967,
  },
  {
    title: "The Catcher in the Rye",
    authorName: "J. D. Salinger",
    releaseYear: 1951,
  },
];

let filteredBooks = books.filter((book) => book.releaseYear < 1950);

filteredBooks.sort(sortByYear);

console.log("Books Written Before 1950 (sorted according to release year)");
filteredBooks.forEach((book) => {
  console.log(`${book.title} by ${book.authorName} (${book.releaseYear})`);
});
```
