---
id: ac6993d51946422351508a41
title: Implémente l'algorithme de troncature de chaîne
challengeType: 26
dashedName: truncate-a-string
---

# --description--

Dans ce lab, tu vas t'entraîner à tronquer une chaîne à une certaine longueur.

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois avoir une fonction `truncateString` qui accepte deux arguments : le premier une chaîne, le second un nombre.
1. Si la longueur de la chaîne est supérieure au nombre donné, la chaîne doit être tronquée pour réduire sa longueur afin qu'elle soit égale au nombre donné, et `...` doit être ajouté à la fin de la chaîne tronquée.
1. Si la longueur de la chaîne est inférieure ou égale au nombre donné, la chaîne doit être retournée sans modification.

# --hints--

`truncateString("A-tisket a-tasket A green and yellow basket", 8)` doit retourner la chaîne `A-tisket...`.

```js
assert.strictEqual(
  truncateString('A-tisket a-tasket A green and yellow basket', 8),
  'A-tisket...'
);
```

`truncateString("Peter Piper picked a peck of pickled peppers", 11)` doit retourner la chaîne `Peter Piper...`.

```js
assert.strictEqual(
  truncateString('Peter Piper picked a peck of pickled peppers', 11),
  'Peter Piper...'
);
```

`truncateString("A-tisket a-tasket A green and yellow basket", "A-tisket a-tasket A green and yellow basket".length)` doit retourner la chaîne `A-tisket a-tasket A green and yellow basket`.

```js
assert.strictEqual(
  truncateString(
    'A-tisket a-tasket A green and yellow basket',
    'A-tisket a-tasket A green and yellow basket'.length
  ),
  'A-tisket a-tasket A green and yellow basket'
);
```

`truncateString("A-tisket a-tasket A green and yellow basket", "A-tisket a-tasket A green and yellow basket".length + 2)` doit retourner la chaîne `A-tisket a-tasket A green and yellow basket`.

```js
assert.strictEqual(
  truncateString(
    'A-tisket a-tasket A green and yellow basket',
    'A-tisket a-tasket A green and yellow basket'.length + 2
  ),
  'A-tisket a-tasket A green and yellow basket'
);
```

`truncateString("A-", 1)` doit retourner la chaîne `A...`.

```js
assert.strictEqual(truncateString('A-', 1), 'A...');
```

`truncateString("Absolutely Longer", 2)` doit retourner la chaîne `Ab...`.

```js
assert.strictEqual(truncateString('Absolutely Longer', 2), 'Ab...');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function truncateString(str, num) {
  if (num >= str.length) {
    return str;
  }

  return str.slice(0, num) + '...';
}

```
