---
id: 67051431a73f1ca25d9a6f25
title: Construire un tableau d'idées de projets
challengeType: 26
dashedName: build-a-project-idea-board
---

# --description--

**Objectif :** Remplir les user stories ci-dessous et faire passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu dois définir une constante objet nommée `projectStatus` avec les trois clés : `PENDING`, `SUCCESS` et `FAILURE`. Chaque statut doit être associé à un objet avec une clé `description` dont la valeur est respectivement `Pending Execution`, `Executed Successfully` et `Execution Failed`.
2. Tu dois définir une classe nommée `ProjectIdea` avec un `constructor` qui prend `title` et `description` comme paramètres. Initialise les propriétés `title` et `description` avec les paramètres fournis. La classe doit aussi avoir une propriété nommée `status` qui est définie par défaut à la valeur `projectStatus.PENDING`.
3. Tu dois définir une méthode nommée `updateProjectStatus` dans la classe `ProjectIdea`. Cette méthode doit accepter un paramètre `newStatus` et mettre à jour la propriété `status` avec la valeur donnée.
4. Tu dois définir une classe `ProjectIdeaBoard` avec un `constructor` qui accepte un `title` et initialise un tableau vide nommé `ideas` pour contenir des instances de la classe `ProjectIdea`.
5. Tu dois définir une méthode nommée `pin` dans la classe `ProjectIdeaBoard` qui accepte une instance de la classe `ProjectIdea` et ajoute l'instance donnée au tableau `ideas`.
6. Tu dois définir une méthode nommée `unpin` dans la classe `ProjectIdeaBoard`. Cette méthode doit accepter une instance de la classe `ProjectIdea` et la retirer du tableau `ideas`.
7. Tu dois définir une méthode nommée `count` qui renvoie le nombre d'idées de projets dans le tableau `ProjectIdeaBoard` donné.
8. Tu dois définir une méthode nommée `formatToString` qui renvoie le nom des idées de projets, leur description et leur statut au format :

```js
<ProjectIdeaBoard title> has <ProjectIdeaBoard count> idea(s)
<ProjectIdea title> (<ProjectIdea status description>) - <ProjectIdea description>
<ProjectIdea title> (<ProjectIdea status description>) - <ProjectIdea description>
.
.
.
```

# --hints--

Tu dois définir une constante objet nommée `projectStatus` avec les bons statuts et descriptions de projet.

```js
assert.isObject(projectStatus);
assert.sameMembers(Object.keys(projectStatus), ["PENDING", "SUCCESS", "FAILURE"]);
assert.deepEqual(projectStatus.PENDING, { description: "Pending Execution" });
assert.deepEqual(projectStatus.SUCCESS, { description: "Executed Successfully" });
assert.deepEqual(projectStatus.FAILURE, { description: "Execution Failed" });
```

Tu dois avoir une classe `ProjectIdea`.

```js
assert.isFunction(ProjectIdea);
assert.isObject(new ProjectIdea());
```

Ta classe `ProjectIdea` doit initialiser les propriétés `title` et `description` d'après les paramètres passés. Elle doit aussi définir le statut à `projectStatus.PENDING`.

```js
const test = new ProjectIdea("sample title", "sample description");
assert.isAbove(test.title.length, 0);
assert.isAbove(test.description.length, 0);
assert.strictEqual(test.status.description, "Pending Execution");
```

Ta classe `ProjectIdea` doit avoir une méthode `updateProjectStatus`.

```js
const projectIdea = new ProjectIdea();
assert.isFunction(projectIdea.updateProjectStatus);
```

Tu dois avoir une classe `ProjectIdeaBoard`.

```js
assert.isFunction(ProjectIdeaBoard);
assert.isObject(new ProjectIdeaBoard());
```

Ta classe `ProjectIdeaBoard` doit initialiser la propriété `title` d'après le paramètre passé et initialiser un tableau vide nommé `ideas` pour contenir des instances de la classe `ProjectIdea`.

```js
const test = new ProjectIdeaBoard("Sample Title");
assert.isArray(test.ideas);
assert.lengthOf(test.ideas, 0);
assert.exists(test.title);
```

Ta classe `ProjectIdeaBoard` doit avoir une méthode `pin`.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.pin);
```

Ta classe `ProjectIdeaBoard` doit avoir une méthode `unpin`.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.unpin);
```

Ta classe `ProjectIdeaBoard` doit avoir une méthode `count`.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.count);
```

Ta classe `ProjectIdeaBoard` doit avoir une méthode `formatToString`.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.formatToString);
```

`new ProjectIdea("Smart Window Locks", "An automation project allowing users to lock, unlock windows automatically based on weather conditions.")` devrait renvoyer `{ title: 'Smart Window Locks', description: 'An automation project allowing users to lock, unlock windows automatically based on weather conditions.', status: { description: 'Pending Execution' } }`.

```js
const windowLocks = new ProjectIdea(
  "Smart Window Locks",
  "An automation project allowing users to lock, unlock windows automatically based on weather conditions."
);
assert.deepEqual(windowLocks, {
  title: "Smart Window Locks",
  description:
    "An automation project allowing users to lock, unlock windows automatically based on weather conditions.",
  status: { description: "Pending Execution" },
});
```

Appeler `updateProjectStatus(projectStatus.SUCCESS)` sur `new ProjectIdea("Fitness Tracker App", "An app that tracks user workouts, diet, and sleep patterns.")` devrait mettre à jour le statut à `{ description: 'Executed Successfully' }`.

```js
let mobileApp = new ProjectIdea(
  "Fitness Tracker App",
  "An app that tracks user workouts, diet, and sleep patterns."
);

// Update the project status to SUCCESS
mobileApp.updateProjectStatus(projectStatus.SUCCESS);
assert.deepEqual(mobileApp, {
  title: "Fitness Tracker App",
  description: "An app that tracks user workouts, diet, and sleep patterns.",
  status: { description: "Executed Successfully" },
});
```

Appeler `updateProjectStatus(projectStatus.FAILURE)` sur `new ProjectIdea("Breakfast Chef Robot", "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone.")` devrait mettre à jour le statut à `{ description: 'Execution Failed' }`.

```js
let breakfastChef = new ProjectIdea(
  "Breakfast Chef Robot",
  "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone."
);

// Update the project status to FAILURE
breakfastChef.updateProjectStatus(projectStatus.FAILURE);

assert.deepEqual(breakfastChef, {
  title: "Breakfast Chef Robot",
  description:
    "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone.",
  status: { description: "Execution Failed" },
});
```

Appeler `updateProjectStatus(projectStatus.SUCCESS)` sur `new ProjectIdea("Online Used Video Games Store", "An online platform where users can buy second hand physical copies of video games from other users.")` devrait mettre à jour le statut à `{ description: 'Executed Successfully' }`.

```js
let usedVideoGamesStore = new ProjectIdea(
  "Online Used Video Games Store",
  "An online platform where users can buy second hand physical copies of video games from other users."
);

// Update the project status set to SUCCESS
usedVideoGamesStore.updateProjectStatus(projectStatus.SUCCESS);

assert.deepEqual(usedVideoGamesStore, {
  title: "Online Used Video Games Store",
  description:
    "An online platform where users can buy second hand physical copies of video games from other users.",
  status: { description: "Executed Successfully" },
});
```

Tu dois pouvoir épingler un objet `ProjectIdea` à ton `ProjectIdeaBoard` avec la méthode `pin`.

```js
let smartHome = new ProjectIdea(
  "Smart Home System",
  "An integrated system to control lighting, temperature, and security devices remotely."
);

// Create a project board and pin the idea
let techProjects = new ProjectIdeaBoard("Tech Projects Board");
techProjects.pin(smartHome);

// Correct the assertion to compare the first idea in the array
assert.deepEqual(techProjects.ideas[0], {
    title: 'Smart Home System',
    description: 'An integrated system to control lighting, temperature, and security devices remotely.',
    status: { description: 'Pending Execution' }
});
```

Tu dois pouvoir désépingler un objet `ProjectIdea` de ton `ProjectIdeaBoard` avec la méthode `unpin`.

```js
let smartHome = new ProjectIdea(
  "Smart Home System",
  "An integrated system to control lighting, temperature, and security devices remotely."
);

// Create a project board and pin ideas
let techProjects = new ProjectIdeaBoard("Tech Projects Board");

techProjects.pin(smartHome);
techProjects.unpin(smartHome);

assert.isEmpty(techProjects.ideas);
```

Quand `new ProjectIdeaBoard("Empty Board")` est vide, `emptyBoard.formatToString()` devrait renvoyer `Empty Board has 0 idea(s)\n`.

```js
let emptyBoard = new ProjectIdeaBoard("Empty Board");

assert.equal(emptyBoard.formatToString(), "Empty Board has 0 idea(s)\n");
```

Quand tu épingles `new ProjectIdea("Smart Home System", "An integrated system to control lighting, temperature, and security devices remotely.")` à `new ProjectIdeaBoard("Tech Projects Board")`, `techProjects.formatToString()` devrait renvoyer `Tech Projects Board has 1 idea(s)\nSmart Home System (Pending Execution) - An integrated system to control lighting, temperature, and security devices remotely.\n`.

```js
let techProjectsBoard = new ProjectIdeaBoard("Tech Projects Board");
let smartHome = new ProjectIdea("Smart Home System", "An integrated system to control lighting, temperature, and security devices remotely.");
techProjectsBoard.pin(smartHome);

assert.equal(techProjectsBoard.formatToString(), "Tech Projects Board has 1 idea(s)\nSmart Home System (Pending Execution) - An integrated system to control lighting, temperature, and security devices remotely.\n");
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const projectStatus = {
  PENDING: { description: "Pending Execution" },
  SUCCESS: { description: "Executed Successfully" },
  FAILURE: { description: "Execution Failed" },
};

class ProjectIdea {
  constructor(title, description) {
    this.title = title;
    this.description = description;
    this.status = projectStatus.PENDING;
  }

  updateProjectStatus(newStatus) {
    this.status = newStatus;
  }
}

class ProjectIdeaBoard {
  constructor(title) {
    this.title = title;
    this.ideas = [];
  }

  pin(newIdea) {
    this.ideas.push(newIdea);
  }

  unpin(idea) {
    this.ideas.splice(this.ideas.indexOf(idea), 1);
  }

  count() {
    return this.ideas.length;
  }

  formatToString() {
    let representation = `${this.title} has ${this.count()} idea(s)\n`;

    this.ideas.forEach(
      (idea) =>
        (representation += `${idea.title} (${idea.status.description}) - ${idea.description}\n`),
    );
    return representation;
  }
}

let windowLocks = new ProjectIdea(
  "Smart Window Locks",
  "An automation project allowing users to lock, unlock windows automatically based on weather conditions."
);

let breakfastChef = new ProjectIdea(
  "Breakfast Chef Robot",
  "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone."
);
breakfastChef.updateProjectStatus(projectStatus.FAILURE);

let usedVideoGamesStore = new ProjectIdea(
  "Online Used Video Games Store",
  "An online platform where users can buy second hand physical copies of video games from other users."
);
usedVideoGamesStore.updateProjectStatus(projectStatus.SUCCESS);

let hardwareProjectIdeas = new ProjectIdeaBoard("Hardware Project Board");
hardwareProjectIdeas.pin(windowLocks);
hardwareProjectIdeas.pin(breakfastChef);

let softwareProjectIdeas = new ProjectIdeaBoard("Software Project Board");
softwareProjectIdeas.pin(usedVideoGamesStore);

console.log(softwareProjectIdeas.formatToString());

console.log(hardwareProjectIdeas.formatToString());
hardwareProjectIdeas.unpin(breakfastChef);
console.log(hardwareProjectIdeas.formatToString());
```
