---
id: 61fd70336ebb3e4f62ee81ba
title: Étape 8
challengeType: 0
dashedName: step-8
---

# --description--

Les tableaux HTML utilisent l'élément `caption` pour décrire le sujet du tableau. L'élément `caption` doit toujours être le premier enfant d'un `table`, mais il peut être positionné avec la propriété CSS `caption-side`.

Ajoute un élément `caption` à ton premier `table`, et donne-lui le texte `Assets`.

# --hints--

Ton premier élément `table` doit avoir un élément `caption`.

```js
assert.strictEqual(
  document.querySelector('table')?.children?.[0]?.localName,
  'caption'
);
```

Ton élément `caption` doit avoir le texte `Assets`.

```js
assert.strictEqual(document.querySelector('caption')?.textContent.trim(), 'Assets');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
--fcc-editable-region--
          <table>
          </table>
--fcc-editable-region--
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
