---
id: 61fd71d596e8f253b9408b39
title: Étape 10
challengeType: 0
dashedName: step-10
---

# --description--

L'élément `tr` sert à indiquer une ligne de tableau. Ajoute un élément `tr` dans ton élément `thead`. Dans ton nouveau `tr`, ajoute un élément `td`, suivi de trois éléments `th`.

L'élément `td` indique une cellule de données, tandis que l'élément `th` indique une cellule d'en-tête.

# --hints--

Ton élément `thead` doit avoir un élément `tr`.

```js
assert.strictEqual(
  document.querySelector('thead')?.children?.[0]?.localName,
  'tr'
);
```

Ton élément `tr` doit avoir un élément `td` comme premier enfant.

```js
assert.strictEqual(
  document.querySelector('tr')?.children?.[0]?.localName,
  'td'
);
```

Ton élément `tr` doit avoir trois éléments `th`, après l'élément `td`.

```js
assert.strictEqual(
  document.querySelector('tr')?.children?.[1]?.localName,
  'th'
);
assert.strictEqual(
  document.querySelector('tr')?.children?.[2]?.localName,
  'th'
);
assert.strictEqual(
  document.querySelector('tr')?.children?.[3]?.localName,
  'th'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
--fcc-editable-region--
            <thead>
            </thead>
--fcc-editable-region--
            <tbody></tbody>
          </table>
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
