---
id: 61fd9b7285bde783ad5b8aac
title: Étape 27
challengeType: 0
dashedName: step-27
---

# --description--

Dans le `tbody`, ajoute un `tr` avec la `class` définie sur `total`. Dans celui-ci, ajoute un `th` avec le texte `Total Net Worth`, et encadre `Net Worth` dans un `span` avec la `class` définie sur `sr-only`.

Ajoute ensuite trois éléments `td`, en donnant au troisième une `class` définie sur `current`, et en donnant à chacun le texte suivant : `$-171`, `$136`, `$334`.

# --hints--

Ton élément `tbody` doit avoir un élément `tr`.

```js
const table = document.querySelectorAll('table')?.[2];
const tbody = table?.querySelector('tbody');
assert.lengthOf(tbody?.querySelectorAll('tr'), 1);
```

Ton élément `tr` doit avoir la `class` définie sur `total`.

```js
const table = document.querySelectorAll('table')?.[2];
const tbody = table?.querySelector('tbody');
assert.isTrue(tbody?.querySelector('tr')?.classList?.contains('total'));
```

Ton `tr` doit avoir un élément `th`.

```js
const table = document.querySelectorAll('table')?.[2];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[0];
assert.isNotNull(tableRow?.querySelector('th'));
```

Ton élément `th` doit avoir le texte `Total Net Worth`.

```js
const table = document.querySelectorAll('table')?.[2];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[0];
assert.strictEqual(tableRow?.querySelector('th')?.innerText.trim(), 'Total Net Worth');
```

Tu dois encadrer le texte `Net Worth` dans un élément `span`.

```js
const table = document.querySelectorAll('table')?.[2];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[0];
assert.strictEqual(
  tableRow?.querySelector('th > span')?.textContent.trim(),
  'Net Worth'
);
```

Ton élément `span` doit avoir l'attribut `class` défini sur `sr-only`.

```js
const table = document.querySelectorAll('table')?.[2];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[0];
assert.isTrue(
  tableRow?.querySelector('th > span')?.classList?.contains('sr-only')
);
```

Tu dois avoir trois éléments `td`.

```js
const table = document.querySelectorAll('table')?.[2];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[0];
assert.lengthOf(tableRow?.querySelectorAll('td'), 3);
```

Ton premier élément `td` doit avoir le texte `$-171`.

```js
const table = document.querySelectorAll('table')?.[2];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[0];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[0]?.textContent.trim(), '$-171');
```

Ton deuxième élément `td` doit avoir le texte `$136`.

```js
const table = document.querySelectorAll('table')?.[2];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[0];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[1]?.textContent.trim(), '$136');
```

Ton troisième élément `td` doit avoir le texte `$334`.

```js
const table = document.querySelectorAll('table')?.[2];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[0];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[2]?.textContent.trim(), '$334');
```

Ton troisième élément `td` doit avoir la `class` définie sur `current`.

```js
const table = document.querySelectorAll('table')?.[2];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[0];
assert.isTrue(
  tableRow?.querySelectorAll('td')?.[2]?.classList?.contains('current')
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>
              Assets
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Cash <span class="description">This is the cash we currently have on hand.</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
              <tr class="data">
                <th>Checking <span class="description">Our primary transactional account.</span></th>
                <td>$54</td>
                <td>$56</td>
                <td class="current">$53</td>
              </tr>
              <tr class="data">
                <th>Savings <span class="description">Funds set aside for emergencies.</span></th>
                <td>$500</td>
                <td>$650</td>
                <td class="current">$728</td>
              </tr>
              <tr class="total">
                <th>Total <span class="sr-only">Assets</span></th>
                <td>$579</td>
                <td>$736</td>
                <td class="current">$809</td>
              </tr>
            </tbody>
          </table>
          <table>
            <caption>
              Liabilities
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only">2019</span></th>
                <th><span class="sr-only">2020</span></th>
                <th><span class="sr-only">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Loans <span class="description">The outstanding balance on our startup loan.</span></th>
                <td>$500</td>
                <td>$250</td>
                <td class="current">$0</td>
              </tr>
              <tr class="data">
                <th>Expenses <span class="description">Annual anticipated expenses, such as payroll.</span></th>
                <td>$200</td>
                <td>$300</td>
                <td class="current">$400</td>
              </tr>
              <tr class="data">
                <th>Credit <span class="description">The outstanding balance on our credit card.</span></th>
                <td>$50</td>
                <td>$50</td>
                <td class="current">$75</td>
              </tr>
              <tr class="total">
                <th>Total <span class="sr-only">Liabilities</span></th>
                <td>$750</td>
                <td>$600</td>
                <td class="current">$475</td>
              </tr>
            </tbody>
          </table>
--fcc-editable-region--
          <table>
            <caption>
              Net Worth
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only">2019</span></th>
                <th><span class="sr-only">2020</span></th>
                <th><span class="sr-only">2021</span></th>
              </tr>
            </thead>
            <tbody>
            </tbody>
          </table>
--fcc-editable-region--
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
