---
id: 61fd6b7c83dbf54a08cf0498
title: Étape 6
challengeType: 0
dashedName: step-6
---

# --description--

Dans ton élément `div`, ajoute trois éléments `span`. Donne à chacun un attribut `class` défini sur `year`, et ajoute le texte suivant (dans l'ordre) : `2019`, `2020` et `2021`.

# --hints--

Ton élément `div` doit avoir trois éléments `span`.

```js
assert.lengthOf(document.querySelector('div')?.children, 3);
```

Chaque élément `span` doit avoir l'attribut `class` défini sur `year`.

```js
const spans = [...document.querySelector('div')?.children];
spans.forEach(span => assert.isTrue(span?.classList?.contains('year')));
```

Ton premier `span` doit avoir le texte `2019`.

```js
assert.strictEqual(
  document.querySelector('div')?.children?.[0]?.textContent.trim(),
  '2019'
);
```

Ton deuxième `span` doit avoir le texte `2020`.

```js
assert.strictEqual(
  document.querySelector('div')?.children?.[1]?.textContent.trim(),
  '2020'
);
```

Ton troisième `span` doit avoir le texte `2021`.

```js
assert.strictEqual(
  document.querySelector('div')?.children?.[2]?.textContent.trim(),
  '2021'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--
        <div id="years" aria-hidden="true">
        </div>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
