---
id: 61fd8fd08af43372f02952d0
title: Étape 19
challengeType: 0
dashedName: step-19
---

# --description--

Donne à chaque élément `th` un élément `span` avec la classe définie sur `sr-only` et le texte suivant, dans l'ordre : `2019`, `2020` et `2021`.

# --hints--

Chacun de tes éléments `th` doit avoir un élément `span`.

```js
const ths = [
  ...document.querySelectorAll('table')?.[1]?.querySelectorAll('th')
];
ths?.forEach(th => {
  assert.lengthOf(th?.children, 1);
  assert.strictEqual(th?.children?.[0]?.localName, 'span');
});
```

Chacun de tes nouveaux éléments `span` doit avoir l'attribut `class` défini sur `sr-only`.

```js
const ths = [
  ...document.querySelectorAll('table')?.[1]?.querySelectorAll('th')
];
ths?.forEach(th => {
  assert.isTrue(th?.children?.[0]?.classList?.contains('sr-only'));
});
```

Ton premier élément `span` doit avoir le texte `2019`.

```js
const table = document.querySelectorAll('table')?.[1];
const th = table.querySelectorAll('th')?.[0];
assert.strictEqual(th.children?.[0]?.textContent.trim(), '2019');
```

Ton deuxième élément `span` doit avoir le texte `2020`.

```js
const table = document.querySelectorAll('table')?.[1];
const th = table.querySelectorAll('th')?.[1];
assert.strictEqual(th?.children?.[0]?.textContent.trim(), '2020');
```

Ton troisième élément `span` doit avoir le texte `2021`.

```js
const table = document.querySelectorAll('table')?.[1];
const th = table.querySelectorAll('th')?.[2];
assert.strictEqual(th?.children?.[0]?.textContent.trim(), '2021');
```

Ton élément `td` doit être vide.

```js
const table = document.querySelectorAll('table')?.[1];
const td = table.querySelectorAll('td')?.[0];
assert.strictEqual(td?.textContent.trim(), '');
assert.lengthOf(td?.children, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>
              Assets
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Cash <span class="description">This is the cash we currently have on hand.</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
              <tr class="data">
                <th>Checking <span class="description">Our primary transactional account.</span></th>
                <td>$54</td>
                <td>$56</td>
                <td class="current">$53</td>
              </tr>
              <tr class="data">
                <th>Savings <span class="description">Funds set aside for emergencies.</span></th>
                <td>$500</td>
                <td>$650</td>
                <td class="current">$728</td>
              </tr>
              <tr class="total">
                <th>Total <span class="sr-only">Assets</span></th>
                <td>$579</td>
                <td>$736</td>
                <td class="current">$809</td>
              </tr>
            </tbody>
          </table>
--fcc-editable-region--
          <table>
            <caption>
              Liabilities
            </caption>
            <thead>
              <tr>
                <td></td>
                <th></th>
                <th></th>
                <th></th>
              </tr>
            </thead>
            <tbody>
            </tbody>
          </table>
--fcc-editable-region--
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
