---
id: 61fd6ab779390f49148773bb
title: Étape 5
challengeType: 0
dashedName: step-5
---

# --description--

Sous ton élément `h1`, crée un élément `div`. Donne-lui un attribut `id` défini sur `years`. Tu veux que cet élément précis soit masqué pour les lecteurs d'écran, donc donne-lui l'attribut `aria-hidden` défini sur `true`.

# --hints--

Tu dois créer un nouvel élément `div` après ton élément `h1`.

```js
assert.strictEqual(
  document.querySelector('h1')?.nextElementSibling?.localName,
  'div'
);
```

Ton nouvel élément `div` doit avoir l'attribut `id` défini sur `years`.

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('id'), 'years');
```

Ton nouvel élément `div` doit avoir l'attribut `aria-hidden` défini sur `true`.

```js
assert.strictEqual(
  document.querySelector('div')?.getAttribute('aria-hidden'),
  'true'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
