---
id: 61fd5a93fd62bb35968adeab
title: Étape 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Commence ton projet en donnant à ton élément `head` un élément `link` pour ta feuille de style.

# --hints--

Ton code doit avoir un élément `link`.

```js
assert.match(code, /<link/);
```

Ton élément `link` doit être dans ton élément `head`.

```js
assert.match(code, /<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i);
```

Ton élément `link` doit avoir un attribut `rel` avec la valeur `stylesheet`.

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/);
```

Ton élément `link` doit avoir un attribut `href` avec la valeur `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
