---
id: 61fd66c687e610436494c6f1
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

Dans ton élément `section`, ajoute un élément `h1` avec un élément `span` imbriqué.

# --hints--

Ton élément `section` doit avoir un élément `h1`.

```js
assert.strictEqual(
  document.querySelector('section')?.children?.[0]?.localName,
  'h1'
);
```

Ton élément `h1` doit avoir un élément `span`.

```js
assert.strictEqual(
  document.querySelector('h1')?.children?.[0]?.localName,
  'span'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
--fcc-editable-region--
      <section>
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```

```css

```
