---
id: 61fd6cc9475a784b7776233e
title: Étape 7
challengeType: 0
dashedName: step-7
---

# --description--

Sous ton élément `div` existant, ajoute un nouvel élément `div` avec une `class` définie sur `table-wrap`. Ce sera le conteneur de tes tableaux.

À l'intérieur, ajoute trois éléments `table`. Tu utiliseras CSS pour les styliser comme un seul tableau, mais utiliser des tableaux séparés aidera les lecteurs d'écran à comprendre le flux du document.

# --hints--

Tu dois créer un nouvel élément `div`.

```js
assert.lengthOf(document.querySelectorAll('div'), 2);
```

Ton nouvel élément `div` doit avoir la `class` définie sur `table-wrap`.

```js
assert.strictEqual(document.querySelector('.table-wrap')?.localName, 'div');
```

Ton élément `.table-wrap` doit venir après ton `div` existant.

```js
assert.isTrue(
  document.querySelectorAll('div')?.[1]?.classList?.contains('table-wrap')
);
```

Ton élément `.table-wrap` doit avoir trois éléments `table`.

```js
const children = [...(document.querySelector('.table-wrap')?.children ?? [])];
assert.lengthOf(children, 3);
children.forEach(child => assert.strictEqual(child?.localName, 'table'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
