---
id: 61fd67a656743144844941cb
title: Étape 4
challengeType: 0
dashedName: step-4
---

# --description--

Les lecteurs d'écran annoncent les éléments HTML selon le flux du document. À terme, nous voulons que le bilan ait un titre "Balance Sheet" et un sous-titre "AcmeWidgetCorp". Cependant, cet ordre n'a pas de sens s'il est annoncé par un lecteur d'écran.

Donne à ton `span` existant l'attribut `class` défini sur `flex`, et ajoute deux éléments `span` à l'intérieur. Donne au premier le texte `AcmeWidgetCorp`. Donne au second le texte `Balance Sheet`. Tu utiliseras CSS pour inverser l'ordre du texte sur la page, mais l'ordre HTML aura plus de sens pour un lecteur d'écran.

# --hints--

Ton élément `span` existant doit avoir l'attribut `class` défini sur `flex`.

```js
assert.isTrue(
  document.querySelector('h1')?.children?.[0]?.classList?.contains('flex')
);
```

Ton élément `span` existant doit contenir deux nouveaux éléments `span`.

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.localName,
  'span'
);
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.localName,
  'span'
);
```

Tes nouveaux éléments `span` ne doivent pas avoir d'attribut `class`.

```js
assert.isEmpty(document.querySelector('.flex')?.children?.[0]?.classList);
assert.isEmpty(document.querySelector('.flex')?.children?.[1]?.classList);
```

Ton premier nouvel élément `span` doit avoir le texte `AcmeWidgetCorp`.

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.textContent.trim(),
  'AcmeWidgetCorp'
);
```

Ton deuxième nouvel élément `span` doit avoir le texte `Balance Sheet`.

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.textContent.trim(),
  'Balance Sheet'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
--fcc-editable-region--
          <span>
          </span>
--fcc-editable-region--
        </h1>
      </section>
    </main>
  </body>
</html>
```

```css

```
