---
id: 61fd75ea7f663457612dba02
title: Étape 11
challengeType: 0
dashedName: step-11
---

# --description--

Dans chacun de tes nouveaux éléments `th`, imbrique un élément `span` avec la `class` définie sur `sr-only year`. Donne-leur le texte suivant (dans l'ordre) : `2019`, `2020` et `2021`.

Donne à ton troisième élément `th` l'attribut `class` défini sur `current`.

Laisse l'élément `td` vide. Cet élément existe seulement pour garantir que ton tableau a une mise en page à quatre colonnes et pour associer les en-têtes aux bonnes colonnes.

# --hints--

Chacun de tes éléments `th` doit avoir un élément `span`.

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert.lengthOf(th?.children, 1);
  assert.strictEqual(th?.children?.[0]?.localName, 'span');
});
```

Chacun de tes nouveaux éléments `span` doit avoir l'attribut `class` défini sur `sr-only year`.

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert.isTrue(th?.children?.[0]?.classList?.contains('sr-only'));
  assert.isTrue(th?.children?.[0]?.classList?.contains('year'));
});
```

Ton premier élément `span` doit avoir le texte `2019`.

```js
assert.strictEqual(
  document.querySelectorAll('th')?.[0]?.children?.[0]?.textContent.trim(),
  '2019'
);
```

Ton deuxième élément `span` doit avoir le texte `2020`.

```js
assert.strictEqual(
  document.querySelectorAll('th')?.[1]?.children?.[0]?.textContent.trim(),
  '2020'
);
```

Ton troisième élément `span` doit avoir le texte `2021`.

```js
assert.strictEqual(
  document.querySelectorAll('th')?.[2]?.children?.[0]?.textContent.trim(),
  '2021'
);
```

Ton troisième élément `th` doit avoir la `class` définie sur `current`.

```js
assert.isTrue(
  document
    .querySelector('table')
    ?.querySelectorAll('th')?.[2]
    ?.classList?.contains('current')
);
```

Ton élément `td` doit être vide.

```js
const table = document.querySelector('table');
const tableData = table.querySelectorAll('td')?.[0];
assert.strictEqual(tableData?.textContent.trim(), '');
assert.isEmpty(tableData?.children);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
--fcc-editable-region--
            <thead>
              <tr>
                <td></td>
                <th></th>
                <th></th>
                <th></th>
              </tr>
            </thead>
--fcc-editable-region--
            <tbody></tbody>
          </table>
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
