---
id: 61fd6343fb42013d99bcd7f3
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

Dans ton élément `body`, imbrique un élément `section` dans un élément `main`.

# --hints--

Le premier élément à l'intérieur de l'élément `body` doit être ton élément `main`.

```js
assert.strictEqual(
  document.querySelector('body')?.children?.[0]?.localName,
  'main'
);
```

Tu dois avoir un élément `section` dans ton élément `main`.

```js
assert.strictEqual(
  document.querySelector('main')?.children?.[0]?.localName,
  'section'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```

```css

```
