---
id: 6720b7f3155a9278d6b6a8b8
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

Il est maintenant temps de styliser ta carte de vœux.

Ajoute un sélecteur pour l'élément `body`, puis :

- change la `font-family` pour utiliser `Arial` suivie de la famille générique `sans-serif`,
- ajoute un padding de `40px` sur tous les côtés,
- définis la propriété `text-align` sur `center`.

# --hints--

Tu devrais avoir un sélecteur `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("body"));
```

Tu devrais définir `font-family` sur `"Arial", sans-serif` dans le sélecteur `body`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("font-family"), ["'Arial', sans-serif", "Arial, sans-serif", '"Arial", sans-serif']);
```

Tu devrais définir `padding` sur `40px` dans le sélecteur `body`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("padding"), "40px");
```

Tu devrais définir `text-align` sur `center` dans le sélecteur `body`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("text-align"), "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
	</div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
