---
id: 6720b1f2680ec706ed534acf
title: Étape 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Dans cet atelier, tu vas t'entraîner à travailler avec différentes pseudo-classes et pseudo-éléments en concevant une carte de vœux. Le code HTML de base t'a déjà été fourni.

Commence l'atelier en liant le fichier `styles.css`.

# --hints--

Ton élément `link` devrait être dans ton élément `head`.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

Ton élément `link` devrait avoir un attribut `rel` avec la valeur `stylesheet`.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

Ton élément `link` devrait avoir un attribut `href` avec la valeur `styles.css`.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.match(linkHrefValue, /^(\.\/)?styles\.css$/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
