---
id: 6720b3191bbafd269bae0ae8
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

Crée un élément `div` qui a un `id` de `greeting-card` et une `class` de `card`.

Dans l'élément `div`, ajoute un `h1` avec le texte `Happy Birthday!`. Ajoute ensuite un élément paragraphe avec une `class` appelée `message` et le texte `Wishing you all the happiness and joy on your special day!`.

# --hints--

Tu devrais avoir un élément `div`.

```js
assert.exists(document.querySelector("div"));
```

L'élément `div` devrait avoir un `id` de `greeting-card`.

```js
assert.strictEqual(document.querySelector("div").id, "greeting-card");
```

L'élément `div` devrait avoir une `class` de `card`.

```js
assert.strictEqual(document.querySelector("#greeting-card").className, "card");
```

Tu devrais avoir un élément `h1` à l'intérieur de l'élément `div`.

```js
assert.exists(document.querySelector("div > h1"));
```

L'élément `h1` devrait contenir le texte `Happy Birthday!`.

```js
assert.strictEqual(document.querySelector("div > h1").textContent.trim(), "Happy Birthday!");
```

Tu devrais avoir un élément `p` à l'intérieur du `div`.

```js
assert.exists(document.querySelector("div > p"));
```

L'élément `p` devrait avoir une `class` de `message`.

```js
assert.exists(document.querySelector("div > p.message"));
```

L'élément `p` devrait contenir le texte `Wishing you all the happiness and joy on your special day!`.

```js
assert.strictEqual(document.querySelector("div > p.message").textContent.trim(), "Wishing you all the happiness and joy on your special day!");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```

```css

```
