---
id: 672b4434dd6d6a1edc405419
title: Étape 8
challengeType: 0
dashedName: step-8
---

# --description--

Ajoute deux éléments `a` à l'intérieur de l'élément `.card-links`.

Le premier devrait avoir le texte `Send Card`, une `class` de `send-link` et un `href` de `#send`.

Le second devrait avoir le texte `Share on Social Media`, une `class` de `share-link` et un `href` de `#share`.

# --hints--

Le premier élément `a` à l'intérieur de `.card-links` devrait avoir le texte `Send Card`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[0].innerText, "Send Card");
```

Le premier élément `a` devrait avoir une `class` avec la valeur `send-link`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[0].className, "send-link");
```

Le premier élément `a` devrait avoir un `href` avec la valeur `#send`.

```js
assert.match(document.querySelectorAll(".card-links > a")[0].href, /#send$/);
```

Le second élément `a` à l'intérieur de `.card-links` devrait avoir le texte `Share on Social Media`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[1].innerText, "Share on Social Media");
```

Le second élément `a` devrait avoir une `class` avec la valeur `share-link`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[1].className, "share-link");
```

Le second élément `a` devrait avoir un `href` avec la valeur `#share`.

```js
assert.match(document.querySelectorAll(".card-links > a")[1].href, /#share$/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
--fcc-editable-region--

--fcc-editable-region--
      </div>
	</div>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
}

```
