---
id: 6733352ce2a5642827b4f7f0
title: Étape 18
challengeType: 0
dashedName: step-18
---

# --description--

Ajoute un sélecteur `.card-links` et définis la propriété `margin-top` sur `20px`.

Tu peux ajouter `display: flex` pour définir un élément afin qu'il utilise flexbox. Tu en apprendras plus sur flexbox plus tard dans le cours ; considère ceci comme un petit aperçu.

Pour espacer les deux liens afin qu'ils aient le même espace autour d'eux, ajoute une propriété `display` définie sur `flex` et une propriété `justify-content` définie sur `space-around`.

# --hints--

Tu devrais avoir un sélecteur `.card-links`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(".card-links"));
```

Le sélecteur `.card-links` devrait avoir une propriété `margin-top` définie sur `20px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card-links")?.getPropertyValue("margin-top"), "20px");
```

Le sélecteur `.card-links` devrait avoir une propriété `display` définie sur `flex`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card-links")?.getPropertyValue("display"), "flex");
```

Le sélecteur `.card-links` devrait avoir une propriété `justify-content` définie sur `space-around`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card-links")?.getPropertyValue("justify-content"), "space-around");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
	</div>
    <section id="send">
      <h2>Sending your card...</h2>
      <p>Card successfully sent to your recipient!</p>
    </section>
    <section id="share">
      <h2>Sharing your card...</h2>
      <p>Your card was shared on social media!</p>
    </section>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
  transition: transform 0.3s, background-color 0.3s ease
}

.card:hover {
  background-color: khaki;
  transform: scale(1.1);
}

h1::before {
  content: "🥳 ";
}

h1::after {
  content: " 🥳";
}

.message {
  font-size: 1.2em;
  margin-bottom: 20px;
}

--fcc-editable-region--

--fcc-editable-region--
```
