---
id: 67333082ab35e414eb4ea70d
title: Étape 14
challengeType: 0
dashedName: step-14
---

# --description--

Quand les éléments `a` sont survolés, la couleur de l'arrière-plan effectue une transition vers une autre couleur. Tu peux contrôler la façon dont cette transition se produit avec la propriété `transition` :

```css
a {
  transition: color 1s linear;
}
```

Les valeurs acceptées par la propriété `transition` sont, dans l'ordre, la propriété à laquelle la transition s'applique, la durée de la transition, puis le timing.

S'il y a plusieurs propriétés avec une transition, tu peux écrire les valeurs de chacune en les séparant par une virgule :

```css
p {
  transition: property1 0.1s, property2 0.6s linear;
}
```

Si une valeur est omise, comme le timing pour la première propriété, une valeur par défaut est appliquée.

Ajoute au sélecteur `.card` `transition: transform 0.3s, background-color 0.3s ease`.

Essaie-le : la transition au survol est terminée.

# --hints--

Le sélecteur `.card` devrait avoir une propriété `transition` définie sur `transform 0.3s, background-color 0.3s ease`.

```js
assert.oneOf(
  new __helpers.CSSHelp(document).getStyle(".card")?.getPropertyValue("transition"),
  [
    "transform 0.3s ease 0s, background-color 0.3s ease 0s", // this is what comes out of `getPropertyValue`
    "transform 0.3s, background-color 0.3s" // different environment got this instead
  ]
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
	</div>
    <section id="send">
      <h2>Sending your card...</h2>
      <p>Card successfully sent to your recipient!</p>
    </section>
    <section id="share">
      <h2>Sharing your card...</h2>
      <p>Your card was shared on social media!</p>
    </section>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
--fcc-editable-region--

--fcc-editable-region--
}

.card:hover {
  background-color: khaki;
  transform: scale(1.1);
}

```
