---
id: 68b94f23f7b3ee14078c89c5
title: Étape 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Dans cet atelier, tu vas utiliser l'élément `iframe` pour afficher une vidéo. Le code de base HTML t'est déjà fourni.

Commence par créer un élément `h1` avec le texte `iframe Video Display`.

# --hints--

Tu dois avoir un élément `h1`.

```js
assert.exists(document.querySelector('h1'))
```

Ton élément `h1` doit avoir le texte `iframe Video Display`.

```js
assert.equal(document.querySelector('h1')?.textContent.trim(), 'iframe Video Display')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Afficher des vidéos dans une iframe</title>
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
