---
id: 68ba9c4f1688914d72876094
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

Dans la première leçon sur l'élément `iframe`, tu as appris que c'est un élément remplacé, comme `img`. Cela signifie qu'il peut aussi recevoir les propriétés `width` et `height` pour déterminer sa largeur et sa hauteur.

Donne à ton élément `iframe` une largeur (`width`) de `560` et une hauteur (`height`) de `315`.

# --hints--

Ton élément `iframe` doit avoir un attribut `width` défini sur `560`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('width'), '560')
```

Ton élément `iframe` doit avoir un attribut `height` défini sur `315`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('height'), '315')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Afficher des vidéos dans une iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>

    <iframe
--fcc-editable-region--

--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
