---
id: 68ba9c4f1688914d72876093
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

Maintenant, crée un élément `iframe`. Ne mets encore rien dedans.

# --hints--

Tu dois avoir un élément `iframe`.

```js
assert.exists(document.querySelector('iframe'))
```

Ton élément `iframe` doit être vide.

```js
assert.equal(document.querySelector('iframe')?.innerHTML.trim(), '')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Afficher des vidéos dans une iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
