---
id: 68ba9c4f1688914d72876096
title: Étape 5
challengeType: 0
dashedName: step-5
---

# --description--

L'un des attributs est `allow`. C'est comme une liste de permissions qui indique au navigateur quelles fonctionnalités l'`iframe` a le droit d'utiliser.

Voici un élément `iframe` avec l'attribut `allow` :

```html
<iframe
  allow="accelerometer autoplay clipboard-write encrypted-media gyroscope picture-in-picture web-share"
></iframe>
```

Ajoute l'attribut `allow` avec les valeurs `accelerometer`, `autoplay` et `clipboard-write`.

`accelerometer` permet à l'`iframe` d'utiliser les capteurs de mouvement pour détecter des choses comme l'inclinaison et la rotation de l'appareil. `autoplay` permet à la vidéo de démarrer automatiquement, et `clipboard-write` permet à l'iframe d'écrire des données dans le presse-papiers de l'utilisateur.

# --hints--

Ton élément `iframe` doit avoir un attribut `allow` défini sur `accelerometer`.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /accelerometer/)
```

L'attribut `allow` de ton élément `iframe` doit contenir `autoplay` parmi ses valeurs.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /autoplay/)
```

L'attribut `allow` de ton élément `iframe` doit contenir `clipboard-write` parmi ses valeurs.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /clipboard-write/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Afficher des vidéos dans une iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
--fcc-editable-region--

--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
