---
id: 68ba9c4f1688914d72876098
title: Étape 7
challengeType: 0
dashedName: step-7
---

# --description--

Le prochain attribut à ajouter est `referrerpolicy`. C'est la règle qui détermine la quantité de détails partagée quand ta page se connecte à une autre page.

Ajoute l'attribut `referrerpolicy` et définis-le sur `strict-origin-when-cross-origin`. Cela partage l'adresse complète sur le même site, seulement le nom du site sur les autres sites, et rien sur les sites non sécurisés.

# --hints--

Ton élément `iframe` doit avoir un attribut `referrerpolicy` défini sur `strict-origin-when-cross-origin`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('referrerpolicy'), 'strict-origin-when-cross-origin')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Afficher des vidéos dans une iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
      allow="accelerometer autoplay clipboard-write encrypted-media gyroscope web-share"
--fcc-editable-region--

--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
