---
id: 68ba9c4f1688914d72876095
title: Étape 4
challengeType: 0
dashedName: step-4
---

# --description--

L'élément `iframe` accepte aussi un attribut `src` avec une valeur qui indique l'URL ou le chemin de la ressource à afficher.

Ajoute un attribut `src` avec la valeur `https://www.youtube.com/embed/I0_951_MPE0` à ton élément `iframe`.

À ce stade, tu devrais voir la vidéo s'afficher sur la page, mais il reste quelques attributs à ajouter.

# --hints--

Ton élément `iframe` doit avoir un attribut `src` défini sur `https://www.youtube.com/embed/I0_951_MPE0`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('src'), 'https://www.youtube.com/embed/I0_951_MPE0')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Afficher des vidéos dans une iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
--fcc-editable-region--

--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
