---
id: 68ba9c4f1688914d72876097
title: Étape 6
challengeType: 0
dashedName: step-6
---

# --description--

Ajoute `encrypted-media`, `gyroscope` et `web-share` aux valeurs existantes de l'attribut `allow`.

Ces trois valeurs autorisent l'utilisation des extensions de médias chiffrés pour protéger la vidéo, donnent accès aux capteurs de mouvement et d'orientation de l'appareil, et permettent de partager le contenu de l'iframe avec les boîtes de dialogue de partage natives de l'appareil.

# --hints--

L'attribut `allow` de ton élément `iframe` doit contenir `encrypted-media` parmi ses valeurs.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /encrypted-media/)
```

L'attribut `allow` de ton élément `iframe` doit contenir `gyroscope` parmi ses valeurs.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /gyroscope/)
```

L'attribut `allow` de ton élément `iframe` doit contenir `web-share` parmi ses valeurs.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /web-share/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Afficher des vidéos dans une iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
--fcc-editable-region--
      allow="accelerometer autoplay clipboard-write"
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
