---
id: 5cd9a70215d3c4e65518328f
title: Construire un compte à rebours
challengeType: 26
dashedName: build-a-countdown
---

# --description--

Dans ce lab, tu vas construire une fonction `countdown` qui prend un nombre de départ et génère une liste de compte à rebours.

Quand tu appelles `countdown(n)`, elle doit renvoyer un tableau qui commence à `n` et compte à rebours de `1` jusqu'à atteindre `1`.

Par exemple :

- `countdown(5)` devrait renvoyer `[5, 4, 3, 2, 1]`
- `countdown(1)` devrait renvoyer `[1]`

Si le nombre de départ est inférieur à `1`, il n'y a rien à compter à rebours, donc renvoie un tableau vide :

- `countdown(0)` devrait renvoyer `[]`
- `countdown(-3)` devrait renvoyer `[]`

Pour terminer le lab, tu dois construire le résultat avec la récursion et tu ne dois pas utiliser de boucles (`for`, `while`), `Array.from()`, ni de méthodes d'itération de tableau (`forEach`, `map`, `filter`, `reduce`). Chaque appel doit produire son propre tableau de résultat (n'utilise pas de variables globales pour stocker le compte à rebours).

**Objectif :** Remplir les user stories ci-dessous et faire passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu dois créer une fonction nommée `countdown`.
2. La fonction `countdown` doit prendre un seul argument `n` (un nombre).
3. Si `n` est inférieur à `1`, `countdown` doit renvoyer un tableau vide.
4. Sinon, `countdown` doit renvoyer un tableau contenant les entiers de `n` jusqu'à `1` dans l'ordre décroissant.
5. Tu dois utiliser la récursion (la fonction doit s'appeler elle-même) et tu ne dois utiliser aucune boucle (`for`, `while`, `Array.from()`, ni de méthodes d'ordre supérieur comme `forEach`, `map`, `filter`, `reduce`).
6. Chaque appel récursif doit utiliser une valeur plus petite que l'appel précédent (par exemple, `n - 1`) pour que la fonction atteigne son cas de base.
7. La solution ne doit pas utiliser de variables globales pour stocker ou mettre en cache le tableau de résultat.
8. Appeler `countdown` plusieurs fois avec des entrées différentes doit toujours renvoyer le résultat correct et indépendant.

# --hints--

Tu dois avoir une fonction nommée `countdown`.

```js
assert.isFunction(countdown);
```

`countdown(-1)` devrait renvoyer un tableau vide.

```js
assert.isEmpty(countdown(-1));
```

`countdown(10)` devrait renvoyer `[10, 9, 8, 7, 6, 5, 4, 3, 2, 1]`

```js
assert.deepStrictEqual(countdown(10), [10, 9, 8, 7, 6, 5, 4, 3, 2, 1]);
```

`countdown(5)` devrait renvoyer `[5, 4, 3, 2, 1]`

```js
assert.deepStrictEqual(countdown(5), [5, 4, 3, 2, 1]);
```

Ton code ne doit pas s'appuyer sur des boucles d'aucune sorte (`for`, `while`, `Array.from()` ou des fonctions d'ordre supérieur telles que `forEach`, `map`, `filter` et `reduce`).

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce|Array\.from/g)
);
```

Tu dois utiliser la récursion pour résoudre ce problème.

```js
assert.isAtLeast(
  __helpers.removeJSComments(countdown.toString()).match(/countdown\s*\(.+\)/g)?.length,
  2
);
```

Les variables globales ne doivent pas être utilisées pour mettre en cache le tableau.

```js
countdown(1)
assert.deepStrictEqual(countdown(5), [5, 4, 3, 2, 1]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function countdown(n) {
  if (n < 1) {
    return [];
  }

  return [n, ...countdown(n - 1)];
}
```
