---
id: 5dc24614f86c76b9248c6ebd
title: Étape 10
challengeType: 0
dashedName: step-10
---

# --description--

Tu peux créer un **lien** vers une autre page grâce à l'élément ancre, écrit `a` (pour *anchor*).

Voici un exemple de lien vers `https://www.freecodecamp.org` :

```html
<a href="https://www.freecodecamp.org"></a>
```

L'attribut `href` indique l'**URL de destination** du lien.

Ajoute un élément `a` **après** le paragraphe, qui pointe vers `https://freecatphotoapp.com`. À ce stade, le lien n'apparaîtra pas encore dans l'aperçu parce qu'il n'a pas de texte cliquable.

# --hints--

Ton élément ancre (`a`) doit avoir une balise ouvrante.

```js
assert(document.querySelector('a'));
```

Ton élément `a` doit avoir une balise fermante. Une balise fermante porte un `/` juste après le `<`.

```js
assert(code.match(/<\/a\>/));
```

Ton élément `a` doit se trouver **en dessous** de l'élément `p`.

```js
const collection = [...document.querySelectorAll('a, p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('P') < collection.indexOf('A'));
```

Ton élément `a` n'a pas d'attribut `href`. Vérifie qu'il y a un espace après le nom de l'élément.

```js
assert(document.querySelector('a').hasAttribute('href'));
```

Ton élément `a` doit pointer vers `https://freecatphotoapp.com`. Vérifie l'URL.

```js
assert(
  document.querySelector('a').getAttribute('href') ===
    'https://freecatphotoapp.com'
);
```

Même si l'`href` a la bonne valeur, il faut **toujours** entourer la valeur d'un attribut par des guillemets.

```js
assert(
  !/\<a\s+href\s*=\s*https:\/\/www.freecodecamp.org\/cat-photos/.test(code)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>
      
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
