---
id: 5dc23991f86c76b9248c6eb8
title: Étape 6
challengeType: 0
dashedName: step-6
---

# --description--

À l'étape précédente, tu as placé les éléments `h1`, `h2`, le commentaire et `p` à l'intérieur de l'élément `main`. On dit que les éléments sont **imbriqués** (*nested*).

Pour rendre le code lisible, on **décale** chaque élément imbriqué de deux espaces vers la droite par rapport à son parent. Ce décalage s'appelle l'**indentation**.

Exemple d'imbrication correctement indentée :

```html
<main>
  <h1>Contenu principal de la page</h1>
  <p>Encore un peu de contenu important...</p>
</main>
```

Dans le code ci-dessous, les éléments `h1`, `h2` et le commentaire sont décalés de deux espaces de plus que `<main>`. Utilise la barre d'espace pour ajouter deux espaces devant l'élément `p` afin qu'il soit indenté comme les autres.

# --hints--

Ton code doit contenir un élément `h2` avec le texte `Cat Photos`. Tu l'as peut-être supprimé sans le vouloir.

```js
assert(
  document.querySelector('h2') &&
    code.match(/<\/h2\>/) &&
    document.querySelector('h2').innerText.toLowerCase() === 'cat photos'
);
```

Tu ne dois pas ajouter d'éléments `ul` ou `li` (ils ne sont là que dans l'exemple précédent, pas dans ton code).

```js
assert(
  !document.querySelector('ul') && !document.querySelector('li')
);
```

Tu ne dois pas changer l'indentation de la ligne du `h2`. Sa balise ouvrante doit commencer 6 espaces après le début de la ligne. Tu peux redémarrer l'étape pour retrouver l'indentation initiale.

```js
assert(code.toLowerCase().match(/<\/h1\>\s*\n\s{6}<h2>/));
```

Ton code doit toujours contenir le commentaire. Ne le supprime pas.

```js
assert(code.match(/<!--.*-->/));
```

Le texte du commentaire doit rester `TODO: Add link to cat photos`. Ne change ni le texte ni les espaces.

```js
assert(code.match(/<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\.?\s*-->/i));
```

Tu ne dois pas changer l'indentation de la ligne du commentaire. Sa balise ouvrante doit commencer 6 espaces après le début de la ligne.

```js
assert(
  code
    .toLowerCase()
    .match(/<\/h2>\s*\n\s{6}<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/)
);
```

Ton code doit toujours contenir l'élément `p`. Ne le supprime pas.

```js
assert(document.querySelector('p'));
```

Le texte de l'élément `p` doit rester `Everyone loves cute cats online!`. Ne modifie ni le texte ni la ponctuation.

```js
assert.match(
  document
    .querySelector('p')
    .innerText.toLowerCase(),
    /everyone\s+loves\s+cute\s+cats\s+online!$/
);
```

La balise ouvrante du `p` doit avoir la même indentation que celles de `h2` et du commentaire : 6 espaces depuis le début de la ligne.

```js
assert(code.toLowerCase().match(/-->\s*\n\s{6}<p>/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
    <p>Everyone loves cute cats online!</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```
