---
id: 5dc23f9bf86c76b9248c6eba
title: Étape 7
challengeType: 0
dashedName: step-7
---

# --description--

Pour ajouter une image à ta page web, on utilise l'élément `img`. Particularité : l'élément `img` n'a **qu'une balise ouvrante** et **pas de balise fermante**. On appelle ce type d'élément un **élément void** (vide).

Ajoute un élément `img` sous l'élément `p`. À ce stade, aucune image n'apparaîtra encore dans le navigateur : ce sera pour l'étape suivante.

# --hints--

Ton élément `img` doit avoir une balise ouvrante. Une balise ouvrante a la forme : `<nomDeLElement>`.

```js
assert(document.querySelector('img'));
```

Ton élément `img` ne doit **pas** avoir de balise fermante `</img>`.

```js
assert(!code.match(/<\/img\>/));
```

Tu ne dois avoir qu'un seul élément `img`. Supprime ceux en trop.

```js
assert(document.querySelectorAll('img').length === 1);
```

Ton élément `img` doit se trouver **sous** l'élément `p`.

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('P') < collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
