---
id: 5dfa37b9eacea3f48c6300b0
title: Étape 22
challengeType: 0
dashedName: step-22
---

# --description--

Pour créer une **liste non ordonnée** (à puces), on utilise l'élément `ul` (pour *unordered list*).

Après l'élément `h3` qui contient le texte `Things cats love:`, ajoute un élément `ul`. À ce stade, rien ne s'affichera dans l'aperçu (la liste est encore vide).

# --hints--

Ton élément `ul` doit avoir une balise ouvrante.

```js
assert(document.querySelector('ul'));
```

Ton élément `ul` doit avoir une balise fermante.

```js
assert(code.match(/<\/ul>/));
```

L'élément `ul` doit se trouver au-dessus de la balise fermante du deuxième `section`.

```js
const secondSectionLastElemNode = document.querySelectorAll('main > section')?.[1]?.lastElementChild;
assert.equal(secondSectionLastElemNode?.nodeName , 'UL');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
--fcc-editable-region--
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
