---
id: 5dfb6a35eacea3f48c6300b4
title: Étape 26
challengeType: 0
dashedName: step-26
---

# --description--

L'élément `figcaption` (pour *figure caption*) sert à ajouter une **légende** à l'image contenue dans un élément `figure`.

Voici un exemple de `figcaption` avec la légende `A cute cat` :

```html
<figure>
  <img src="image.jpg" alt="A description of the image">
  <figcaption>A cute cat</figcaption>
</figure>
```

Après l'image imbriquée dans le `figure`, ajoute un élément `figcaption` avec ce texte :

`Cats love lasagna.`

# --hints--

Ton élément `figcaption` doit avoir une balise ouvrante.

```js
assert(document.querySelector('figcaption'));
```

Ton élément `figcaption` doit avoir une balise fermante.

```js
assert(code.match(/<\/figcaption\>/));
```

L'élément `figcaption` doit être imbriqué dans l'élément `figure`.

```js
assert(
  document.querySelector('figure > figcaption') &&
    document.querySelector('figure > figcaption')
);
```

L'élément `img` de la lasagne doit être imbriqué dans le `figure`.

```js
assert(
  document.querySelector('figure > img') &&
    document.querySelector('figure > img').getAttribute('src').toLowerCase() ===
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

L'élément `figcaption` imbriqué dans le `figure` doit être **en dessous** de l'élément `img`.

```js
assert(
  document.querySelector('figcaption').previousElementSibling.nodeName === 'IMG'
);
```

Le texte de ton `figcaption` doit être `Cats love lasagna.`.

```js
assert(
  document.querySelector('figcaption').innerText.trim().match(/^Cats love lasagna\.?$/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          
        </figure>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
