---
id: 5ef9b03c81a63668521804e9
title: Étape 66
challengeType: 0
dashedName: step-66
---

# --description--

Transforme le texte `freeCodeCamp.org` existant en **lien** en l'entourant d'un `a`. L'attribut `href` doit valoir `https://www.freecodecamp.org`.

# --hints--

Ton `a` doit être imbriqué dans le `p` du `footer` (avec balises ouvrante et fermante).

```js
assert(document.querySelectorAll('footer > p > a').length);
```

Ton `a` doit avoir une balise fermante.

```js
const aElemClosingTags = code.match(/<\/a\>/g);
assert.lengthOf(aElemClosingTags, 4);
```

L'attribut `href` doit valoir `https://www.freecodecamp.org`.

```js
const nestedAnchor = document.querySelector('footer > p > a');
assert.strictEqual(nestedAnchor?.getAttribute('href')?.toLowerCase(), 'https://www.freecodecamp.org');
```

Le texte du lien doit être `freeCodeCamp.org`.

```js
const nestedAnchor = document.querySelector('footer > p > a');
assert.strictEqual(
  nestedAnchor?.innerText?.toLowerCase()?.replace(/\s+/g, ' '),
  'freecodecamp.org'
);
```

Après avoir imbriqué le lien, le seul texte visible du `p` dans le navigateur doit être `No Copyright - freeCodeCamp.org`.

```js
const pText = document.querySelector('footer > p')?.innerText?.toLowerCase()?.replace(/\s+/g, ' ');
assert.match(pText, /^no copyright - freecodecamp.org$/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor" checked> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
            <input id="loving" type="checkbox" name="personality" value="loving" checked> <label for="loving">Loving</label>
            <input id="lazy" type="checkbox" name="personality" value="lazy"> <label for="lazy">Lazy</label>
            <input id="energetic" type="checkbox" name="personality" value="energetic"> <label for="energetic">Energetic</label>
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
    <footer>
--fcc-editable-region--
      <p>No Copyright - freeCodeCamp.org</p>
--fcc-editable-region--
    </footer>
  </body>
</html>
```
