---
id: 5ef9b03c81a63668521804dc
title: Étape 46
challengeType: 0
dashedName: step-46
---

# --description--

Les **boutons radio** servent quand tu veux qu'une seule réponse soit choisie parmi plusieurs options (par exemple : « Homme / Femme / Autre »).

Exemple :

```html
<input type="radio"> cat
```

Rappelle-toi : un `input` est un élément **void** (pas de balise fermante).

**Avant** le champ texte, ajoute un bouton radio suivi du texte :

`Indoor`

# --hints--

Tu dois créer un élément `input` pour ton bouton radio.

```js
assert(document.querySelectorAll('form > input').length >= 2);
```

Ton `input` doit avoir une balise ouvrante, **pas** de balise fermante.

```js
assert(document.querySelector('form > input') && !code.match(/<\/input\>/g));
```

Tu ne dois ajouter qu'un seul nouvel `input` radio.

```js
assert(document.querySelectorAll('form > input'), 2);
```

Le nouveau `input` doit être **au-dessus** de l'`input type="text"`.

```js
const existingInputElem = document.querySelector('form > input[type="text"]');
assert(
  existingInputElem &&
    existingInputElem.previousElementSibling.nodeName === 'INPUT'
);
```

Ton `input` doit avoir un attribut `type`.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('type'));
```

Ton `input` ne doit avoir qu'**un seul** attribut `type`.

```js
assert(document.querySelector('input')
  ?.getAttributeNames()
  .filter(attr => attr === 'type')
  .length === 1
);
```

L'attribut `type` doit valoir `radio`.

```js
assert.match(
  document.querySelector('input')?.getAttribute('type'),
  /^radio$/i
);
```

Pense à toujours entourer la valeur d'un attribut par des guillemets.

```js
assert(!/\<\s*input\s+type\s*=\s*radio/i.test(code));
```

Le texte ` Indoor` doit être **après** le bouton radio, pas avant.

```js
const radioInputElem = document.querySelector('input');
assert.notMatch(radioInputElem?.previousSibling?.nodeValue, /Indoor/i);
```

Le texte ` Indoor` doit être directement à droite du bouton radio.

```js
const radioInputElem = document.querySelector('input');
assert.match(radioInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '), /\s*Indoor/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
 
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
