---
id: 5ef9b03c81a63668521804d6
title: Étape 37
challengeType: 0
dashedName: step-37
---

# --description--

Tu vas maintenant ajouter un **formulaire web** pour récolter des informations auprès de l'utilisateur.

L'élément `form` sert à collecter des données saisies par l'utilisateur : nom, e-mail, message, etc.

Après le titre `Cat Form`, ajoute un élément `form` (vide pour l'instant).

# --hints--

Ton élément `form` doit avoir une balise ouvrante et une balise fermante.

```js
assert(document.querySelector('form') && code.match(/<\/form>/g));
```

Les balises du `form` doivent être dans le bon ordre (ouvrante avant fermante).

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.indexOf('<form>') < noSpaces.indexOf('</form>'));
```

Le `form` imbriqué dans le dernier `section` doit se trouver **sous** le `h2`.

```js
assert(document.querySelector('form').previousElementSibling.nodeName === 'H2');
```

Le `form` doit être **vide** (pas de contenu entre les balises pour l'instant).

```js
assert.lengthOf(document.querySelector('form')?.innerHTML?.trim(), 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
--fcc-editable-region--
        <h2>Cat Form</h2>
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
