---
id: 5ef9b03c81a63668521804dd
title: Étape 47
challengeType: 0
dashedName: step-47
---

# --description--

L'élément `label` (étiquette) sert à **associer** un texte à un `input`. Très important pour l'**accessibilité** : les lecteurs d'écran annoncent le texte de l'étiquette, et l'utilisateur peut cliquer sur le texte pour sélectionner l'`input`.

Exemple :

```html
<label><input type="radio"> cat</label>
```

Dans cet exemple, cliquer sur le mot `cat` sélectionne aussi le bouton radio.

Imbrique ton bouton radio à l'intérieur d'un élément `label`.

# --hints--

Le bouton radio doit toujours être présent.

```js
assert(document.querySelector('input[type="radio"]'));
```

Le texte ` Indoor` doit toujours se trouver juste à droite du bouton radio.

```js
const radioInputElem = document.querySelector('input');
assert.match(
  radioInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '),
  / Indoor/i
);
```

Ton élément `label` doit avoir une balise ouvrante.

```js
assert(document.querySelector('label'));
```

Ton élément `label` doit avoir une balise fermante.

```js
assert(code.match(/<\/label\>/));
```

Le bouton radio et son texte doivent être entre les balises ouvrante et fermante du `label`.

```js
const labelChildNodes = [...document.querySelector('form > label').childNodes];
assert(
  labelChildNodes.filter((childNode) => childNode.nodeName === 'INPUT').length
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="radio"> Indoor
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
