---
id: 5ef9b03c81a63668521804d7
title: Étape 38
challengeType: 0
dashedName: step-38
---

# --description--

L'attribut `action` indique **où** les données du formulaire doivent être envoyées (l'URL de destination).

Exemple :

```html
<form action="/submit-url"></form>
```

Dans cet exemple, `action="/submit-url"` dit au navigateur d'envoyer les données du formulaire à l'adresse `/submit-url`.

Ajoute un attribut `action` qui vaut `https://freecatphotoapp.com/submit-cat-photo` à ton élément `form`.

# --hints--

Ton `form` doit avoir une balise ouvrante et une balise fermante, dans le bon ordre.

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

Le `form` doit toujours se trouver sous le `h2`.

```js
assert.equal(document.querySelector('form')?.previousElementSibling?.nodeName, 'H2');
```

Le `form` doit toujours être vide pour l'instant.

```js
assert.lengthOf(document.querySelector('form')?.innerHTML?.trim(), 0);
```

Ton `form` doit avoir un attribut `action`. Vérifie qu'il y a un espace entre `form` et `action`.

```js
const form = document.querySelector('form');
assert(form.hasAttribute('action'));
```

L'attribut `action` doit valoir `https://freecatphotoapp.com/submit-cat-photo`.

```js
const form = document.querySelector('form');
assert(
  form
    .getAttribute('action')
    .match(/^https:\/\/freecatphotoapp\.com\/submit-cat-photo$/)
);
```

Pense à toujours entourer la valeur d'un attribut par des guillemets.

```js
assert(
  !/\<form\s+action\s*=\s*https:\/\/freecatphotoapp\.com\/submit-cat-photo/.test(
    code
  )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form>
        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
