---
id: 5ef9b03c81a63668521804d8
title: Étape 39
challengeType: 0
dashedName: step-39
---

# --description--

L'élément `input` (pour *entrée*) permet à l'utilisateur de **saisir** une donnée dans le formulaire (du texte, un nombre, une case à cocher, etc.). Comme l'élément `img`, c'est un **élément void** — il n'a **pas** de balise fermante.

Imbrique un élément `input` dans le `form`.

# --hints--

Ton `form` doit toujours avoir ses deux balises dans le bon ordre.

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

La balise ouvrante du `form` ne doit avoir que l'attribut `action`. Retire ce que tu aurais pu ajouter en plus.

```js
assert([...document.querySelector('form').attributes].length < 2);
```

Tu dois créer un élément `input`.

```js
assert(document.querySelector('input'));
```

L'`input` doit avoir une balise ouvrante mais **pas** de balise fermante.

```js
assert(document.querySelector('input') && !code.match(/<\/input\>/g));
```

L'`input` doit être imbriqué à l'intérieur du `form`.

```js
assert(document.querySelector('form > input'));
```

Le `form` ne doit contenir **que** l'`input` (rien d'autre, pas de texte).

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          
        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
