---
id: 5efb2c990dc218d6c85f89b2
title: Étape 45
challengeType: 0
dashedName: step-45
---

# --description--

Même si tu as placé le bouton **sous** le champ texte, ils apparaissent côte à côte dans l'aperçu. C'est parce que `input` et `button` sont des **éléments en ligne** (*inline*) — ils ne provoquent pas de retour à la ligne.

Le bouton soumet déjà le formulaire par défaut, mais s'appuyer sur ce comportement implicite peut prêter à confusion. Ajoute l'attribut `type` avec la valeur `submit` à ton `button` pour rendre son rôle **explicite**.

# --hints--

Ton élément `button` doit avoir une balise ouvrante.

```js
assert(document.querySelector('button'));
```

Ton élément `button` doit avoir une balise fermante.

```js
assert(code.match(/<\/button\>/));
```

Ton `button` doit avoir un attribut `type`.

```js
assert.isTrue(document.querySelector('button')?.hasAttribute('type'));
```

L'attribut `type` doit valoir `submit`. N'oublie pas les guillemets.

```js
assert.match(
  document.querySelector('button')?.getAttribute('type'),
  /^submit$/i
);
```

Pense à toujours entourer la valeur d'un attribut par des guillemets.

```js
assert(!/\<\s*button\s+type\s*=\s*submit/i.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button>Submit</button>
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```
