---
id: 5efc575c8d6a74d05e68af77
title: Étape 61
challengeType: 0
dashedName: step-61
---

# --description--

Ajoute une **troisième** case à cocher après la précédente avec `id="energetic"`. Le `name` doit être le même que la case précédente.

Ajoute aussi un `label` à droite de cette nouvelle case avec le texte `Energetic`. Pense à l'**associer** à la case via l'attribut `for`.

# --hints--

Tu dois ajouter une nouvelle case à cocher.

```js
assert.lengthOf(document.querySelectorAll('input[type="checkbox"]'),3);
```

La nouvelle case doit avoir `id="energetic"` et `name="personality"`.

```js
const checkboxes = [...document.querySelectorAll('input[type="checkbox"]')];
assert(
  checkboxes.some(
    (checkbox) =>
      checkbox.id === 'energetic' &&
      checkbox.getAttribute('name') === 'personality'
  )
);
```

La nouvelle case doit être **après** la deuxième.

```js
const checkboxes = [...document.querySelectorAll('input[type="checkbox"]')].map(
  (checkbox) => checkbox.id
);
assert(checkboxes.indexOf('lazy') < checkboxes.indexOf('energetic'));
```

À droite de la nouvelle case, il doit y avoir un `label` avec le texte `Energetic`.

```js
const nextElementSibling = document.querySelectorAll('input[type="checkbox"]')?.[2]?.nextElementSibling;
assert.isOk(
  nextElementSibling?.nodeName === 'LABEL' &&
  nextElementSibling?.innerText?.trim().match(/^Energetic$/i)
);
```

Le nouveau `label` doit avoir `for` avec la même valeur que l'`id` de la nouvelle case.

```js
assert.strictEqual(
  document.querySelectorAll('input[type="checkbox"]')?.[2]?.nextElementSibling?.getAttribute('for'),
  'energetic'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
--fcc-editable-region--
            <input id="loving" type="checkbox" name="personality"> <label for="loving">Loving</label>
            <input id="lazy" type="checkbox" name="personality"> <label for="lazy">Lazy</label>
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
