---
id: 5f05a1d8e233dff4a68508d8
title: Étape 49
challengeType: 0
dashedName: step-49
---

# --description--

Crée un **second bouton radio** sous le premier. Imbrique-le dans un nouveau `label` avec le texte `Outdoor`. Donne-lui un attribut `id` qui vaut `outdoor`.

# --hints--

Seul le bouton radio `Indoor` doit avoir `id="indoor"`. Tu peux redémarrer l'étape pour récupérer le code initial si besoin.

```js
assert(document.querySelectorAll('#indoor').length < 2);
```

Tu ne dois rien changer au bouton radio `Indoor` existant.

```js
const indoorInput = document.querySelectorAll('#indoor');
assert(
  indoorInput.length == 1 &&
  indoorInput[0]?.nodeName?.toUpperCase() === 'INPUT' &&
  indoorInput[0]?.type === 'radio' &&
  code.match(/<label\s*>\s*<input [^>]*(id=["']?indoor["']?)[^>]*>\s*Indoor\s*<\/label>/)
);
```

Tu dois ajouter **un seul** nouveau `input` imbriqué dans un nouveau `label`.

```js
assert(document.querySelectorAll('label input').length === 2);
```

Tes balises `<label>` ne doivent pas avoir d'attribut.

```js
assert(code.match(/<label\s*>/g).length === 2);
```

Le nouveau `input` doit avoir `id="outdoor"`.

```js
assert(document.querySelector('label input[id="outdoor"]'));
```

Le nouveau `input` doit avoir `type="radio"`.

```js
assert(document.querySelector('label input[id="outdoor"][type="radio"]'));
```

Tu ne dois ajouter que l'`input` à l'intérieur du `label` (pas d'autres éléments).

```js
assert(document.querySelector('label input[id="outdoor"]:only-child'));
```

Il ne doit pas y avoir de texte **à gauche** du nouvel `input`.

```js
const outdoorBtn = document.querySelector('label input[id="outdoor"]');
assert(
    outdoorBtn?.previousSibling?.nodeName !== '#text' ||
    outdoorBtn?.previousSibling?.nodeValue?.replace(/\s/g, '') === ''
);
```

Le texte du nouveau radio doit être exactement `Outdoor` (avec un O majuscule).

```js
assert(document.querySelector('label > input[id="outdoor"]')?.nextSibling?.nodeValue?.replace(/^\s+|\s+$/g, '') === 'Outdoor');
```

Le nouveau radio doit être **juste après** le label `Indoor`, sans rien entre les deux.

```js
assert(code.match(/<\/label>\s*<label\s*>\s*<input [^>]+>\s*Outdoor/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input id="indoor" type="radio"> Indoor</label>
          
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
