---
id: 5f07be6ef7412fbad0c5626b
title: Étape 18
challengeType: 0
dashedName: step-18
---

# --description--

Avant d'ajouter du nouveau contenu, on va utiliser un élément `section` pour **séparer** le contenu des photos de chats du contenu à venir.

L'élément `section` sert à définir des **sections** dans un document : un chapitre, un en-tête, un pied de page, ou n'importe quelle zone thématique. C'est un élément **sémantique** qui aide au référencement (SEO) et à l'accessibilité.

```html
<section>
  <h2>Titre de la section</h2>
  <p>Contenu de la section…</p>
</section>
```

Prends ton `h2`, tes éléments `p` et ton élément ancre (`a`), et **imbrique-les** dans un élément `section`.

# --hints--

Ton élément `section` doit avoir une balise ouvrante.

```js
assert(document.querySelector('section'));
```

Ton élément `section` doit avoir une balise fermante.

```js
assert(code.match(/<\/section\s*>/i));
```

Tout l'élément `section` doit être placé entre les balises ouvrante et fermante du `main`.

```js
assert(document.querySelector('section').parentNode.nodeName === 'MAIN');
```

Le `h2` existant, les deux `p` et l'élément ancre doivent être à l'intérieur du `section`.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const foundElements = childrenOfSection.filter((child) => {
  return ['H2', 'A', 'P'].includes(child.nodeName);
});
assert.lengthOf(foundElements, 4);
```

L'élément `h1` ne doit **pas** être imbriqué dans le `section`.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const includesH1 = childrenOfSection.some((child) => child.nodeName === 'H1');
assert.isFalse(includesH1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>

      <h2>Cat Photos</h2>
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
      <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

    </main>
--fcc-editable-region--
  </body>
</html>
```
