---
id: 671141d8e32fe934c26fa1be
title: Étape 12
challengeType: 0
dashedName: step-12
---

# --description--

Ajoute les mots `See more ` **avant** l'élément ancre et ` in our gallery` **après** l'élément ancre.

# --hints--

Tu dois utiliser l'élément ancre existant, n'en crée pas un nouveau.

```js
assert.lengthOf(document.querySelectorAll('a'), 1);
```

Le texte à l'intérieur de l'ancre ne doit pas changer : il doit rester `cat photos`.

```js
assert.strictEqual(document.querySelector('a').innerText, "cat photos");
```

Ton code doit contenir `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery`.

```js
assert.match(code, /See more <a href="https:\/\/freecatphotoapp\.com">cat photos<\/a> in our gallery/)
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
