---
id: 671141feba228a35cefba82d
title: Étape 14
challengeType: 0
dashedName: step-14
---

# --description--

Transforme le texte existant `cute cats` en un élément ancre qui pointe vers :

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg`

# --hints--

Il doit y avoir un nouvel élément ancre dans le premier élément `p`.

```js
assert.exists(document.querySelector('p')?.querySelector('a'));
```

Le texte de l'ancre doit être `cute cats`.

```js
assert.strictEqual(document.querySelector('a')?.innerText.trim(), "cute cats");
```

L'attribut `href` du nouvel élément ancre doit être `https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg`.

```js
assert.strictEqual(document.querySelector('a')?.href, "https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg");
```

Il ne doit y avoir qu'**une seule** occurrence de `cute cats` dans ton code.

```js
assert.strictEqual(code.indexOf("cute cats"),code.lastIndexOf("cute cats"));
```

Le texte de l'élément `p` doit toujours être `Everyone loves cute cats online!`.

```js
assert.strictEqual(document.querySelector('p')?.innerText.trim(), "Everyone loves cute cats online!");
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
