---
id: 7cf9b03d81a65668421804c3
title: Étape 41
challengeType: 0
dashedName: step-41
---

# --description--

Pour que les données envoyées par le formulaire puissent être lues à l'arrivée (à l'URL définie par `action`), il faut donner un **nom** à chaque champ via l'attribut `name`. Cette valeur sert d'étiquette à la donnée envoyée.

Exemple d'`input` avec un attribut `name` :

```html
<input type="text" name="name">
```

Ajoute l'attribut `name` avec la valeur `catphotourl` à ton champ texte.

# --hints--

Tu dois toujours avoir ton `input`. Pense aux guillemets autour des valeurs.

```js
assert(document.querySelectorAll('input').length);
```

Le `form` ne doit contenir que l'`input`.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
  document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

Ton `input` doit avoir un attribut `name`. Vérifie qu'il y a un espace entre les attributs.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('name'));
```

L'attribut `name` doit valoir `catphotourl`. N'oublie pas les guillemets.

```js
assert.match(
  document.querySelector('input')?.getAttribute('name'),
  /^catphotourl$/i
);
```

Pense à toujours entourer la valeur d'un attribut par des guillemets.

```js
assert(!/\<\s*input\s+.*\s*=\s*catphotourl/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```
