---
id: 68813fbfc9f56dde3f21da08
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

Quand un utilisateur édite une note et quitte la carte avec Tab ou un clic à l'extérieur, un message d'état devrait s'afficher pour indiquer que la note a bien été enregistrée.

Pour commencer, crée un élément `div` avec un attribut `id` défini à `"status"`.

# --hints--

Tu devrais avoir un deuxième élément `div`.

```js
assert.lengthOf(document.querySelectorAll("div"), 2);
```

Ton deuxième élément `div` devrait avoir un attribut `id`.

```js
assert.isTrue(document.querySelectorAll("div")[1].hasAttribute("id"));
```

Ton deuxième élément `div` devrait avoir un attribut `id` défini à `"status"`.

```js
assert.equal(document.querySelectorAll("div")[1].getAttribute("id"), "status");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Note taking app</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <p class="helper-text">Click or tap on the card to edit your note.</p>

    <div id="note" class="note" contenteditable="true" aria-label="Note editor">
      Many languages have words that carry meanings so specific or culturally rooted that they can't be neatly translated into English.

      One example is the Japanese word "tsundoku", which refers to the habit of acquiring books and letting them pile up unread, something many book lovers can relate to. Another is the Portuguese word "saudade", describing a deep, bittersweet longing for something or someone that is absent. Meanwhile, the French word "Dépaysement" captures the disorienting yet exciting feeling of being in a new place, far from home.

      These unique words remind us that language is more than vocabulary: it's a window into the values, habits, and emotions of the cultures that create it.
    </div>

    --fcc-editable-region--

    --fcc-editable-region--

    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  margin: 2em;
  max-width: 700px;
  background-color: #f5f5f5;
}

.note {
  background-color: #ffffff;
  border: 1px solid #ddd;
  border-radius: 4px;
  box-shadow: 0 2px 5px rgba(0, 0, 0, 0.1);
  padding: 1.5em;
  margin-bottom: 1em;
  line-height: 1.5;
  min-height: 250px;
  font-size: 16px;
  /* This is needed to preserve line breaks in the div */
  white-space: pre-wrap;
}

.note[contenteditable="true"] {
  caret-color: black;
}

.note:hover {
  background-color: #fff;
  box-shadow: 0 0 5px rgba(0,0,0,0.2);
}

.helper-text {
  font-size: 0.9rem;
  color: #666;
  margin-top: 0.5em;
  user-select: none;
  font-style: italic;
}

#status {
  color: #00471b;
  padding: 0 1em;
}
```
