---
id: 68814221fc4752e17b9696bd
title: Étape 8
challengeType: 0
dashedName: step-8
---

# --description--

Quand l'utilisateur clique sur la note, fait des modifications, puis clique à l'extérieur de la note, tu auras besoin d'un moyen d'obtenir cette version nouvellement éditée. Tu peux le faire avec l'événement `blur`. Cet événement se déclenche quand un élément perd le focus. Voici un exemple :

```js
element.addEventListener("blur", () => {
  // do something here
});
```

Attache une méthode `addEventListener` à la variable `noteEl`. Le premier argument devrait être l'événement `"blur"` et le deuxième argument devrait être une fonction fléchée. Dans le corps de cette fonction fléchée, utilise `const` pour créer une nouvelle variable appelée `newContent` et assigne-lui la valeur de `noteEl.innerHTML`.

`innerHTML` est utilisé ici pour aider à préserver les espaces et le formatage de la note.

# --hints--

Tu devrais attacher une méthode `addEventListener` à la variable `noteEl`.

```js
assert.match(code, /noteEl\.addEventListener/);
```

Ta méthode `addEventListener` devrait avoir comme premier argument l'événement `"blur"`.

```js
assert.match(code, /noteEl\.addEventListener\(['"]blur['"],/);
```

Dans le corps de la fonction fléchée, tu devrais utiliser `const` pour créer une nouvelle variable appelée `newContent` et lui assigner la valeur de `noteEl.innerHTML`.

```js
assert.match(code, /noteEl\.addEventListener\(['"]blur['"],\s*\(\)\s*=>\s*{[^}]*const\s+newContent\s*=\s*noteEl\.innerHTML[^}]*}\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Note taking app</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <p class="helper-text">Click or tap on the card to edit your note.</p>

    <div id="note" class="note" contenteditable="true" aria-label="Note editor">
      Many languages have words that carry meanings so specific or culturally rooted that they can't be neatly translated into English.

      One example is the Japanese word "tsundoku", which refers to the habit of acquiring books and letting them pile up unread, something many book lovers can relate to. Another is the Portuguese word "saudade", describing a deep, bittersweet longing for something or someone that is absent. Meanwhile, the French word "Dépaysement" captures the disorienting yet exciting feeling of being in a new place, far from home.

      These unique words remind us that language is more than vocabulary: it's a window into the values, habits, and emotions of the cultures that create it.
    </div>

    <div id="status" aria-live="polite"></div>

    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  margin: 2em;
  max-width: 700px;
  background-color: #f5f5f5;
}

.note {
  background-color: #ffffff;
  border: 1px solid #ddd;
  border-radius: 4px;
  box-shadow: 0 2px 5px rgba(0, 0, 0, 0.1);
  padding: 1.5em;
  margin-bottom: 1em;
  line-height: 1.5;
  min-height: 250px;
  font-size: 16px;
  /* This is needed to preserve line breaks in the div */
  white-space: pre-wrap;
}

.note[contenteditable="true"] {
  caret-color: black;
}

.note:hover {
  background-color: #fff;
  box-shadow: 0 0 5px rgba(0,0,0,0.2);
}

.helper-text {
  font-size: 0.9rem;
  color: #666;
  margin-top: 0.5em;
  user-select: none;
  font-style: italic;
}

#status {
  color: #00471b;
  padding: 0 1em;
}
```

```js
const noteEl = document.getElementById("note");
const statusEl = document.getElementById("status");

let currentContent = "";

--fcc-editable-region--

--fcc-editable-region--

window.addEventListener("DOMContentLoaded", () => {
  currentContent = noteEl.textContent;
});
```
