---
id: 694c84330cbe604821b9deef
title: Étape 3
challengeType: 1
dashedName: step-3
---

# --description--

Le contrôle de mission a besoin d'une méthode pour ajouter de nouveaux membres à la liste. Crée une fonction vide nommée `addCrewMember` qui accepte deux paramètres, `crew` et `astronaut`.

# --hints--

Tu dois créer une fonction nommée `addCrewMember`.

```js
assert.isFunction(addCrewMember);
```

Ta fonction `addCrewMember` doit avoir les paramètres `crew` et `astronaut`.

```js
const regex = __helpers.functionRegex('addCrewMember', ['crew', 'astronaut']);
assert.match(__helpers.removeJSComments(addCrewMember.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

--fcc-editable-region--

--fcc-editable-region--
```
