---
id: 694c9c186c089f7bdff0f3c6
title: Étape 4
challengeType: 1
dashedName: step-4
---

# --description--

En général, c'est une bonne pratique de valider tes entrées à l'intérieur d'une fonction. Utilise une boucle `for` pour parcourir le tableau `crew` et vérifier si un membre a déjà le même `id` que l'`astronaut` fourni. Si un doublon est trouvé, affiche le message d'erreur `console.log("Duplicate ID: " + astronaut.id)` puis appelle `return` pour sortir de la fonction de façon anticipée.

Voici un exemple de cette technique couramment utilisée :

```js
if (existingMember.id === newMember.id) {
    console.log("Duplicate ID: " + newMember.id);
    return;
}
```

# --hints--

Ta fonction doit utiliser une boucle `for` pour parcourir les membres existants de `crew` (toute forme valide de `for (...)` convient).

```js
assert.match(__helpers.removeJSComments(code), /for\s*\([^)]*\)/);
```

À l'intérieur de la boucle `for`, tu dois vérifier s'il y a un `id` en double. Si c'en est le cas, affiche `"Duplicate ID: " + astronaut.id"` dans la console puis appelle `return` pour sortir de la fonction.

```js
const spy = __helpers.spyOn(console, "log");

try {
  const crewWithDuplicate = [
    { id: 1 },
    { id: 2 },
    { id: 3 },
    { id: 4 },
    { id: 3 }
  ];
  addCrewMember(crewWithDuplicate, { id: 3 });

  assert.lengthOf(
    spy.calls,
    1,
    "Duplicate ID input should trigger exactly one console.log and then exit"
  );
  assert.match(
    spy.calls[0][0],
    /duplicate\s+id\s*:\s*3/i,
    "Console message should mention 'duplicate id'"
  );

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

Lorsqu'on lui fournit des entrées valides, ta fonction `addCrewMember` ne doit pas déclencher de message d'erreur `console.log()` ni de sortie anticipée par `return`.

```js
const spy = __helpers.spyOn(console, "log");

try {
  const crewNoDuplicate = [
    { id: 1 },
    { id: 2 },
    { id: 3 },
    { id: 4 }
  ];
  addCrewMember(crewNoDuplicate, { id: 5 });

  assert.lengthOf(
    spy.calls,
    0,
    "A truly new astronaut should produce zero console.log calls"
  );

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
--fcc-editable-region--

--fcc-editable-region--
}
```
