---
id: 694c84106e9023f289c03b84
title: Étape 2
challengeType: 1
dashedName: step-2
---

# --description--

Ta liste commence avec un seul astronaute commandant. Crée un objet nommé `firstAstronaut` avec les propriétés suivantes :

| Clé | Valeur |
| ----------- | ------- |
|`id`|`1`|
|`name`|`"Andy"`|
|`role`|`"Commander"`|
|`isEVAEligible`|`true`|
|`priority`|`3`|

# --hints--

Tu dois créer une variable `firstAstronaut`.

```js
assert.isDefined(firstAstronaut);
```

Ton `firstAstronaut` doit être un objet.

```js
assert.isObject(firstAstronaut);
```

Ton objet `firstAstronaut` doit inclure les bonnes valeurs de propriété.

```js
assert.deepEqual(firstAstronaut, { id: 1, name: "Andy", role: "Commander", isEVAEligible: true, priority: 3 });
```

# --seed--

## --seed-contents--

```js
const squad = [];

--fcc-editable-region--

--fcc-editable-region--
```
