---
id: 694c40d86fe6a78d20b6559c
title: Étape 1
challengeType: 1
dashedName: step-1
---

# --description--

Dans ce workshop, tu vas utiliser JavaScript pour construire une liste d'astronautes pour une mission spatiale.

Pour commencer, crée un tableau vide pour ta liste, nommé `squad`.

# --hints--

Tu dois créer une variable `squad`.

```js
assert.isDefined(squad);
```

Ta variable `squad` doit être un tableau.

```js
assert.isArray(squad);
```

Ton tableau `squad` doit être vide.

```js
assert.deepEqual(squad, []);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
