---
id: 694d8ddd106340237efb0c80
title: Étape 12
challengeType: 1
dashedName: step-12
---

# --description--

La méthode `splice()` peut modifier les tableaux en ajoutant ou en supprimant des éléments à n'importe quelle position, y compris au milieu. Comme `splice()` retourne un tableau contenant les éléments supprimés, tu peux l'utiliser pour échanger deux éléments d'un tableau sans modifier l'original. Cela peut se faire en une ligne en utilisant la technique suivante :

```js
// swap elements at i and j without mutating the original
const copy = array.slice();
copy[i] = copy.splice(j, 1, copy[i])[0];
```

Voici un exemple :

```js
const originalArray = [12, 97, 68, 55];
const copyArray = originalArray.slice();
copyArray[1] = copyArray.splice(3, 1, copyArray[1])[0];
console.log(copyArray); // [12, 55, 68, 97]
```

La technique, appliquée à la troisième ligne, fonctionne comme suit :

- `splice(3, 1, copyArray[1])` supprime l'élément à l'index `3` (`55`)

- Elle insère l'élément de `copyArray[1]` (`97`) à l'index `3`

- `splice()` retourne `[55]`, un tableau contenant l'élément supprimé

- `[0]` extrait `55` de ce tableau

- Cette valeur (`55`) est réaffectée à `copyArray[1]`, ce qui complète l'échange

Utilise cette technique avec le tableau `updatedCrew` pour échanger les astronautes situés à `fromIndex` et `toIndex`.

# --hints--

Tu dois effectuer l'échange en appelant la méthode `splice()`.

```js
assert.match(swapCrewMembers.toString(), /splice\s*\(/);
```

Tu dois appeler `splice()` avec le tableau `updatedCrew` pour effectuer l'échange. Tu dois utiliser la technique en une ligne.

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(swapCrewMembers.toString()));
const regex = /updatedCrew\[(fromIndex|toIndex)\]=updatedCrew\.splice\((fromIndex|toIndex),1,updatedCrew\[\1\]\)\[0\];?\s*\}$/;
assert.match(cleaned, regex);
```

Ta fonction `swapCrewMembers` ne doit pas modifier le tableau `crew` fourni en entrée.

```js
const demo = [{ id: 1 }, { id: 2 }, { id: 3 }];
swapCrewMembers(demo, 0, 2);
assert.deepEqual(demo.map(a => a.id), [1, 2, 3]);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

function swapCrewMembers(crew, fromIndex, toIndex) {
  if (
    fromIndex < 0 ||
    toIndex < 0 ||
    fromIndex >= crew.length ||
    toIndex >= crew.length
  ) {
    console.log("Invalid crew indices");
    return;
  }

  const updatedCrew = crew.slice();
--fcc-editable-region--

--fcc-editable-region--
}
```
