---
id: 68552464ef378e1051a55c28
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

À l’intérieur du second `div`, crée deux éléments d’ancre. Le premier doit avoir le texte `Learn more`, et le second le texte `Start listening`.

Pour l’attribut `href` des deux ancres, définis-les toutes les deux à `#`.

# --hints--

Tu dois créer deux éléments `a` à l’intérieur du second élément `div`

```js
assert.lengthOf(document.querySelectorAll("div > div:nth-child(2) > a"), 2)
```

Ton premier élément `a` doit avoir le texte `Learn more` et un attribut `href` égal à `#` à l’intérieur du second élément `div`.

```js
const links = document.querySelectorAll("div > div:nth-child(2) > a")

assert.equal(links[0]?.textContent, "Learn more")
assert.equal(links[0]?.getAttribute("href"), "#")
```

Ton second élément `a` doit avoir le texte `Start listening` et un attribut `href` égal à `#` à l’intérieur du second élément `div`.

```js
const links = document.querySelectorAll("div > div:nth-child(2) > a")

assert.equal(links[1]?.textContent, "Start listening")
assert.equal(links[1]?.getAttribute("href"), "#")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        --fcc-editable-region--

        --fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
