---
id: 68551b677a5f6898af364b85
title: Étape 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Dans cet atelier, tu vas approfondir Tailwind CSS en utilisant ses classes utilitaires Flexbox pour construire un composant d’appel à l’action (CTA) responsive.

Le boilerplate HTML de base t’a été fourni, donc tu dois d’abord construire les éléments HTML.

Commence par créer un élément `div`. À l’intérieur de ce `div`, crée deux autres éléments `div` séparés.

# --hints--

Tu dois créer un élément `div`.

```js
assert.exists(document.querySelector("div"));
```

Tu dois avoir deux autres éléments `div` à l’intérieur de l’élément `div`.

```js
assert.lengthOf(document.querySelectorAll("div > div"), 2)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>
```
