---
id: 6858ebf2552046ab3de9f17c
title: Étape 7
challengeType: 0
dashedName: step-7
---

# --description--

À ce stade, tu remarqueras qu’il y a trop d’espace à droite du contenu dans le `div`. Tu peux réduire cela en ajoutant une largeur.

Ajoute la classe `md:w-1/2` au `div` pour réduire la largeur à 50 % sur les écrans moyens et plus grands. C’est-à-dire les écrans de 768 pixels et plus.

Pour voir l’effet, agrandis l’onglet d’aperçu.

# --hints--

Ton élément `div` extérieur doit avoir la classe utilitaire `md:w-1/2`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("md:w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
