---
id: 6858ef974c8ad5dcd7abb9f1
title: Étape 9
challengeType: 0
dashedName: step-9
---

# --description--

Le contenu à l’intérieur du `div` peut être mieux espacé et arrangé. Pour le reste de l’espacement et cet arrangement, tu dois utiliser Flexbox.

Ajoute la classe `flex` au `div` pour que les deux `div` enfants soient côte à côte.

Ajoute aussi `justify-around` pour que l’espace autour et entre le contenu soit réparti uniformément, et `items-center` pour que tout le contenu soit enfin centré dans le `div`.

# --hints--

Ton élément `div` extérieur doit avoir la classe utilitaire `flex`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("flex"));
```

Ton élément `div` extérieur doit avoir la classe utilitaire `justify-around`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("justify-around"));
```

Ton élément `div` extérieur doit avoir la classe utilitaire `items-center`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("items-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
