---
id: 6858e7e23bf9e475e14cb931
title: Étape 4
challengeType: 0
dashedName: step-4
---

# --description--

Maintenant que la structure HTML est en place, il est temps de passer au style avec Tailwind.

Commence par l’élément `div` qui enveloppe tout le contenu. Définis sa couleur de fond en lui donnant une `class` de `bg-indigo-600`.

# --hints--

Ton `div` extérieur doit avoir un attribut `class` égal à `bg-indigo-600`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("bg-indigo-600"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div>
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
   --fcc-editable-region--
  </body>
</html>
```
