---
id: 6858f3a67aa8580ff37016c7
title: Étape 10
challengeType: 0
dashedName: step-10
---

# --description--

Si tu ajustes la largeur de l’onglet d’aperçu à celle d’un appareil mobile, tu remarqueras que le contenu reste toujours côte à côte. Ce n’est pas ce que tu veux sur mobile.

Pour corriger cela, ajoute la classe `flex-col` au `div`. Cela change la direction flex en colonne, donc les éléments s’empilent les uns sur les autres.

Voici comment utiliser la classe utilitaire `flex-col` :

```html
<div class="flex flex-col">
  <element></element>
  <element></element>
</div>
```

# --hints--

Ton élément `div` ouvrant doit avoir la classe utilitaire `flex-col`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("flex-col"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex justify-around items-center">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
