---
id: 6858f90c54647a48ddb8822d
title: Étape 11
challengeType: 0
dashedName: step-11
---

# --description--

À ce stade, le contenu, c’est-à-dire les deux éléments `div` à l’intérieur du `div` extérieur, sera empilé quelle que soit la taille d’écran.

C’est parce que tu as la classe `flex-col` dans l’élément `div` enveloppant. Pour corriger cela, utilise le préfixe `lg:` pour ajouter la classe `flex-row` au `div`.

Cela signifie que sur les appareils avec de plus grands écrans, la direction flex sera `row`. Pour voir l’effet, ajuste le navigateur à une plus grande largeur.

# --hints--

Ton élément `div` extérieur doit avoir la classe utilitaire `lg:flex-row`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("lg:flex-row"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col justify-around items-center">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
