---
id: 68597114730d7c6b8838d204
title: Étape 15
challengeType: 0
dashedName: step-15
---

# --description--

Le paragraphe sous le `h1` est bien, donc il n’a pas besoin de plus de style. Tu peux passer aux éléments d’ancre à l’intérieur du second élément `div`.

Donne au premier élément d’ancre les classes `bg-white`, `hover:bg-gray-100` et `text-indigo-600`. Cela lui donne un fond blanc, une couleur de texte indigo, et un fond légèrement plus clair au survol.

Pour le second élément d’ancre, donne-lui les classes `bg-pink-500`, `hover:bg-pink-600` et `text-white`.

# --hints--

Ton premier élément d’ancre doit avoir la classe utilitaire `bg-white`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("bg-white"));
```

Ton premier élément d’ancre doit avoir la classe utilitaire `hover:bg-gray-100`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("hover:bg-gray-100"));
```

Ton premier élément d’ancre doit avoir la classe utilitaire `text-indigo-600`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("text-indigo-600"));
```

Ton second élément d’ancre doit avoir la classe utilitaire `bg-pink-500`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("bg-pink-500"));
```

Ton second élément d’ancre doit avoir la classe utilitaire `hover:bg-pink-600`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("hover:bg-pink-600"));
```

Ton second élément d’ancre doit avoir la classe utilitaire `text-white`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("text-white"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        <span class="uppercase">Soundflow</span>
        <h1 class="font-bold text-4xl my-4">Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        --fcc-editable-region--
        <a href="#">
          Learn more
        </a>
        <a href="#">
          Start listening
        </a>
        --fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
