---
id: 66c06d618d075c7f7f1b890a
title: Crée une diseuse de bonne aventure
challengeType: 26
dashedName: build-a-fortune-teller
---

# --description--

Dans ce lab, tu vas définir cinq prédictions et en sélectionner une au hasard pour afficher la prédiction à l'utilisateur.

**Objectif :** Réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu devrais initialiser les cinq variables `fortune1`, `fortune2`, `fortune3`, `fortune4` et `fortune5` avec une valeur de chaîne de ton choix. Tu peux utiliser les options ci-dessous si tu le souhaites :

   - `"Ton chat aura l'air très câlin aujourd'hui."`
   - `"Il fera beau demain."`
   - `"Méfie-toi de tes nouveaux voisins."`
   - `"Tu trouveras bientôt un nouveau passe-temps."`
   - `"Il serait sage d'éviter la couleur rouge aujourd'hui."`

2. Tu devrais sélectionner un nombre aléatoire entre 1 et 5 inclus, et l'affecter à la variable `randomNumber`.

3. Tu devrais créer une variable `selectedFortune` et lui affecter la prédiction appropriée selon ces règles :

   - Si `randomNumber` vaut 1, affecte la valeur de `fortune1` à `selectedFortune`.
   - Si `randomNumber` vaut 2, affecte la valeur de `fortune2` à `selectedFortune`.
   - Si `randomNumber` vaut 3, affecte la valeur de `fortune3` à `selectedFortune`.
   - Si `randomNumber` vaut 4, affecte la valeur de `fortune4` à `selectedFortune`.
   - Si `randomNumber` vaut 5, affecte la valeur de `fortune5` à `selectedFortune`.

4. Tu devrais afficher `selectedFortune` dans la console.

# --hints--

Tu devrais initialiser `fortune1` avec une valeur de chaîne.

```js
assert.isNotNull(fortune1);
assert.isString(fortune1);
```

Tu devrais initialiser `fortune2` avec une valeur de chaîne.

```js
assert.isNotNull(fortune2);
assert.isString(fortune2);
```

Tu devrais initialiser `fortune3` avec une valeur de chaîne.

```js
assert.isNotNull(fortune3);
assert.isString(fortune3);
```

Tu devrais initialiser `fortune4` avec une valeur de chaîne.

```js
assert.isNotNull(fortune4);
assert.isString(fortune4);
```

Tu devrais initialiser `fortune5` avec une valeur de chaîne.

```js
assert.isNotNull(fortune5);
assert.isString(fortune5);
```

Tu devrais utiliser la méthode `Math.random()` pour générer un nombre aléatoire.

```js
assert.match(__helpers.removeJSComments(code), /Math\.random\(\)/);
```

Tu devrais générer un nombre aléatoire entre 1 et 5 inclus, et l'affecter à la variable `randomNumber`.

```js
assert.isNotNull(randomNumber);
const originalRandom__test_code__ = Math.random;
try {
  const mockValues__test_code__ = [0.0, 0.2, 0.4, 0.6, 0.8];
  let callCount__test_code__ = 0;
  Math.random = () => mockValues__test_code__[callCount__test_code__++ % mockValues__test_code__.length]
  const userCode__test_code__ = new Function(`${code};return randomNumber`);
  for (let i = 0; i < 5; i++) {
    const result = userCode__test_code__()
    assert.strictEqual(result, i + 1);
  }
} finally {
  Math.random = originalRandom__test_code__;
}
```

Tu devrais avoir une variable `selectedFortune` à laquelle une valeur est affectée en fonction de la valeur de `randomNumber`.

```js
assert.isNotNull(selectedFortune);
```

Le `randomNumber` devrait correspondre à sa prédiction. Par exemple, si `randomNumber` vaut 1, `selectedFortune` devrait être égal à `fortune1`, et ainsi de suite.

```js
const originalRandom__test_code__ = Math.random;
try {
  const mockValues__test_code__ = [0.0, 0.2, 0.4, 0.6, 0.8];
  let callCount__test_code__ = 0;
  Math.random = () => mockValues__test_code__[callCount__test_code__++ % mockValues__test_code__.length];
  const fortunes__test_code__ = [fortune1, fortune2, fortune3, fortune4, fortune5,];
  const userCode__test_code__ = new Function(`${code};return selectedFortune`);
  for (let i = 0; i < 5; i++) {
    const result = userCode__test_code__();
    assert.strictEqual(result, fortunes__test_code__[i]);
  }
} finally {
  Math.random = originalRandom__test_code__;
}
```

Tu devrais afficher `selectedFortune` dans la console.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*selectedFortune\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const fortune1 = "Your cat will look very cuddly today.";
const fortune2 = "The weather will be nice tomorrow.";
const fortune3 = "Be cautious of your new neighbors.";
const fortune4 = "You will find a new hobby soon.";
const fortune5 = "It would be wise to avoid the color red today.";

let randomNumber = Math.floor(Math.random() * 5) + 1;

let selectedFortune;

if (randomNumber === 1) {
  selectedFortune = fortune1;
} else if (randomNumber === 2) {
  selectedFortune = fortune2;
} else if (randomNumber === 3) {
  selectedFortune = fortune3;
} else if (randomNumber === 4) {
  selectedFortune = fortune4;
} else {
  selectedFortune = fortune5;
}

console.log(selectedFortune);
```
