---
id: 60a3e3396c7b40068ad6996b
title: Étape 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Commence ton projet en ajoutant un élément `img` avec un `src` de `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` dans l’élément `body`.

# --hints--

Ton code devrait contenir un élément `img`.

```js
assert(document.querySelectorAll('img').length === 1);
```

L’élément `img` devrait avoir un `src` de `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`.

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
