---
id: 60a3e3396c7b40068ad69980
title: Étape 22
challengeType: 0
dashedName: step-22
---

# --description--

Ajoute un autre `div` avec une valeur de `class` de `two` juste sous ton élément `one`. Ce sera ton deuxième rectangle.

# --hints--

Tu ne devrais pas modifier l’élément `.one` existant.

```js
assert(document.querySelectorAll('.one').length === 1);
```

Tu devrais avoir un deuxième élément `div` dans ton élément `.canvas`.

```js
assert(document.querySelector('.canvas').children[1].tagName === 'DIV');
```

Ton deuxième élément `div` devrait avoir une valeur de `class` de `two`.

```js
assert(document.querySelector('.canvas').children[1].className.split(' ').includes('two'));
```

Ton élément `.two` devrait venir après ton élément `.one`.

```js
assert(document.querySelector('.two').previousElementSibling.className.split(' ').includes('one'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
--fcc-editable-region--

--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
