---
id: 60a3e3396c7b40068ad69972
title: Étape 8
challengeType: 0
dashedName: step-8
---

# --description--

Ajoute la propriété `height` avec la valeur `600px` à ta règle `.canvas`.

# --hints--

Tu devrais définir la propriété `height` sur `600px`.

```js
const hasHeight = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.height === '600px');
assert(hasHeight);
```

Ton élément `.canvas` devrait avoir une propriété `height` avec la valeur `600px`.

```js
const canvasHeight = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('height');
assert(canvasHeight === '600px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
