---
id: 60a3e3396c7b40068ad6997f
title: Étape 21
challengeType: 0
dashedName: step-21
---

# --description--

Ajouter 1 pixel de padding en haut, en bas, à gauche et à droite de la toile a changé ses dimensions à 502 pixels par 602 pixels.

Remplace la propriété `padding` par `overflow` défini sur `hidden`, ce qui ramènera la toile à ses dimensions d’origine.

# --hints--

Tu devrais supprimer la propriété `padding` du sélecteur `.canvas`.

```js
const canvasPadding = new __helpers.CSSHelp(document).getStyle('.canvas').getPropertyValue('padding');
assert(!canvasPadding);
```

Tu devrais définir la propriété `overflow` sur `hidden`.

```js
const hasOverflow = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.overflow === 'hidden');
assert(hasOverflow);
```

Ton élément `.canvas` devrait avoir la propriété `overflow` avec la valeur `hidden`.

```js
const canvasOverflow = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('overflow');
assert(canvasOverflow === 'hidden')
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
--fcc-editable-region--
  padding: 1px;
--fcc-editable-region--
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
