---
id: 60a3e3396c7b40068ad69973
title: Étape 9
challengeType: 0
dashedName: step-9
---

# --description--

Remplace la propriété `background-color` de la toile par `#4d0f00`.

# --hints--

Tu devrais définir la propriété `background-color` sur `#4d0f00`.

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-color'] === 'rgb(77, 15, 0)');
assert(hasBackground);
```

Ton élément `.canvas` devrait avoir une propriété `background-color` avec la valeur `#4d0f00`.

```js
const canvasBackground = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('background-color');
assert(canvasBackground === 'rgb(77, 15, 0)');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
