---
id: 60a3e3396c7b40068ad69979
title: Étape 15
challengeType: 0
dashedName: step-15
---

# --description--

Ajoute un nouvel élément `div` dans ton élément `.canvas`.

Donne au nouveau `div` l’attribut `class` avec la valeur `one`. Ce sera ton premier rectangle.

# --hints--

Tu devrais créer un nouvel élément `div`.

```js
assert(document.querySelectorAll('div').length === 3);
```

Tu devrais imbriquer le nouvel élément `div` dans ton élément `.canvas`.

```js
assert(document.querySelector('.canvas').children[0].tagName === 'DIV');
```

Ton nouveau `div` devrait avoir un attribut `class` avec la valeur `one`.

```js
assert(document.querySelector('.canvas').children[0].className.split(' ').includes('one'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
--fcc-editable-region--

--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
