---
id: 60a3e3396c7b40068ad69971
title: Étape 7
challengeType: 0
dashedName: step-7
---

# --description--

Place au CSS.

Même si ton `<div>` ne contient pas de texte, il est quand même traité comme une boîte avec du contenu.
Écris une règle CSS qui utilise le sélecteur de classe `.canvas` et définis sa `width` sur 500 pixels.
Voici une règle CSS qui définit la largeur de la classe `card` sur 300 pixels :

```css
.card {
  width: 300px;
}
```

# --hints--

Ton code devrait contenir un sélecteur `.canvas`.

```js
const hasCanvas = new __helpers.CSSHelp(document).getStyle('.canvas');
assert(hasCanvas)
```

Tu devrais définir la propriété `width` sur `500px`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '500px')
assert(hasWidth);
```

Ton élément `.canvas` devrait avoir une propriété `width` avec la valeur `500px`.

```js
const width = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('width');
assert(width === '500px');
```

# --seed--

## --seed-contents--

```css
--fcc-editable-region--

--fcc-editable-region--

```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
