---
id: 60a3e3396c7b40068ad6998b
title: Étape 32
challengeType: 0
dashedName: step-32
---

# --description--

Il est utile que tes marges poussent dans une seule direction.

Dans ce cas, la marge inférieure de l’élément `.one` pousse `.two` de 20 pixels vers le bas.

Dans le sélecteur `.two`, utilise la propriété raccourcie `margin` pour définir la marge supérieure sur `0`, la marge horizontale sur `auto` et la marge inférieure sur `20px`. Cela supprimera sa marge supérieure, le centrera horizontalement et définira sa marge inférieure à 20 pixels.

# --hints--

Tu devrais définir la propriété `margin` sur `0 auto 20px`.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '0px auto 20px');
assert(hasMargin);
```

Ton élément `.two` devrait avoir la propriété `margin` avec la valeur `0 auto 20px`.

```js
const twoMargin = new __helpers.CSSHelp(document).getStyle('.two')?.getPropertyValue('margin');
assert(twoMargin === '0px auto 20px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}

.two {
  width: 475px;
  height: 200px;
  background-color: #8f0401;
--fcc-editable-region--
  margin: auto;
--fcc-editable-region--
}

.three {
  width: 91%;
  height: 28%;
  background-color: #b20403;
  margin: auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
        <div class="three"></div>
      </div>
    </div>
  </body>
</html>
```
