---
id: 60a3e3396c7b40068ad69986
title: Étape 27
challengeType: 0
dashedName: step-27
---

# --description--

Crée un nouveau `div` avec une valeur de `class` de `three` juste sous l’élément `.two`. Ce sera ton troisième rectangle.

# --hints--

Tes éléments `.one` et `.two` existants ne devraient pas être modifiés.

```js
assert(document.querySelectorAll('.one').length === 1);
assert(document.querySelectorAll('.two').length === 1);
```

Ton nouveau `div` devrait être imbriqué dans ton élément `.canvas`.

```js
assert(document.querySelector('.canvas').children[2].tagName === 'DIV');
```

Ton nouveau `div` devrait venir après ton élément `.two`.

```js
assert(document.querySelector('.two').nextElementSibling.tagName === 'DIV');
```

Ton nouvel élément `div` devrait avoir une classe avec la valeur `three`.

```js
assert(document.querySelector('.canvas').children[2].className.split(' ').includes('three'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}

.two {
  width: 475px;
  height: 200px;
  background-color: #8f0401;
  margin: auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
--fcc-editable-region--

--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
