---
id: 60a3e3396c7b40068ad69975
title: Étape 11
challengeType: 0
dashedName: step-11
---

# --description--

Écris une nouvelle règle avec le sélecteur de classe `.frame`.

Utilise la déclaration raccourcie `border` pour donner à l’élément `.frame` une bordure noire pleine d’une largeur de `50px`.

# --hints--

Ton code devrait contenir un sélecteur `.frame`.

```js
const hasFrame = new __helpers.CSSHelp(document).getStyle('.frame');
assert(hasFrame);
```

Tu devrais définir la propriété `border` sur `50px solid black`.

```js
const hasBorder = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.border === '50px solid black');
assert(hasBorder);
```

Ton élément `.frame` devrait avoir la propriété `border` avec la valeur `50px solid black`.

```js
const frameBorder = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('border');
assert(frameBorder === '50px solid black');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

--fcc-editable-region--

--fcc-editable-region--
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
