---
id: 60a3e3396c7b40068ad6996c
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

Dans les leçons précédentes sur le modèle de boîte CSS, tu as appris que chaque élément HTML est traité comme une boîte avec quatre zones.

Imagine que tu reçoives un colis de ta boutique en ligne préférée : le contenu est l’objet dans la boîte ou, dans ton cas, un élément d’en-tête, de paragraphe ou d’image.

Remplace l’attribut `src` dans le `<img>`, de `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` à `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`.

# --hints--

L’élément `img` devrait avoir un `src` de `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`.

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png">
--fcc-editable-region--
  </body>
</html>
```
