---
id: 60a3e3396c7b40068ad6997e
title: Étape 20
challengeType: 0
dashedName: step-20
---

# --description--

Maintenant, `.one` est centré horizontalement, mais sa marge supérieure dépasse de la toile et pousse sur la bordure du cadre, ce qui décale toute la toile de 20 pixels vers le bas.

Ajoute un `padding` de `1px` à l’élément `.canvas` pour donner à l’élément `.one` une surface contre laquelle pousser.

# --hints--

Tu devrais définir la propriété `padding` sur `1px`.

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '1px');
assert(hasPadding);
```

Ton élément `.canvas` devrait avoir la propriété `padding` avec la valeur `1px`.

```js
const canvasPadding = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('padding');
assert(canvasPadding === '1px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
--fcc-editable-region--

--fcc-editable-region--
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
