---
id: 60a3e3396c7b40068ad6996e
title: Étape 4
challengeType: 0
dashedName: step-4
---

# --description--

La marge est la zone située à l’extérieur de la boîte et peut servir à contrôler l’espace entre d’autres boîtes ou éléments.

Ici, l’élément du bas a une marge supérieure plus grande, ce qui le pousse plus bas dans la page.

Maintenant que tu as rapidement revu le modèle de boîte CSS, il est temps de commencer la peinture de style Rothko.

Supprime l’élément `img`.

# --hints--

Tu ne devrais pas avoir d’élément `img` dans ton code.

```js
assert(document.querySelector('img') === null);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png">
--fcc-editable-region--
  </body>
</html>
```
