---
id: 60a3e3396c7b40068ad69974
title: Étape 10
challengeType: 0
dashedName: step-10
---

# --description--

Chaque peinture a besoin d’un cadre.

Enveloppe l’élément `.canvas` dans un autre `div`. Donne à ce `div` la classe `frame`.
# --hints--

Tu devrais ajouter un nouvel élément `div`.

```js
assert(document.querySelectorAll('div').length === 2)
```

Ton élément `.canvas` devrait être imbriqué dans le nouvel élément `div`.

```js
assert(document.querySelector('.canvas').parentElement.tagName === 'DIV');
```

Ton nouveau `div` devrait avoir une classe avec la valeur `frame`.

```js
assert(document.querySelector('div').className.split(' ').includes('frame'));
```

Ton nouveau `div` devrait se trouver dans ton élément `body`.

```js
assert(document.querySelector('div').parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
--fcc-editable-region--

    <div class="canvas">
    </div>

--fcc-editable-region--
  </body>
  ```
