---
id: 60a3e3396c7b40068ad69970
title: Étape 6
challengeType: 0
dashedName: step-6
---

# --description--

Avant de commencer à styliser le `div` que tu as ajouté, tu dois lier ton CSS à ton HTML.

Ajoute un élément `link` pour lier ton fichier `styles.css`. Définis `href` sur `styles.css` et pense à définir l’attribut `rel` sur `stylesheet`.

# --hints--

Ton code devrait contenir un élément `link`.

```js
assert(/<link/.test(code))
```

Ton élément `link` devrait avoir un attribut `rel` avec la valeur `stylesheet`.

```js
assert.match(code, /<link.*?rel=('|"|`)stylesheet\1/)
```

Ton élément `link` devrait avoir un attribut `href` avec la valeur `styles.css`.

```js
assert.match(code, /<link.*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
