---
id: 60a3e3396c7b40068ad6996d
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

Le contenu est entouré d’un espace appelé padding, un peu comme le papier bulle sépare un objet de la boîte qui l’entoure.

Pense à la bordure comme à la boîte en carton dans laquelle ton objet a été expédié.

Remplace l’attribut `src` par `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png`.

# --hints--

L’élément `img` devrait avoir un `src` de `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png`.

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png">
--fcc-editable-region--
  </body>
</html>
```
