---
id: 60a3e3396c7b40068ad69977
title: Étape 13
challengeType: 0
dashedName: step-13
---

# --description--

Utilise le padding pour ajuster l’espacement à l’intérieur d’un élément.

Dans `.frame`, utilise la propriété raccourcie `padding` pour augmenter de `50px` l’espace entre les éléments `.frame` et `.canvas`. Le raccourci augmentera l’espace en haut, en bas, à gauche et à droite entre la bordure de l’élément et la toile à l’intérieur.

# --hints--

Tu devrais définir la propriété `padding` sur `50px`.

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '50px');
assert(hasPadding);
```

Ton élément `.frame` devrait avoir la propriété `padding` avec la valeur `50px`.

```js
const framePadding = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('padding');
assert(framePadding === '50px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
