---
id: 60a3e3396c7b40068ad6996f
title: Étape 5
challengeType: 0
dashedName: step-5
---

# --description--

Ajoute un élément `div` dans le `body`.

Définis l’attribut `class` sur `canvas`.

Cela servira de toile pour ta peinture.

# --hints--

Ton code devrait contenir un élément `div`.

```js
assert(document.querySelectorAll('div').length === 1)
```

L’élément `div` devrait avoir une classe avec la valeur `canvas`.

```js
assert(document.querySelector('div').className.split(' ').includes('canvas'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
