---
id: 60a3e3396c7b40068ad69978
title: Étape 14
challengeType: 0
dashedName: step-14
---

# --description--

Utilise les marges pour ajuster l’espacement à l’extérieur d’un élément.

Avec la propriété `margin`, donne à l’élément `.frame` une marge verticale de `20px` et une marge horizontale de `auto`. Cela déplacera le cadre de 20 pixels vers le bas et le centrera horizontalement sur la page.

# --hints--

Tu devrais définir la propriété `margin` sur `20px auto`.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '20px auto');
assert(hasMargin);
```

Ton élément `.frame` devrait avoir la propriété `margin` avec la valeur `20px auto`.

```js
const frameMargin = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('margin');
assert(frameMargin === '20px auto');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
