---
id: 5dc174fcf86c76b9248c6eb2
title: Étape 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

En HTML, presque tous les éléments sont composés d'une **balise ouvrante** comme `<h1>` et d'une **balise fermante** comme `</h1>`. Le texte que l'élément affiche se place entre les deux.

L'élément `h1` représente le titre principal d'une page — c'est le titre le plus important.

Dans le code ci-dessous, remplace le texte `Hello World` par `CatPhotoApp`, puis regarde le résultat dans l'aperçu à droite.

Quand c'est fait, clique sur le bouton « Vérifier votre code » pour valider l'étape.

# --hints--

Le texte `CatPhotoApp` doit apparaître dans ton code. Vérifie bien l'orthographe.

```js
assert(code.match(/catphotoapp/i));
```

Ton élément `h1` doit avoir une balise ouvrante. Une balise ouvrante a la forme : `<nomDeLElement>`.

```js
assert(document.querySelector('h1'));
```

Tes balises `h1` doivent être écrites en minuscules. Par convention, toutes les balises HTML s'écrivent en minuscules.

```js
assert.notMatch(code, /<\/?H1>/);
```

Ton élément `h1` doit avoir une balise fermante. Une balise fermante porte un `/` juste après le `<`, comme dans `</h1>`.

```js
assert(code.match(/<\/h1\>/));
```

Tu as plusieurs éléments `h1` dans le code. Garde-en un seul et supprime les autres.

```js
assert(document.querySelectorAll('h1').length === 1);
```

Le texte de ton élément `h1` doit être exactement `CatPhotoApp`. Vérifie qu'il n'y a pas de faute de frappe et qu'il se trouve bien entre les balises `<h1>` et `</h1>`.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Hello World</h1>
--fcc-editable-region--
  </body>
</html>
```
