---
id: 5dc17dc8f86c76b9248c6eb5
title: Étape 4
challengeType: 0
dashedName: step-4
---

# --description--

Les **commentaires** permettent d'écrire des notes dans ton code sans que le navigateur ne les affiche. Ils servent à :

- expliquer ce que fait une partie du code,
- laisser des rappels (TODO),
- ou désactiver temporairement du code sans le supprimer.

En HTML, un commentaire commence par `<!--`, contient le texte que tu veux, et finit par `-->`.

Voici un exemple :

```html
<!-- TODO: Remove h1 -->
```

Ajoute un commentaire **au-dessus** de l'élément `p` avec ce texte :

`TODO: Add link to cat photos`

# --hints--

Ton commentaire doit commencer par `<!--`. Tu as oublié un ou plusieurs caractères d'ouverture.

```js
assert(code.match(/<!--/));
```

Ton commentaire doit finir par `-->`. Tu as oublié un ou plusieurs caractères de fermeture.

```js
assert(code.match(/-->/));
```

Tu ne dois pas avoir de caractères de commentaire en trop. Il y a un `<!--` ou un `-->` qui apparaît dans le navigateur.

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.match(/<!--/g).length < 2 && noSpaces.match(/-->/g).length < 2);
```

Ton commentaire doit être **au-dessus** de l'élément `p`. Vérifie l'ordre.

```js
assert.match(
  code.replace(/\s/g, ''),
  /<!--(.*?)--><p>everyonelovescutecatsonline!<\/p>/i
);
```

Ton commentaire doit contenir le texte `TODO: Add link to cat photos`.

```js
assert.match(code, /<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--
    
    <p>Everyone loves cute cats online!</p>
    
--fcc-editable-region--
  </body>
</html>
```
