---
id: 5dc1798ff86c76b9248c6eb3
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

HTML propose six niveaux de titres, de `h1` à `h6`. Plus le chiffre est petit, plus le titre est important : un `h1` est plus important qu'un `h2`, qui est plus important qu'un `h3`, et ainsi de suite.

```html
<h1>titre le plus important</h1>
<h2>deuxième niveau d'importance</h2>
<h3>troisième niveau d'importance</h3>
<h4>quatrième niveau d'importance</h4>
<h5>cinquième niveau d'importance</h5>
<h6>titre le moins important</h6>
```

**Bonne pratique :** on n'utilise qu'**un seul** `h1` par page, et on place les titres moins importants en dessous des plus importants.

Sous l'élément `h1`, ajoute un élément `h2` avec ce texte :

`Cat Photos`

# --hints--

Ton élément `h1` doit avoir une balise ouvrante. Une balise ouvrante a la forme : `<nomDeLElement>`.

```js
assert(document.querySelector('h1'));
```

Ton élément `h1` doit avoir une balise fermante. Une balise fermante porte un `/` juste après le `<`.

```js
assert(code.match(/<\/h1\>/));
```

Tu ne dois avoir qu'un seul élément `h1`. Supprime ceux en trop.

```js
assert(
  document.querySelector('h1') && document.querySelectorAll('h1').length === 1
);
```

Le texte de ton `h1` doit rester `CatPhotoApp`. Vérifie qu'il n'a pas changé.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

Tes balises `h2` doivent être en minuscules.

```js
assert.notMatch(code, /<\/?H2>/);
```

Ton élément `h2` doit avoir une balise ouvrante.

```js
assert(document.querySelector('h2'));
```

Ton élément `h2` doit avoir une balise fermante.

```js
assert(code.match(/<\/h2\>/));
```

Le texte de ton `h2` doit être `Cat Photos`. Place ce texte entre les balises ouvrante et fermante du `h2`.

```js
assert(document.querySelector('h2').innerText.toLowerCase() === 'cat photos');
```

Ton élément `h2` doit se trouver **en dessous** du `h1`. Le `h1` est plus important et doit donc apparaître en premier dans la page.

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H1') < collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>CatPhotoApp</h1>
    
--fcc-editable-region--
  </body>
</html>
```
