---
id: 5dc17d3bf86c76b9248c6eb4
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

L'élément `p` (pour *paragraph*, paragraphe) sert à créer un **paragraphe de texte** sur une page web. C'est l'élément le plus utilisé pour afficher du texte courant.

Crée un élément `p` **sous** ton élément `h2`, et donne-lui ce texte :

`Everyone loves cute cats online!`

# --hints--

Ton élément `p` doit avoir une balise ouvrante. Une balise ouvrante a la forme : `<nomDeLElement>`.

```js
assert(document.querySelector('p'));
```

Tes balises `p` doivent être en minuscules.

```js
assert.notMatch(code, /<\/?P>/);
```

Ton élément `p` doit avoir une balise fermante. Une balise fermante porte un `/` juste après le `<`.

```js
assert(code.match(/<\/p\>/));
```

Le texte de ton `p` doit être `Everyone loves cute cats online!`. Vérifie qu'il n'y a pas de faute de frappe et que tu n'as rien oublié.

```js
const extraSpacesRemoved = document
  .querySelector('p')
  .innerText.replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /everyone loves cute cats online!$/i);
```

Ton élément `p` doit se trouver **en dessous** du `h2`. L'ordre des éléments compte.

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H2') < collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <h2>Cat Photos</h2>
    
--fcc-editable-region--
  </body>
</html>
```
