---
id: 5dc2385ff86c76b9248c6eb7
title: Étape 5
challengeType: 0
dashedName: step-5
---

# --description--

HTML5 fournit des éléments dits **sémantiques** qui décrivent le rôle de chaque zone de la page. Ces éléments rendent ton HTML plus lisible et améliorent le **référencement** (SEO) ainsi que l'**accessibilité** (lecteurs d'écran).

L'élément `main` représente le **contenu principal** de la page. Tout ce qui est à l'intérieur doit être unique à cette page (pas l'en-tête, ni le menu, ni le pied de page qui se répètent d'une page à l'autre).

```html
<main>
  <h1>Contenu principal de la page</h1>
  <p>Un peu plus de contenu important...</p>
</main>
```

Identifie la zone principale de la page en ajoutant une balise ouvrante `<main>` **avant** l'élément `h1`, et une balise fermante `</main>` **après** l'élément `p`.

# --hints--

Ton élément `main` doit avoir une balise ouvrante.

```js
assert(document.querySelector('main'));
```

Tes balises `main` doivent être en minuscules.

```js
const mainTagMatches = code.matchAll(/<\/?main>/gi);
for (const match of mainTagMatches) {
  const tag = match[0];
  assert.strictEqual(
    tag,
    tag.toLowerCase()
  );
}
```

Ton élément `main` doit avoir une balise fermante. Une balise fermante porte un `/` juste après le `<`.

```js
assert(code.match(/<\/main\>/));
```

La balise ouvrante `<main>` doit se trouver **juste après** la balise ouvrante `<body>`, sans aucun autre élément entre les deux.

```js
const main = document.querySelector('main');
assert.equal(main?.previousElementSibling, null);
```

La balise ouvrante `<main>` doit être **au-dessus** de l'élément `h1`.

```js
const collection = [...document.querySelectorAll('main,h1')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('MAIN') < collection.indexOf('H1'));
```

La balise fermante `</main>` doit être **en dessous** de l'élément `p`, de façon à ce que le `p` soit bien contenu dans le `main`.

```js
const mainNode = document.querySelector('main');
const pNode = document.querySelector('p');
assert.isTrue(mainNode.contains(pNode));
assert.match(pNode.textContent.toLowerCase(), /everyone loves cute cats online!/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
  
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>Everyone loves cute cats online!</p>
    
--fcc-editable-region--
  </body>
</html>
```
