---
id: 5dc24073f86c76b9248c6ebb
title: Étape 8
challengeType: 0
dashedName: step-8
---

# --description--

Les **attributs** HTML sont des mots spéciaux qu'on place dans la balise ouvrante d'un élément pour modifier son comportement. Sur un élément `img`, l'attribut `src` indique l'**URL** de l'image (l'adresse où elle se trouve).

Voici un exemple d'élément `img` avec un attribut `src` qui pointe sur le logo de freeCodeCamp :

```html
<img src="https://cdn.freecodecamp.org/platform/universal/fcc_secondary.svg">
```

Dans l'élément `img` existant, ajoute un attribut `src` avec cette URL :

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`

# --hints--

Ton code doit toujours contenir un élément `img`. Tu l'as peut-être supprimé, ou tu as oublié les guillemets autour de la valeur de `src`.

```js
assert(document.querySelector('img'));
```

Ton élément `img` doit avoir un attribut `src`. Vérifie qu'il y a bien un espace entre le nom de l'élément et le nom de l'attribut.

```js
assert(document.querySelector('img').src);
```

L'attribut `src` de l'`img` doit valoir exactement `https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`. Attention aux majuscules/minuscules.

```js
assert.strictEqual(document.querySelector('img').src, 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg');
```

Même si la valeur de `src` est correcte, il est **fortement recommandé** d'entourer toujours la valeur d'un attribut par des guillemets.

```js
assert.notMatch(code, /\<img\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg/);
```

Ton élément `img` doit être fermé soit par `>` soit par `/>`.

```js
assert.match(code, /<img\s+src\s*=\s*("|')\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg\s*\1\s*\/?>/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img>
--fcc-editable-region--
    </main>
  </body>
</html>
```
