---
id: 5dc24165f86c76b9248c6ebc
title: Étape 9
challengeType: 0
dashedName: step-9
---

# --description--

Tous les éléments `img` doivent avoir un attribut `alt`. Le texte de l'attribut `alt` (pour *alternative*) est lu par les **lecteurs d'écran** pour aider les personnes malvoyantes, et il s'affiche aussi à la place de l'image si celle-ci ne peut pas être chargée.

Voici un exemple d'élément `img` avec un attribut `alt` :

```html
<img src="cat.jpg" alt="Un chat">
```

Dans ton élément `img`, ajoute un attribut `alt` avec ce texte :

`A cute orange cat lying on its back`

# --hints--

Ton code doit toujours contenir un élément `img`. Ne le supprime pas.

```js
assert(document.querySelector('img'));
```

Ton élément `img` n'a pas d'attribut `alt`. Vérifie qu'il y a bien un espace après le nom de l'élément et avant le nom de chaque attribut.

```js
assert(document.querySelector('img').hasAttribute('alt'));
```

La valeur de `alt` doit être exactement `A cute orange cat lying on its back`. Pense aux guillemets autour de la valeur.

```js
const altText = document
  .querySelector('img')
  .alt.toLowerCase()
  .replace(/\s+/g, ' ');
assert(altText.match(/A cute orange cat lying on its back\.?$/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg">
--fcc-editable-region--
    </main>
  </body>
</html>
```
