---
id: 5ddbd81294d8ddc1510a8e56
title: Étape 11
challengeType: 0
dashedName: step-11
---

# --description--

Le texte qui s'affiche pour un lien doit être placé **entre** la balise ouvrante et la balise fermante de l'élément ancre (`a`).

Voici un exemple de lien avec le texte `click here to go to freeCodeCamp.org` :

```html
<a href="https://www.freecodecamp.org">click here to go to freeCodeCamp.org</a>
```

Ajoute le texte `cat photos` à l'intérieur de ton élément `a`. Ce texte sera le texte cliquable du lien.

# --hints--

Ton élément `a` doit avoir une balise ouvrante.

```js
assert(document.querySelector('a'));
```

Ton élément `a` doit avoir une balise fermante.

```js
assert(code.match(/<\/a\>/));
```

Le texte de ton élément `a` doit être `cat photos`. Place le texte du lien entre la balise ouvrante et la balise fermante.

```js
assert.strictEqual(
  document.querySelector('a').innerText.toLowerCase().replace(/\s+/g, ' '),
    'cat photos'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com"></a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
