---
id: 5dfa3589eacea3f48c6300ae
title: Étape 20
challengeType: 0
dashedName: step-20
---

# --description--

À l'intérieur du **deuxième** élément `section`, ajoute un nouvel élément `h2` avec le texte `Cat Lists`.

# --hints--

Ton élément `section` doit avoir une balise ouvrante.

```js
assert(
  document.querySelectorAll('section').length === 2 &&
    code.match(/<\/section>/g).length === 2
);
```

Ton élément `h2` doit avoir une balise ouvrante.

```js
assert(document.querySelectorAll('h2').length === 2);
```

Ton élément `h2` doit avoir une balise fermante.

```js
assert(code.match(/<\/h2\>/g).length === 2);
```

Ton deuxième élément `h2` doit se trouver juste **au-dessus** de la balise fermante du deuxième `section`.

```js
const secondSection = document.querySelectorAll('section')[1];
assert(secondSection.lastElementChild.nodeName === 'H2');
```

Le texte du deuxième `h2` doit être `Cat Lists`. Vérifie l'orthographe.

```js
assert(
  document
    .querySelectorAll('main > section')[1]
    .lastElementChild.innerText.toLowerCase() === 'cat lists'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      <section>
        
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```
