---
id: 5dfa2407b521be39a3de7be1
title: Étape 15
challengeType: 0
dashedName: step-15
---

# --description--

Pour ouvrir un lien dans un **nouvel onglet**, on utilise l'attribut `target` sur l'élément `a`.

L'attribut `target` précise **où** ouvrir le document lié. La valeur `target="_blank"` ouvre le lien dans un nouvel onglet (ou une nouvelle fenêtre).

Voici la syntaxe générale :

```html
<a href="https://www.freecodecamp.org" target="_blank">freeCodeCamp</a>
```

Ajoute un attribut `target` avec la valeur `_blank` à la balise ouvrante de ton élément `a`, pour que le lien s'ouvre dans un nouvel onglet.

# --hints--

Ton élément `p` doit toujours contenir un élément `a` imbriqué avec le texte `cat photos`.

```js
const anchor = document.querySelectorAll('p > a');
assert.isNotEmpty(anchor);
assert.strictEqual(anchor[1]?.innerText?.toLowerCase().replace(/\s+/g, ' '), 'cat photos')
```

Ton élément `a` n'a pas d'attribut `target`. Vérifie qu'il y a bien un espace entre le nom de l'élément et le nom de l'attribut.

```js
assert.isTrue(document.querySelectorAll('a')[1]?.hasAttribute('target'));
```

La valeur de `target` doit être `_blank`. N'oublie pas les guillemets autour de la valeur.

```js
assert.strictEqual(document.querySelectorAll('a')[1]?.getAttribute('target'), '_blank');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
