---
id: 5dfb6250eacea3f48c6300b2
title: Étape 24
challengeType: 0
dashedName: step-24
---

# --description--

Après la liste non ordonnée, ajoute une **nouvelle image** avec l'attribut `src` qui pointe sur :

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`

Et l'attribut `alt` avec ce texte :

`A slice of lasagna on a plate.`

# --hints--

Il doit y avoir un élément `img` juste après la balise fermante `</ul>`.

```js
assert.equal(document.querySelectorAll('section')[1].querySelector('ul').nextElementSibling?.nodeName, 'IMG');
```

La nouvelle image n'a pas d'attribut `alt`. Vérifie qu'il y a bien un espace entre le nom de l'élément et chaque attribut.

```js
assert.isTrue(document.querySelectorAll('section')[1].querySelector('ul').nextElementSibling?.hasAttribute('alt'));
```

La valeur de `alt` de la nouvelle image doit être `A slice of lasagna on a plate.` Pense aux guillemets.

```js
assert.match(
  document.querySelectorAll('section')[1].querySelector('ul').nextElementSibling?.getAttribute('alt')?.replace(/\s+/g, ' '), /^A slice of lasagna on a plate\.?$/i 
);
```

La nouvelle image n'a pas d'attribut `src`.

```js
assert.isTrue(document.querySelectorAll('section')[1].querySelector('ul').nextElementSibling?.hasAttribute('src'));
```

La valeur de `src` doit être `https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`.

```js
assert.strictEqual(
  document.querySelectorAll('section')[1].querySelector('ul').nextElementSibling?.getAttribute('src'),
    'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

Pense à toujours entourer la valeur d'un attribut par des guillemets.

```js
assert.isTrue(!/\<img\s+.+\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/lasagna\.jpg/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
--fcc-editable-region--
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
