---
id: 5ef9b03c81a63668521804d0
title: Étape 27
challengeType: 0
dashedName: step-27
---

# --description--

Pour **mettre en valeur** un mot ou une phrase, on utilise l'élément `em` (pour *emphasis*). À l'affichage, le texte est généralement mis en italique.

Mets en valeur le mot `love` dans l'élément `figcaption` en l'entourant d'un élément `em`.

# --hints--

Ton élément `em` doit avoir une balise ouvrante.

```js
assert(document.querySelector('em'));
```

Ton élément `em` doit avoir une balise fermante.

```js
assert(code.match(/<\/em\>/));
```

Tu ne dois pas supprimer le `figcaption`.

```js
assert(document.querySelector('figcaption') && code.match(/<\/figcaption\>/));
```

Ton élément `em` doit entourer le texte `love`.

```js
assert(
  document.querySelector('figcaption > em').innerText.toLowerCase() === 'love'
);
```

Le texte du `figcaption` doit toujours être `Cats love lasagna`. Pense à laisser les espaces autour du `em`.

```js
assert(
  document
    .querySelector('figcaption')
    .innerText.replace(/\s+/gi, ' ')
    .match(/cats love lasagna\.?/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
--fcc-editable-region--
          <figcaption>Cats love lasagna.</figcaption>
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```
