---
id: 5ef9b03c81a63668521804d2
title: Étape 29
challengeType: 0
dashedName: step-29
---

# --description--

Le code d'une **liste ordonnée** (`ol` pour *ordered list*) ressemble à celui d'une liste non ordonnée, sauf que les éléments sont **numérotés** automatiquement à l'affichage (1, 2, 3, …).

Sous l'élément `h3`, ajoute une liste ordonnée avec ces trois éléments :

`flea treatment`
`thunder`
`other cats`

# --hints--

Ton élément `ol` doit avoir une balise ouvrante.

```js
assert(document.querySelector('ol'));
```

Ton élément `ol` doit avoir une balise fermante.

```js
assert(code.match(/<\/ol>/));
```

L'élément `ol` doit se trouver au-dessus de la balise fermante du deuxième `section`.

```js
assert.equal(document.querySelectorAll('main > section')?.[1]?.lastElementChild?.nodeName, 'OL');
```

Les trois `li` doivent être imbriqués dans l'élément `ol`.

```js
assert(
  [...document.querySelectorAll('li')].filter(
    (item) => item.parentNode.nodeName === 'OL'
  ).length === 3
);
```

Tu dois avoir trois `li` avec les textes `flea treatment`, `thunder` et `other cats` (peu importe l'ordre).

```js
assert.deepStrictEqual(
  [...document.querySelectorAll('li')]
    .filter((item) => item.parentNode.nodeName === 'OL')
    .map((item) => item.innerText.toLowerCase())
    .sort((a, b) => a.localeCompare(b)),
  ['flea treatment', 'other cats', 'thunder']
);
```

Tu ne dois avoir qu'**un seul** élément `ol`.

```js 
assert([...document.querySelectorAll('ol')].length == 1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
--fcc-editable-region--
        <h3>Top 3 things cats hate:</h3>
        
--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```
