---
id: 5f07c98cdb9413cbd4b16750
title: Étape 19
challengeType: 0
dashedName: step-19
---

# --description--

Il est temps d'ajouter une nouvelle section. Ajoute un **deuxième** élément `section` **en dessous** du `section` existant.

# --hints--

Ton élément `section` doit avoir une balise ouvrante.

```js
assert(document.querySelectorAll('section').length >= 2);
```

Tu ne dois ajouter qu'**une seule** balise ouvrante `<section>` supplémentaire.

```js
assert(document.querySelectorAll('section').length === 2);
```

Ton élément `section` doit avoir une balise fermante.

```js
assert(code.match(/<\/section>/g).length >= 2);
```

Tu ne dois ajouter qu'une seule balise fermante `</section>` supplémentaire.

```js
assert(code.match(/<\/section>/g).length === 2);
```

Le deuxième `section` ne doit **pas** être imbriqué dans le premier `section`.

```js
const childrenOf1stSection = [
  ...document.querySelector('main > section').children
];
const foundElems = childrenOf1stSection.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(foundElems.length === 0);
```

Les deux éléments `section` doivent être à l'intérieur du `main`.

```js
const childrenOfMain = [...document.querySelector('main').children];
const foundElems = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(foundElems.length === 2);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
--fcc-editable-region--
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
