---
id: 671141f948cbab359e74cc93
title: Étape 13
challengeType: 0
dashedName: step-13
---

# --description--

Ajoute des balises `p` autour de `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.` pour en faire un paragraphe.

# --hints--

Tu dois entourer `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.` par un élément `p`.

```js
const P = document.querySelectorAll('p')[1];
assert.exists(P);
assert.strictEqual(P.innerText.trim(), "See more cat photos in our gallery.");
```

Tu dois toujours avoir l'élément ancre à l'intérieur du nouveau paragraphe.

```js
const P = document.querySelectorAll('p')[1];
assert.exists(P.querySelector('a'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
