---
id: 671b6e873249bb35c9debfcf
title: Étape 12
challengeType: 0
dashedName: step-12
---

# --description--

Ajoute les mots `See more ` avant l'élément d'ancre et ` in our gallery` après l'élément d'ancre.

# --hints--

Tu dois utiliser l'élément d'ancre existant, ne crée pas un nouvel élément.

```js
assert.lengthOf(document.querySelectorAll('a'), 1);
```

Le texte à l'intérieur de l'élément d'ancre ne doit pas changer, il doit être `cat photos`.

```js
assert.strictEqual(document.querySelector('a')?.innerText?.replace(/\s+/g, ' ').trim(), "cat photos");
```

Tu dois avoir `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery` dans ton code.

```js
assert.match(code, /See more\s+<a href="https:\/\/freecatphotoapp\.com">cat photos<\/a>\s+in our gallery/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

