---
id: 66c057041df6394ca796bf33
title: Crée un générateur de phrases
challengeType: 26
dashedName: build-a-sentence-maker
---

# --description--

Dans ce lab, tu vas créer deux histoires différentes à partir d'un modèle de phrase. Tu utiliseras des variables pour stocker différentes parties de l'histoire, puis tu afficheras les histoires dans la console.

**Objectif :** Réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu devrais déclarer les variables suivantes avec `let` :
    - `adjective`
    - `noun`
    - `verb`
    - `place`
    - `adjective2`
    - `noun2`

1. Tu devrais affecter aux variables ci-dessus des valeurs de type chaîne de ton choix.

1. Tu devrais déclarer une variable `firstStory`.
1. Tu devrais utiliser le modèle d'histoire suivant pour créer la première histoire et l'affecter à la variable `firstStory` : `"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`;

1. Tu devrais afficher ta première histoire dans la console avec le message `"First story: [firstStory]"`.

1. Tu devrais affecter de nouvelles valeurs à tes variables `adjective`, `noun`, `verb`, `place`, `adjective2` et `noun2`.

1. Tu devrais déclarer une variable `secondStory`.

1. Crée une autre histoire à partir du même modèle et affecte-la à la variable `secondStory`.

1. Tu devrais afficher ta deuxième histoire dans la console avec le message `"Second story: [secondStory]"`.


# --hints--

Tu devrais déclarer une variable `adjective`.

```js
assert.isNotNull(adjective);
```

Tu devrais déclarer une variable `noun`.

```js
assert.isNotNull(noun);
```

Tu devrais déclarer une variable `verb`.

```js
assert.isNotNull(verb);
```

Tu devrais déclarer une variable `place`.

```js
assert.isNotNull(place);
```

Tu devrais déclarer une variable `adjective2`.

```js
assert.isNotNull(adjective2);
```

Tu devrais déclarer une variable `noun2`.

```js
assert.isNotNull(noun2);
```

Tu devrais affecter une valeur de type chaîne à la variable `adjective`.

```js
assert.isString(adjective);
```

Tu devrais affecter une valeur de type chaîne à la variable `noun`.

```js
assert.isString(noun);
```

Tu devrais affecter une valeur de type chaîne à la variable `verb`.

```js
assert.isString(verb);
```

Tu devrais affecter une valeur de type chaîne à la variable `place`.

```js
assert.isString(place);
```

Tu devrais affecter une valeur de type chaîne à la variable `adjective2`.

```js
assert.isString(adjective2);
```

Tu devrais affecter une valeur de type chaîne à la variable `noun2`.

```js
assert.isString(noun2);
```

Tu devrais déclarer une variable `firstStory`.

```js
assert.isNotNull(firstStory);
```

Tu devrais utiliser le bon format d'histoire pour la première histoire : `"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`. Fais attention aux espaces.

```js
const _initialValues = {}
for (const name of ['adjective', 'noun', 'noun2', 'place', 'verb', 'adjective2']) {
    const match = code.match(new RegExp(String.raw`${name}\s*=\s*('|"|${'`'})(?<${name}>.*?)\1(?:,|;|\s|$)`, 'm'));
    _initialValues[name] = match ? match.groups[name] : null;
}

const expected = `Once upon a time, there was a(n) ${_initialValues['adjective']} ${_initialValues['noun']} who loved to eat ${_initialValues['noun2']}. The ${_initialValues['noun']} lived in a ${_initialValues['place']} and had ${_initialValues['adjective2']} nostrils that blew fire when it was ${_initialValues['verb']}.`;

assert.strictEqual(firstStory, expected);
```

Tu devrais assembler ta première histoire en utilisant les variables que tu as déclarées, dans le bon ordre.

```js
assert.match(
  __helpers.removeJSComments(code),
  /firstStory\s*=\s*.*?adjective.*?noun.*?noun2.*?noun.*?place.*?adjective2.*?verb/
);
```

Tu devrais afficher ta première histoire avec le message `"First story: [firstStory]"`.

```js
const condition1 = /console\.log\(\s*["']First\s+story:\s+["']\s*\+\s*firstStory\s*\);?/gm.test(code);
const condition2 = /console\.log\(\s*`First\s+story:\s+\$\{firstStory\}`\s*\);?/gm.test(code);
assert.isTrue(condition1 || condition2);
```

Tu devrais déclarer une variable `secondStory`.

```js
assert.isNotNull(secondStory);
```

Tu devrais réaffecter la variable `adjective` pour la deuxième histoire.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/adjective\s*=\s*/g), 2);
```

Tu devrais réaffecter la variable `noun` pour la deuxième histoire.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/noun\s*=\s*/g), 2);
```

Tu devrais réaffecter la variable `verb` pour la deuxième histoire.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/verb\s*=\s*/g), 2);
```

Tu devrais réaffecter la variable `place` pour la deuxième histoire.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/place\s*=\s*/g), 2);
```

Tu devrais réaffecter la variable `adjective2` pour la deuxième histoire.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/adjective2\s*=\s*/g), 2);
```

Tu devrais réaffecter la variable `noun2` pour la deuxième histoire.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/noun2\s*=\s*/g), 2);
```

Tu devrais utiliser le bon format d'histoire pour la deuxième histoire : `"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`. Fais attention aux espaces.

```js
const expected = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;
assert.strictEqual(secondStory, expected);
```

Tu devrais assembler ta deuxième histoire en utilisant les variables que tu as déclarées, dans le bon ordre.

```js
assert.match(
  __helpers.removeJSComments(code),
  /secondStory\s*=\s*.*?adjective.*?noun.*?noun2.*?noun.*?place.*?adjective2.*?verb/
);
```

Tu devrais afficher ta deuxième histoire avec le format `"Second story: [secondStory]"`.

```js
const condition1 = /console\.log\(\s*["']Second\s+story:\s+["']\s*\+\s*secondStory\s*\);?/gm.test(code);
const condition2 = /console\.log\(\s*`Second\s+story:\s+\$\{secondStory\}`\s*\);?/gm.test(code);
assert.isTrue(condition1 || condition2);
```

`firstStory` ne devrait pas être identique à `secondStory`.

```js
assert.notEqual(firstStory, secondStory);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
let adjective = "funny";
let noun = "dragon";
let verb = "jumping";
let place = "garden";
let adjective2 = "sparkling";
let noun2 = "cakes";

const firstStory = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;

console.log("First story: " + firstStory);

adjective = "cute";
noun = "puppy";
verb = "barking";
place = "park";
adjective2 = "colorful";
noun2 = "flower";

const secondStory = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;

console.log("Second story: " + secondStory);
```
