---
id: 60f02e7361b68405e27b62a5
title: Étape 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Dans ce workshop, tu vas apprendre à styliser des formulaires en concevant un formulaire d'inscription.

Tout le code HTML de base (`DOCTYPE`, `html`, `head` et `body`) t'a déjà été fourni.

Dans le `body`, ajoute un titre pour donner du contexte au formulaire en utilisant un élément `h1` avec le texte `Registration Form`.

# --hints--

Tu dois ajouter le `h1` dans le `body`.

```js
assert.exists(document.querySelector('body > h1'));
```

Tu dois donner au `h1` le texte `Registration Form`.

```js
assert.equal(document.querySelector('body > h1')?.textContent, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
--fcc-editable-region--

--fcc-editable-region--
</body>
</html>
```
