---
id: 60f030d388cb74067cf291c3
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

Sous le titre, utilise le texte suivant dans un élément de paragraphe pour encourager les utilisateurs à s'inscrire :

```md
Please fill out this form with the required information
```

# --hints--

Tu dois ajouter un élément `p` dans le `body`.

```js
assert.exists(document.querySelector('body > p'));
```

Tu dois ajouter l'élément `p` sous le `h1`.

```js
assert.exists(document.querySelector('h1 + p'));
```

Tu dois donner à l'élément `p` le texte `Please fill out this form with the required information`.

```js
assert.equal(document.querySelector('p')?.innerText, 'Please fill out this form with the required information');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
--fcc-editable-region--
<body>
  <h1>Registration Form</h1>

</body>
--fcc-editable-region--
</html>
```
