---
id: 60f1a9cbd23023082e149fee
title: Étape 5
challengeType: 0
dashedName: step-5
---

# --description--

C'est mieux. Maintenant, rends l'arrière-plan plus agréable à lire en changeant le `background-color` de `body` en `#1b1b32`. Ensuite, pour voir le texte, change la `color` en `#f5f6f7`.

# --hints--

Tu dois ajouter le `background-color` dans le sélecteur d'élément `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor);
```

Tu dois donner à `background-color` la valeur `#1b1b32`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor, 'rgb(27, 27, 50)');
```

Tu dois ajouter la `color` dans le sélecteur d'élément `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.color);
```

Tu dois donner à `color` la valeur `#f5f6f7`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.color, 'rgb(245, 246, 247)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
--fcc-editable-region--

--fcc-editable-region--
}
```
