---
id: 60f1922fcbd2410527b3bd89
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

Dans les leçons précédentes, tu as appris à travailler avec les unités de viewport.

Souviens-toi que l'unité `vh` signifie viewport height, et qu'elle est égale à 1 % de la `height` du viewport. Elle est donc relative à la hauteur du viewport.

Il est temps d'améliorer le projet avec un peu de CSS. Commence par donner au `body` une `width` de `100%` et une `height` de `100vh`.

# --hints--

Tu dois utiliser le sélecteur d'élément `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

Tu dois donner au `body` une `width` de `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.width, '100%');
```

Tu dois donner au `body` une `height` de `100vh`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
