---
id: 60f5c3e399ff1a05629964e4
title: Étape 6
challengeType: 0
dashedName: step-6
---

# --description--

Comme le suggère le titre, tu crées un formulaire. Donc, après l'élément `p`, insère un `form` avec un attribut `action` ciblant `https://register-demo.freecodecamp.org`.

# --hints--

Tu dois ajouter un élément `form` adjacent à l'élément `p`.

```js
assert.exists(document.querySelector('p + form'));
```

Tu dois donner à `form` un attribut `action`.

```js
// Default action points to window location
assert.notEqual(document.querySelector('form')?.action, window?.location?.href);
```

Tu dois donner à `action` la valeur `https://register-demo.freecodecamp.org`.

```js
assert.equal(document.querySelector('form')?.action, 'https://register-demo.freecodecamp.org/');
```

Ton élément `form` doit avoir une balise de fermeture `</form>`.

```js
assert.match(code, /<\/form\>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
--fcc-editable-region--

--fcc-editable-region--
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
