---
id: 60f5d2776c854e069560fbe6
title: Étape 9
challengeType: 0
dashedName: step-9
---

# --description--

Le premier `fieldset` contiendra les champs de nom, d'e-mail et de mot de passe. Commence par ajouter quatre éléments `label` au premier `fieldset`.

# --hints--

Il doit y avoir trois éléments `fieldset` au total.

```js
assert.equal(document.querySelectorAll('fieldset')?.length, 3);
```

Les éléments `fieldset` doivent tous être des enfants directs de l'élément `form`.

```js
assert.equal(document.querySelectorAll('form > fieldset')?.length, 3);
```

Les quatre éléments `label` doivent tous être des enfants directs du premier élément `fieldset`. Assure-toi de fermer les éléments `label`.

```js
assert.equal(document.querySelectorAll('form fieldset:first-child > label')?.length, 4);
```

Les deux derniers éléments `fieldset` doivent être vides.

```js
const LastTwoFieldsetElements = Array.from(document.querySelectorAll("form fieldset:not(fieldset:first-child)"));
assert.isTrue(LastTwoFieldsetElements.filter((ele) => ele.innerHTML?.replace(/\s/g, ""))?.length === 0);
```

Assure-toi de fermer les éléments `label`.

```js
assert.lengthOf(document.querySelector('fieldset')?.innerHTML?.match(/<\/label>/g), 4)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--

--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
