---
id: 60f803d5241e6a0433a523a1
title: Étape 11
challengeType: 0
dashedName: step-11
---

# --description--

Dans les leçons précédentes, tu as appris à travailler avec les unités `rem`. Souviens-toi que l'unité `rem` signifie root `em`, et qu'elle est relative à la taille de police de l'élément `html`.

Comme les éléments `label` sont en ligne par défaut, ils s'affichent tous côte à côte sur la même ligne, ce qui rend leur texte difficile à lire.

Pour les afficher sur des lignes séparées, ajoute `display: block` à l'élément `label`, et ajoute une `margin` de `0.5rem 0` pour les séparer les uns des autres.

# --hints--

Tu dois utiliser le sélecteur `label`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('label'));
```

Tu dois ajouter une propriété `display` avec la valeur `block`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.display, 'block');
```

Tu dois ajouter une propriété `margin` avec la valeur `0.5rem 0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.margin, '0.5rem 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label>Enter Your First Name:</label>
      <label>Enter Your Last Name:</label>
      <label>Enter Your Email:</label>
      <label>Create a New Password:</label>
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--
```
