---
id: 60f81616cff80508badf9ad5
title: Étape 16
challengeType: 0
dashedName: step-16
---

# --description--

À ce stade, tu devrais pouvoir soumettre le formulaire. Cependant, tu remarqueras peut-être qu'il ne se passe pas grand-chose.

Pour rendre le formulaire plus interactif, ajoute l'attribut `required` aux éléments `input` du premier `fieldset`.

Maintenant, si tu essaies de soumettre le formulaire sans remplir les champs obligatoires, tu verras un message d'erreur.

# --hints--

Tu dois donner au premier élément `input` un attribut `required`.

```js
assert.isTrue(document.querySelector('input')?.required);
```

Tu dois donner au deuxième élément `input` un attribut `required`.

```js
assert.isTrue(document.querySelectorAll('input')?.[1]?.required);
```

Tu dois donner au troisième élément `input` un attribut `required`.

```js
assert.isTrue(document.querySelectorAll('input')?.[2]?.required);
```

Tu dois donner au quatrième élément `input` un attribut `required`.

```js
assert.isTrue(document.querySelectorAll('input')?.[3]?.required);
```

Tu ne dois pas donner d'attribut `required` à l'`input` dont le type est `submit`.

```js
assert.isFalse(document.querySelector('input[type="submit"]')?.required);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" /></label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
