---
id: 60f83e7bfc09900959f41e20
title: Étape 17
challengeType: 0
dashedName: step-17
---

# --description--

Certaines valeurs de l'attribut `type` incluent une validation de formulaire intégrée. Par exemple, `type="email"` exige que la valeur soit une adresse e-mail valide.

Ajoute une validation personnalisée à l'élément `input` du mot de passe en ajoutant un attribut `minlength` avec une valeur de `8`. Cela empêche la soumission des saisies de moins de 8 caractères.

# --hints--

Tu dois donner à l'élément `input` du mot de passe un attribut `minlength`.

```js
assert.notEqual(document.querySelector('input[type="password"]')?.minLength, -1);
```

Tu dois donner à l'attribut `minlength` une valeur de `8`.

```js
assert.equal(document.querySelector('input[type="password"]')?.minLength, 8);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
--fcc-editable-region--
      <label for="new-password">Create a New Password: <input id="new-password" type="password" required /></label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
