---
id: 60f84ec41116b209c280ba91
title: Étape 18
challengeType: 0
dashedName: step-18
---

# --description--

Avec `type="password"`, tu peux utiliser l'attribut `pattern` pour définir une expression régulière à laquelle le mot de passe doit correspondre pour être considéré comme valide.

Ajoute un attribut `pattern` à l'élément `input` du mot de passe pour exiger que la saisie corresponde à : `[a-z0-5]{8,}`

Ce qui précède est une expression régulière qui correspond à huit lettres minuscules ou plus, ou aux chiffres de `0` à `5`. Ensuite, supprime l'attribut `minlength` et essaie.

# --hints--

Tu dois donner à l'élément `input` du mot de passe un attribut `pattern`.

```js
assert.isNotEmpty(document.querySelector('input[type="password"]')?.pattern);
```

Tu dois donner à l'attribut `pattern` une valeur de `[a-z0-5]{8,}`.

```js
assert.equal(document.querySelector('input[type="password"]')?.pattern, '[a-z0-5]{8,}');
```

Tu dois supprimer l'attribut `minlength` de l'élément `input` du mot de passe.

```js
assert.equal(document.querySelector('input[type="password"]')?.minLength, -1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
--fcc-editable-region--
      <label for="new-password">Create a New Password: <input id="new-password" type="password" minlength="8" required /></label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
