---
id: 60f81167d0d4910809f88945
title: Étape 15
challengeType: 0
dashedName: step-15
---

# --description--

Le premier élément `input` avec un `type` de `submit` est automatiquement configuré pour soumettre l'élément `form` parent le plus proche.

Pour gérer la soumission du formulaire, ajoute après le dernier élément `fieldset` un élément `input` avec l'attribut `type` défini sur `submit` et l'attribut `value` défini sur `Submit`.

# --hints--

Tu dois ajouter l'élément `input` après le dernier élément `fieldset`.

```js
assert.exists(document.querySelectorAll('fieldset')?.[2]?.nextElementSibling?.tagName, 'input');
```

Tu dois donner à l'élément `input` un attribut `type` de `submit`.

```js
assert.exists(document.querySelector('fieldset + input[type="submit"]'));
```

Tu dois donner à l'élément `input` un attribut `value` de `Submit`.

```js
assert.exists(document.querySelector('fieldset + input[value="Submit"]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" /></label>
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
--fcc-editable-region--

--fcc-editable-region--
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
