---
id: 60fabf0dd4959805dbae09e6
title: Étape 29
challengeType: 0
dashedName: step-29
---

# --description--

Ajoute un autre `label` après le premier, avec le texte `Input your age (years): `. Ensuite, imbrique un `input` avec le `type` de `number`.

Ensuite, ajoute un attribut `min` à l'`input` avec une valeur de `13`, car les utilisateurs de moins de 13 ans ne devraient pas s'inscrire. De plus, les utilisateurs n'auront probablement pas plus de 120 ans ; ajoute un attribut `max` avec une valeur de `120`.

Maintenant, si quelqu'un essaie de soumettre le formulaire avec des valeurs en dehors de cette plage, un avertissement apparaîtra et le formulaire ne sera pas soumis. Essaie.

# --hints--

Tu dois ajouter un `label` au troisième `fieldset`, après le `label` existant.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label + label'));
```

Tu dois donner au `label` le texte `Input your age (years): `.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2)')?.textContent?.trim(), 'Input your age (years):');
```

Tu dois donner au `label` un `input` avec un `type` de `number`.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]'));
```

Tu dois donner à l'`input` un attribut `min` avec une valeur de `13`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]')?.min, '13');
```

Tu dois donner à l'`input` un attribut `max` avec une valeur de `120`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]')?.max, '120');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
      <legend>Account type (required)</legend>
      <label for="personal-account"><input id="personal-account" type="radio" name="account-type" value="personal" checked /> Personal</label>
      <label for="business-account"><input id="business-account" type="radio" name="account-type" value="business" /> Business</label>
    </fieldset>
    <fieldset>
      <label>Upload a profile picture: <input type="file" /></label>
--fcc-editable-region--

--fcc-editable-region--
    </fieldset>
    <a href="https://www.freecodecamp.org/news/terms-of-service">Read our terms and conditions</a>
    <label for="terms-and-conditions">
      <input id="terms-and-conditions" type="checkbox" required />
      I accept the terms and conditions
    </label>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
