---
id: 60fad1cafcde010995e15306
title: Étape 38
challengeType: 0
dashedName: step-38
---

# --description--

Lors des soumissions de formulaire, il est utile, et recommandé, de fournir à chaque élément pouvant être soumis un attribut `name`. Cet attribut sert à identifier l'élément dans la soumission du formulaire.

À l'exception des deux entrées `radio` (que tu as déjà nommées), donne à chaque élément pouvant être soumis un attribut `name` unique de ton choix.

# --hints--

Tu dois donner à l'`input` qui attend un prénom un attribut `name`. _PS J'aurais choisi `first-name`_

```js
assert.isNotEmpty(document.querySelector('fieldset:nth-child(1) > label:nth-child(1) > input')?.name);
```

Tu dois donner à l'`input` qui attend un nom de famille un attribut `name`. _PS J'aurais choisi `last-name`_

```js
assert.isNotEmpty(document.querySelector('fieldset:nth-child(1) > label:nth-child(2) > input')?.name);
```

Tu dois donner à l'`email` un attribut `name`. _PS J'aurais choisi `email`_

```js
assert.isNotEmpty(document.querySelector('input[type="email"]')?.name);
```

Tu dois donner au `password` un attribut `name`. _PS J'aurais choisi `password`_

```js
assert.isNotEmpty(document.querySelector('input[type="password"]')?.name);
```

Tu dois donner à la `checkbox` un attribut `name`. _PS J'aurais choisi `terms`_

```js
assert.isNotEmpty(document.querySelector('input[type="checkbox"]')?.name);
```

Tu dois donner au `file` un attribut `name`. _PS J'aurais choisi `file`_

```js
assert.isNotEmpty(document.querySelector('input[type="file"]')?.name);
```

Tu dois donner au `number` un attribut `name`. _PS J'aurais choisi `age`_

```js
assert.isNotEmpty(document.querySelector('input[type="number"]')?.name);
```

Tu dois donner à l'élément `select` un attribut `name`. _PS J'aurais choisi `referrer`_

```js
assert.isNotEmpty(document.querySelector('select')?.name);
```

Tu dois donner à l'élément `textarea` un attribut `name`. _PS J'aurais choisi `bio`_

```js
assert.isNotEmpty(document.querySelector('textarea')?.name);
```

Tu ne dois donner d'attribut `name` à aucun des éléments `option`.

```js
[...document.querySelectorAll('option')]?.forEach(option => assert.isUndefined(option?.name));
```

Tu ne dois donner d'attribut `name` à aucun des éléments `label`.

```js
[...document.querySelectorAll('label')]?.forEach(label => assert.isUndefined(label?.name));
```

Tu ne dois donner d'attribut `name` à aucun des éléments `fieldset`.

```js
[...document.querySelectorAll('fieldset')]?.forEach(fieldset => assert.isEmpty(fieldset?.name));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
      <legend>Account type (required)</legend>
      <label for="personal-account"><input id="personal-account" type="radio" name="account-type" value="personal" checked /> Personal</label>
      <label for="business-account"><input id="business-account" type="radio" name="account-type" value="business" /> Business</label>
    </fieldset>
    <fieldset>
      <label for="profile-picture">Upload a profile picture: <input id="profile-picture" type="file"/></label>
      <label for="age">Input your age (years): <input id="age" type="number" min="13" max="120" /></label>
      <label for="referrer">How did you hear about us?
        <select id="referrer">
          <option value="">(select one)</option>
          <option value="1">freeCodeCamp News</option>
          <option value="2">freeCodeCamp YouTube Channel</option>
          <option value="3">freeCodeCamp Forum</option>
          <option value="4">Other</option>
        </select>
      </label>
      <label for="bio">Provide a bio:
        <textarea id ="bio" rows="3" cols="30" placeholder="I like coding on the beach..."></textarea>
      </label>
    </fieldset>
    <a href="https://www.freecodecamp.org/news/terms-of-service">Read our terms and conditions</a>
    <label for="terms-and-conditions">
      <input id="terms-and-conditions" type="checkbox" required />
      I accept the terms and conditions
    </label>
--fcc-editable-region--
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
