---
id: 60fab4a123ce4b04526b082b
title: Étape 24
challengeType: 0
dashedName: step-24
---

# --description--

Actuellement, les utilisateurs peuvent soumettre le formulaire sans cocher les entrées radio. Même si tu as précédemment utilisé l'attribut `required` pour indiquer qu'une entrée est obligatoire, il ne fonctionnera pas dans ce cas, car ajouter `required` aux deux entrées transmettrait une mauvaise information aux utilisateurs.

Pour résoudre cela, tu peux fournir le contexte de ce qui est nécessaire en ajoutant un élément `legend` avec le texte `Account type (required)` avant les éléments `label` dans le deuxième `fieldset`. Ensuite, ajoute l'attribut `checked` à l'entrée `Personal` pour garantir que le formulaire est soumis avec les données requises.

# --hints--

Tes éléments `input` doivent rester de `type` `radio`.

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-child(2) input[type="radio"]'), 2);
```

Tu dois ajouter un élément `legend` dans le deuxième élément `fieldset`.

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-of-type(2) > legend'), 1);
```

Tu dois ajouter le texte `Account type (required)` à l'élément `legend`.

```js
assert.equal(document.querySelector('fieldset:nth-of-type(2) > legend')?.textContent, 'Account type (required)');
```

Tu dois donner un attribut `checked` à l'entrée `Personal`.

```js
assert.isTrue(document.querySelector('fieldset:nth-child(2) input[type="radio"]')?.hasAttribute('checked'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
--fcc-editable-region--

--fcc-editable-region--
      <label><input type="radio" name="account-type" value="personal" /> Personal</label>
      <label><input type="radio" name="account-type" value="business" /> Business</label>
    </fieldset>
    <fieldset></fieldset>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
