---
id: 62ff8e998d3e7eae14d6ae3b
title: Étape 26
challengeType: 0
dashedName: step-26
---

# --description--

Tu dois confirmer que l'utilisateur a lu les conditions générales.

Ajoute un élément `label`. Dans l'élément `label` nouvellement créé, ajoute un élément `input` et définis l'attribut `type` sur `checkbox`. Rends cet élément `input` obligatoire avec `required` afin que les utilisateurs ne puissent pas s'inscrire sans accepter les conditions générales.

Ajoute un attribut `id` et un attribut `for` avec la valeur `terms-and-conditions` aux éléments pour l'accessibilité.

# --hints--

Tu dois ajouter un `label` après le troisième élément `fieldset`.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label'));
```

Tu dois ajouter un `input` à l'élément `label`.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label > input'));
```

Tu dois ajouter un attribut `type` avec la valeur `checkbox` à l'élément `input`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.type, 'checkbox');
```

Tu dois ajouter un attribut `required` à l'élément `input`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.required, true);
```

L'élément `input` doit avoir un `id` de `terms-and-conditions`.

```js
assert.isTrue(document.querySelector('fieldset:nth-child(3) + label > input')?.matches('#terms-and-conditions'))
```

L'élément `label` doit avoir un attribut `for` avec une valeur de `terms-and-conditions`.

```js
assert.isTrue(document.querySelector('fieldset:nth-child(3) + label')?.matches('label[for="terms-and-conditions"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
      <legend>Account type (required)</legend>
      <label for="personal-account"><input id="personal-account" type="radio" name="account-type" value="personal" checked /> Personal</label>
      <label for="business-account"><input id="business-account" type="radio" name="account-type" value="business" /> Business</label>
    </fieldset>
    <fieldset></fieldset>
--fcc-editable-region--

--fcc-editable-region--
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
