---
id: 62ff8b9dab5ac88e4d3d43a3
title: Étape 13
challengeType: 0
dashedName: step-13
---

# --description--

En suivant les bonnes pratiques d'accessibilité, relie les éléments `input` et les éléments `label` ensemble avec l'attribut `for`.

Utilise `first-name`, `last-name`, `email` et `new-password` comme valeurs pour les attributs `id` respectifs.

# --hints--

Le premier élément `input` doit avoir un `id` de `first-name`.

```js
assert.isTrue(document.querySelectorAll('input')?.[0]?.matches('#first-name'))
```

Le deuxième élément `input` doit avoir un `id` de `last-name`.

```js
assert.isTrue(document.querySelectorAll('input')?.[1]?.matches('#last-name'))
```

Le troisième élément `input` doit avoir un `id` de `email`.

```js
assert.isTrue(document.querySelectorAll('input')?.[2]?.matches('#email'))
```

Le quatrième élément `input` doit avoir un `id` de `new-password`.

```js
assert.isTrue(document.querySelectorAll('input')?.[3]?.matches('#new-password'))
```

Le premier élément `label` doit avoir un attribut `for` avec une valeur de `first-name`.

```js
assert.isTrue(document.querySelectorAll('label')?.[0]?.matches('label[for="first-name"]'))
```

Le deuxième élément `label` doit avoir un attribut `for` avec une valeur de `last-name`.

```js
assert.isTrue(document.querySelectorAll('label')?.[1]?.matches('label[for="last-name"]'))
```

Le troisième élément `label` doit avoir un attribut `for` avec une valeur de `email`.

```js
assert.isTrue(document.querySelectorAll('label')?.[2]?.matches('label[for="email"]'))
```

Le quatrième élément `label` doit avoir un attribut `for` avec une valeur de `new-password`.

```js
assert.isTrue(document.querySelectorAll('label')?.[3]?.matches('label[for="new-password"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      <label>Enter Your First Name: <input /></label>
      <label>Enter Your Last Name: <input /></label>
      <label>Enter Your Email: <input /></label>
      <label>Create a New Password: <input /></label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
