---
id: 63541ef4f96cd82e8e6c788a
title: Étape 7
challengeType: 0
dashedName: step-7
---

# --description--

L'attribut `method` précise comment envoyer les données du formulaire vers l'URL indiquée dans l'attribut `action`. Les données du formulaire peuvent être envoyées via une requête `GET` sous forme de paramètres d'URL (avec `method="get"`) ou via une requête `POST` sous forme de données dans le corps de la requête (avec `method="post"`).

Définis l'attribut `method` pour envoyer les données de ton formulaire via une requête `POST`.

**NOTE** : Tu en apprendras plus sur les méthodes `GET` et `POST` dans les prochaines leçons et workshops.

# --hints--

Tu ne dois pas ajouter de nouvel élément `form`.

```js
assert.equal(document.querySelectorAll('form').length, 1 )
```

Ton élément `form` doit avoir un attribut `method`.

```js
assert.exists(document.querySelector('form')?.getAttribute('method'));
```

Ton attribut `method` doit être défini sur `post`.

```js
assert.match(document.querySelector('form')?.getAttribute('method'), /^post$/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
--fcc-editable-region--
  <form action='https://register-demo.freecodecamp.org'></form>
--fcc-editable-region--
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
