---
id: 66c63bfa4685e4e3843fa0dc
title: Étape 1
challengeType: 1
dashedName: step-1
---

# --description--

Dans ce workshop, tu continueras d'apprendre les tableaux en construisant une liste de courses.

Commence par ajouter un `console.log` qui affiche la chaîne `"Grocery shopping list"` dans la console.

# --hints--

Tu dois avoir un `console.log` dans ton code.

```js
assert.match(code, /console\.log(.*)/);
```

Ton `console.log` doit afficher le message `"Grocery shopping list"`. Vérifie bien les fautes d'orthographe et de casse.

```js
assert.match(code, /console\.log\((["'])(Grocery\s+shopping\s+list)\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
