---
id: 66c644c08ebcd2ef90c750bd
title: Étape 4
challengeType: 1
dashedName: step-4
---

# --description--

Dans les leçons précédentes, tu as appris à ajouter des éléments à la fin d'un tableau en utilisant la méthode `push` comme ceci :

```js
fruits.push('pear');
```

En utilisant la méthode `push`, ajoute la chaîne `"Apples"` au tableau `shoppingList`.

# --hints--

Ton tableau `shoppingList` doit contenir un élément.

```js
assert.lengthOf(shoppingList, 1);
```

Ton tableau `shoppingList` doit contenir la chaîne `"Apples"`. Vérifie bien les fautes d'orthographe et de casse.

```js
assert.includeMembers(shoppingList, ['Apples']);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

--fcc-editable-region--

--fcc-editable-region--
```
