---
id: 66c64ad16796c7f2419b45c5
title: Étape 8
challengeType: 1
dashedName: step-8
---

# --description--

Maintenant, il est temps de commencer à ajouter des articles en haut de la liste de courses.

Commence par ajouter un `console.log` qui affiche le message `"It looks like we need to get some cooking oil."` dans la console.

# --hints--

Tu dois afficher le message `"It looks like we need to get some cooking oil."` dans la console.

```js
assert.match(code, /console\.log\((['"])(It\s+looks\s+like\s+we\s+need\s+to\s+get\s+some\s+cooking\s+oil\.)\1\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
