---
id: 66c64095f61166ede6e39a84
title: Étape 2
challengeType: 1
dashedName: step-2
---

# --description--

Pour cette liste de courses, tu utiliseras un tableau pour représenter les articles que tu dois acheter.

Dans les leçons précédentes, tu as appris à créer des tableaux comme ceci :

```js
const fruits = ['apple', 'banana', 'orange'];
```

Dans cette étape, crée une variable appelée `shoppingList` et affecte-lui un tableau vide.

# --hints--

Tu dois avoir une variable appelée `shoppingList`.

```js
assert.isNotNull(shoppingList);
```

Tu dois affecter un tableau à la variable `shoppingList`.

```js
assert.isArray(shoppingList);
```

Ton tableau doit être vide.

```js
assert.isEmpty(shoppingList);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

--fcc-editable-region--

--fcc-editable-region--
```
