---
id: 66c6491fe4c8e0f16845425f
title: Étape 7
challengeType: 1
dashedName: step-7
---

# --description--

Maintenant, il est temps d'ajouter un autre fruit à la liste.

En utilisant la même méthode de tableau qu'auparavant, ajoute la chaîne `"Grapes"` à la fin du tableau `shoppingList`.

Puis, ajoute un `console.log` et appelle la fonction `getShoppingListMsg` avec le tableau `shoppingList` comme argument à l'intérieur de `console.log` pour voir la liste mise à jour affichée dans la console.

# --hints--

Ton tableau `shoppingList` doit maintenant contenir deux éléments.

```js
assert.lengthOf(shoppingList, 2);
```

Ton tableau `shoppingList` doit contenir la chaîne `"Grapes"`.

```js
assert.equal(shoppingList[1], "Grapes");
```

Tu dois appeler la fonction `getShoppingListMsg(shoppingList)` à l'intérieur du `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 2);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
