---
id: 66c645b345a39bf04864dc50
title: Étape 5
challengeType: 1
dashedName: step-5
---

# --description--

Pour voir la liste de courses mise à jour, tu devras afficher la liste de courses actuelle et un court message dans la console.

Comme ce message sera utilisé de façon répétée tout au long du workshop, il est préférable de créer une fonction réutilisable.

Crée une fonction appelée `getShoppingListMsg` qui prend un tableau en paramètre et retourne la chaîne `"Current Shopping List: "` suivie du contenu du tableau fourni.

Tu peux utiliser des littéraux de gabarit (template literals) ou la concaténation de chaînes avec l'opérateur `+` pour combiner la chaîne avec le tableau passé à la fonction.

# --hints--

Tu dois avoir une fonction appelée `getShoppingListMsg` dans ton code.

```js
assert.isFunction(getShoppingListMsg);
```

Ta fonction `getShoppingListMsg` doit retourner une chaîne.

```js
assert.isString(getShoppingListMsg(shoppingList));
```

Ta fonction `getShoppingListMsg` doit accepter un tableau et retourner le message `"Current Shopping List: "` suivi du contenu de ce tableau. Vérifie bien les espaces et les fautes d'orthographe.

```js
assert.oneOf(getShoppingListMsg(["Anchovies", "Olives"]), ["Current Shopping List: Anchovies,Olives", "Current Shopping List: Anchovies, Olives"]);
assert.strictEqual(getShoppingListMsg(["Apples"]), "Current Shopping List: Apples");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

--fcc-editable-region--

--fcc-editable-region--
```
