---
id: 66c64210413532ee9d3bd342
title: Étape 3
challengeType: 1
dashedName: step-3
---

# --description--

Pour cette prochaine partie du workshop, tu t'entraîneras à ajouter des aliments à la liste de courses.

Commence par utiliser `console.log` pour afficher le message `"It will be nice to have some fruit to eat."`.

# --hints--

Tu dois afficher le message `"It will be nice to have some fruit to eat."` dans la console.

```js
assert.match(code, /console\.log\s*\(\s*(['"])(It\s+will\s+be\s+nice\s+to\s+have\s+some\s+fruit\s+to\s+eat\.)\1\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

--fcc-editable-region--

--fcc-editable-region--
```
