---
id: 66c726c34ecf1e238aa9d7d5
title: Étape 9
challengeType: 1
dashedName: step-9
---

# --description--

Dans les leçons précédentes, tu as appris à ajouter des éléments au début d'un tableau en utilisant la méthode `unshift()`.

Voici un rappel de comment utiliser la méthode `unshift()` :

```js
array.unshift(item1, item2, ..., itemX);
```

Utilise la méthode `unshift()` pour ajouter la chaîne `"Vegetable Oil"` au début du tableau `shoppingList`.

# --hints--

Tu dois avoir trois éléments dans ton tableau `shoppingList`.

```js
assert.lengthOf(shoppingList, 3);
```

Tu dois utiliser la méthode `unshift()` pour ajouter `"Vegetable Oil"` au début du tableau `shoppingList`.

```js
assert.equal(shoppingList[0], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");
function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

--fcc-editable-region--

--fcc-editable-region--
```
