---
id: 66c73a7798f6f62b2ae58f22
title: Étape 16
challengeType: 1
dashedName: step-16
---

# --description--

Maintenant, il est temps d'ajouter d'autres articles au début de la liste de courses.

Commence par ajouter une instruction `console.log` qui affiche le message `"It might be nice to get a dessert."`

En dessous de cette instruction `console`, utilise la bonne méthode de tableau pour ajouter la chaîne `"Chocolate Cake"` au début du tableau `shoppingList`.

Enfin, ajoute un `console.log` et appelle la fonction `getShoppingListMsg` avec le tableau `shoppingList` comme argument à l'intérieur de `console.log` pour voir la liste mise à jour affichée dans la console.

# --hints--

Tu dois avoir un `console.log` qui affiche le message `"It might be nice to get a dessert."`.

```js
assert.match(code, /console\.log\((['"])(It\s+might\s+be\s+nice\s+to\s+get\s+a\s+dessert\.)\1\)/);
```

Ton tableau `shoppingList` doit contenir un total de `6` éléments.

```js
assert.lengthOf(shoppingList, 6);
```

Tu dois utiliser la méthode `unshift` pour ajouter la chaîne `"Chocolate Cake"` au début du tableau `shoppingList`.

```js
assert.strictEqual(shoppingList[0], "Chocolate Cake");
```

Tu dois appeler `getShoppingListMsg(shoppingList)` à l'intérieur du `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 6);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
