---
id: 66c72a55418cc9247b710827
title: Étape 10
challengeType: 1
dashedName: step-10
---

# --description--

Ensuite, ajoute un `console.log` et appelle la fonction `getShoppingListMsg` avec le tableau `shoppingList` comme argument à l'intérieur de `console.log` pour voir la liste mise à jour affichée dans la console.

# --hints--

Tu dois appeler `getShoppingListMsg(shoppingList)` à l'intérieur du `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 3);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");

--fcc-editable-region--

--fcc-editable-region--
```
