---
id: 66c72b0ffbc5522525768558
title: Étape 11
challengeType: 1
dashedName: step-11
---

# --description--

Dans les étapes précédentes, tu as révisé comment ajouter un élément à la fin du tableau en utilisant la méthode `push`.

Mais la méthode `push` accepte plusieurs arguments, donc tu peux ajouter plusieurs éléments à la fin du tableau comme ceci :

```js
array.push(item1, item2, item3);
```

Dans cette étape, utilise la méthode `push` pour ajouter les chaînes `"Popcorn", "Beef Jerky", "Potato Chips"` au tableau `shoppingList`.

L'ordre est important, alors veille à ajouter les éléments dans l'ordre où ils sont listés.

# --hints--

Ton tableau `shoppingList` doit contenir un total de `6` chaînes.

```js
assert.lengthOf(shoppingList, 6);
```

Ton tableau `shoppingList` doit contenir la chaîne `"Popcorn"`.

```js
assert.strictEqual(shoppingList[3], "Popcorn");
```

Ton tableau `shoppingList` doit contenir la chaîne `"Beef Jerky"`.

```js
assert.strictEqual(shoppingList[4], "Beef Jerky");
```

Ton tableau `shoppingList` doit contenir la chaîne `"Potato Chips"`.

```js
assert.strictEqual(shoppingList[5], "Potato Chips");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
