---
id: 66c730ee6ae076281721d0b9
title: Étape 14
challengeType: 1
dashedName: step-14
---

# --description--

Dans les leçons précédentes, tu as appris à supprimer des éléments à la fin d'un tableau en utilisant la méthode `pop`.

Voici un rappel de comment utiliser la méthode `pop` :

```js
let array = [1, 2, 3, 4, 5];
array.pop();

// [1, 2, 3, 4]
console.log(array);
```

Utilise la méthode `pop` pour supprimer le dernier élément du tableau `shoppingList`.

# --hints--

Ton tableau `shoppingList` doit maintenant contenir `5` éléments.

```js
assert.lengthOf(shoppingList, 5);
```

Ton tableau `shoppingList` ne doit plus contenir la chaîne `"Potato Chips"`.

```js
assert.notInclude(shoppingList, "Potato Chips");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

--fcc-editable-region--

--fcc-editable-region--
```
