---
id: 66c72f4d0528bd268a82107b
title: Étape 12
challengeType: 1
dashedName: step-12
---

# --description--

Maintenant, il est temps d'afficher le tableau `shoppingList` mis à jour dans la console.

Ajoute un autre `console.log` et appelle la fonction `getShoppingListMsg` avec le tableau `shoppingList` comme argument à l'intérieur de `console.log` pour voir la liste mise à jour affichée dans la console.

# --hints--

Tu dois appeler `getShoppingListMsg(shoppingList)` à l'intérieur du `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 4);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");

--fcc-editable-region--

--fcc-editable-region--
```
