---
id: 66cbe2319d3845545a293a0b
title: Étape 6
challengeType: 1
dashedName: step-6
---

# --description--

Maintenant, il est temps de voir le message affiché dans la console.

Ajoute un `console.log` et appelle la fonction `getShoppingListMsg` avec le tableau `shoppingList` comme argument à l'intérieur de `console.log` pour voir le message affiché dans la console.

# --hints--

Tu dois appeler la fonction `getShoppingListMsg(shoppingList)` à l'intérieur du `console.log`.

```js
assert.match(code, /console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

--fcc-editable-region--

--fcc-editable-region--
```
