---
id: 6808fe91a5cb0c3355f23027
title: Étape 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Dans cet atelier, tu vas apprendre à travailler avec les formulaires en React en construisant un formulaire de candidature de super-héros.

Un peu de code de base a été fourni, tu n'as qu'à créer le reste du composant React.

Commence par renvoyer un élément `div` avec un `className` de `form-wrap`. Dans le `div`, crée un élément `h2` avec le texte `Superhero Application Form`. En dessous, crée un élément `p` avec le texte `Please complete all fields`.

# --hints--

Tu dois renvoyer un élément `div` avec le `className` de `form-wrap`.

```js
assert.exists(document?.querySelector("div.form-wrap"));
```

Tu dois créer un élément `h2` à l'intérieur de ton élément `div`.

```js
assert.exists(document?.querySelector(".form-wrap > h2"));
```

Ton élément `h2` doit avoir le texte `Superhero Application Form`.

```js
const h2El = document?.querySelector("h2")
assert.equal(h2El?.textContent, "Superhero Application Form");
```

Tu dois créer un élément `p` à l'intérieur de ton élément `div`. Assure-toi qu'il vient après le `h2`.

```js
assert.exists(document?.querySelector(".form-wrap > h2 + p"));
```

Ton élément `p` doit avoir le texte `Please complete all fields`.

```js
const pEl = document?.querySelector("p")
assert.equal(pEl?.textContent, "Please complete all fields");
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Superhero Application Form</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { SuperheroForm } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<SuperheroForm />);
    </script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  min-height: 100vh;
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  background: linear-gradient(30deg, #ff9999 50%, #6699ff 50%)
}

.form-wrap {
  background-color: white;
  width: 400px;
  padding: 20px;
  border: 1px solid black;
  box-shadow: 5px 5px 10px black;
}

.form-wrap h2,
.form-wrap p {
  text-align: center;
}

.form-wrap p {
  position: relative;
  top: -18px;
}

.section {
  display: flex;
  margin-bottom: 30px;
}

.column {
  flex-direction: column;
}

.submit-wrap {
  text-align: center;
}

.submit-btn {
  display: block;
  margin: 0 auto;
  padding: 0.4rem 0.5rem;
  border: 1px solid black
}

.submit-btn:hover {
  cursor: pointer;
  background-color: #f3f3f3;
}

.submit-btn:disabled {
  cursor: not-allowed;
}
```

```jsx
const { useState } = React;

export const SuperheroForm = () => {

  --fcc-editable-region--

  --fcc-editable-region--

};
```
