---
id: 69860157f7fce39f9df399a3
title: 'Implémenter l''algorithme des N-reines'
challengeType: 26
dashedName: implement-the-n-queens-algorithm
---

# --description--

Le problème des N-reines te demande de placer N reines sur un échiquier N×N de sorte qu'aucune paire de reines ne s'attaque (aucune ne partage une ligne, une colonne ou une diagonale).

Par exemple, s'il y a un plateau 4x4, un arrangement valide est :

```md
[1, 3, 0, 2]
```

Cela signifie qu'en ligne 0, la reine est placée en colonne 1 ; en ligne 1, la reine est placée en colonne 3 ; en ligne 2, la reine est placée en colonne 0 ; et en ligne 3, la reine est placée en colonne 2.

Visuellement, cet arrangement ressemble à :

```md
. Q . .
. . . Q
Q . . .
. . Q .
```

Où `Q` représente une reine et `.` représente une case vide.

Dans ce labo, tu vas implémenter le solveur du problème des N-reines avec l'approche du parcours en profondeur.


**Objectif** : Remplis les user stories ci-dessous et fais passer tous les tests pour terminer le labo.

**User Stories :**

1. Tu dois avoir une fonction nommée `dfsNQueens`.
2. La fonction doit accepter exactement un argument : un entier `n`.
3. Si `n` est inférieur à `1`, la fonction doit renvoyer un tableau vide (`[]`).
4. La fonction doit renvoyer une liste de solutions ; chaque solution est elle-même une liste de longueur `n`, où l'élément à l'indice `i` est l'indice de colonne (base 0) de la reine en ligne `i`.

# --hints--

Tu dois avoir une fonction nommée `dfsNQueens` qui prend un argument.

```js
assert.isFunction(dfsNQueens);
assert.lengthOf(dfsNQueens, 1);
```

Si `n` est inférieur à `1`, la fonction doit renvoyer un tableau vide.

```js
assert.sameDeepOrderedMembers(dfsNQueens(0), []);
assert.sameDeepOrderedMembers(dfsNQueens(-1), []);
assert.sameDeepOrderedMembers(dfsNQueens(-5), []);
```

La fonction doit renvoyer un tableau de solutions, où chaque solution est un tableau de longueur `n`.

```js
const output = dfsNQueens(4);
assert.isArray(output);
output.forEach(solution => {
    assert.isArray(solution);
    assert.lengthOf(solution, 4);
})
```

`dfsNQueens(1)` doit renvoyer `[[0]]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(1), [[0]]);
```

`dfsNQueens(2)` doit renvoyer `[]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(2), []);
```

`dfsNQueens(3)` doit renvoyer `[]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(3), []);
```

`dfsNQueens(4)` doit renvoyer `[[1, 3, 0, 2], [2, 0, 3, 1]]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(4), [[1, 3, 0, 2], [2, 0, 3, 1]]);
```

`dfsNQueens(5)` doit renvoyer `[[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(5), [[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]);
```

`dfsNQueens(5).length` doit valoir `10`.

```js
assert.lengthOf(dfsNQueens(5), 10);
```

`dfsNQueens(8).length` doit valoir `92`.

```js
assert.lengthOf(dfsNQueens(8), 92);
```

`dfsNQueens` doit renvoyer le résultat correct.

```js
assert.lengthOf(dfsNQueens(7), 40);
assert.sameDeepOrderedMembers(dfsNQueens(6), [[1, 3, 5, 0, 2, 4], [2, 5, 1, 4, 0, 3], [3, 0, 4, 1, 5, 2], [4, 2, 0, 5, 3, 1]]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dfsNQueens(n) {
  if (n < 1) return [];

  const results = [];

  function isSafe(queens, row, col) {
    for (let r = 0; r < queens.length; r++) {
      const c = queens[r];
      if (c === col || Math.abs(row - r) === Math.abs(col - c)) {
        return false;
      }
    }
    return true;
  }

  function dfs(queens) {
    const row = queens.length;
    if (row === n) {
      results.push(queens.slice());
      return;
    }

    for (let col = 0; col < n; col++) {
      if (isSafe(queens, row, col)) {
        queens.push(col);
        dfs(queens);
        queens.pop();
      }
    }
  }

  dfs([]);
  return results;
}
```
