---
id: 669fd6fd12918e3de87854d4
title: Étape 4
challengeType: 0
dashedName: step-4
---

# --description--

Pour ton blog, il faudrait que les utilisateurs puissent naviguer vers les différentes sections de la page.

L'élément `nav` est utilisé pour fournir des liens de navigation vers d'autres sections du document ou d'autres sections du site web. Souvent, tu verras l'élément `nav` utilisé pour les menus ou les sommaires. 

Voici un exemple qui utilise l'élément `nav` :

```html
<nav>
  <ul>
    <li><a href="#">Home</a></li>
    <li><a href="#about">About</a></li>
    <li><a href="#contact">Contact</a></li>
  </ul>
</nav>
```

Sous ton élément `figure`, ajoute un élément `nav` avec un élément `ul` imbriqué à l'intérieur.

À l'intérieur de l'élément `ul`, ajoute trois éléments `li`.

# --hints--

Ton élément `nav` devrait avoir une balise ouvrante. 

```js
assert.match(code, /<nav\>/);
```

Ton élément `nav` devrait avoir une balise fermante. 

```js
assert.match(code, /<\/nav\>/);
```

Ton élément `nav` devrait être à l'intérieur de ton élément `header`.

```js
const navElement = document.querySelector('nav');
assert.strictEqual(navElement?.parentElement.tagName, 'HEADER');
```

Ton élément `ul` devrait avoir une balise ouvrante. 

```js
assert.match(code, /<ul\>/);
```

Ton élément `ul` devrait avoir une balise fermante. 

```js
assert.match(code, /<\/ul\>/);
```

Ton élément `ul` devrait être à l'intérieur de ton élément `nav`.

```js
const ulElement = document.querySelector('ul');
assert.strictEqual(ulElement?.parentElement.tagName, 'NAV');
```

Ton élément `ul` devrait contenir trois éléments `li`.

```js
const ulElement = document.querySelector("ul");

const liElements = ulElement?.querySelectorAll("li");
assert.strictEqual(liElements?.length, 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
--fcc-editable-region--
      
--fcc-editable-region--
    </header>
  </body>
</html>
```
