---
id: 669fdd3965f36f40af9615e5
title: Étape 6
challengeType: 0
dashedName: step-6
---

# --description--

Maintenant que tu as fini de construire l'en-tête de la page, tu vas commencer à ajouter ton contenu principal. 

Sous ton `header`, ajoute un élément `main`. 

# --hints--

Ton élément `main` devrait avoir une balise ouvrante. 

```js
assert.match(code, /<main\>/);
```

Ton élément `main` devrait avoir une balise fermante. 

```js
assert.match(code, /<\/main\>/);
```

Ton élément `main` devrait être sous ton élément `header`.

```js
assert.exists(document.querySelector('header + main'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>
```
