---
id: 669fd003cb89ee3c2402e041
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

Dans ce contenu d'introduction, tu veux afficher une image de M. Moustaches avec une légende.

Sous l'élément `h1`, commence par ajouter un élément `figure`.

À l'intérieur de l'élément `figure`, ajoute un élément `img`.

L'attribut `src` du `img` devrait avoir la valeur `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"` et le texte `alt` devrait avoir la valeur `"a cat in the garden"`.

Sous ton élément `img`, ajoute un `figcaption` avec le texte `Mr. Whiskers in the Garden`.

# --hints--

Ton élément `figure` devrait avoir une balise ouvrante. 

```js
assert.match(code, /<figure\>/);
```

Ton élément `figure` devrait avoir une balise fermante. 

```js
assert.match(code, /<\/figure\>/);
```

Ton élément `figure` devrait être à l'intérieur de ton élément `header`.

```js
const figureElement = document.querySelector('figure');
assert.equal(figureElement?.parentElement.tagName, 'HEADER');
```

Tu devrais avoir un élément `img`.

```js
assert.exists(document.querySelector("img"));
```

Ton élément `img` devrait être à l'intérieur de ton élément `figure`.

```js
const imgElement = document.querySelector('img');
assert.strictEqual(imgElement?.parentElement.tagName, 'FIGURE');
```

Ton élément `img` devrait avoir un attribut `src` mis à `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"`.

```js
const imgElement = document.querySelector('img');
const src = imgElement?.getAttribute('src');

assert.strictEqual(src, 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

Ton élément `img` devrait avoir un attribut `alt` mis à `"a cat in the garden"`.

```js
const imgElement = document.querySelector('img');
const alt = imgElement?.getAttribute('alt');

assert.strictEqual(alt, 'a cat in the garden');
```

Ton élément `figcaption` devrait avoir une balise ouvrante. 

```js
assert.match(code, /<figcaption\>/);
```

Ton élément `figcaption` devrait avoir une balise fermante. 

```js
assert.match(code, /<\/figcaption\>/);
```

Ton élément `figcaption` devrait être à l'intérieur de ton élément `figure`.

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.parentElement.tagName, 'FIGURE');
```

Ton élément `figcaption` devrait avoir le texte `Mr. Whiskers in the Garden`.

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.innerText.trim(), "Mr. Whiskers in the Garden")
```

Ton `figcaption` devrait être sous l'élément `img`. Tu les as dans le mauvais ordre.

```js
const figureElement = document.querySelector('figure');
assert.strictEqual(figureElement?.lastElementChild?.tagName, "FIGCAPTION")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
--fcc-editable-region--
      
--fcc-editable-region--
    </header>
  </body>
</html>
```
