---
id: 669fdfc9a5e3da42d2376609
title: Étape 9
challengeType: 0
dashedName: step-9
---

# --description--

Maintenant que tu as ajouté la section « à propos », essaie de cliquer sur le lien `About` pour voir la page sauter vers cette section.

La prochaine section dans la page du blog sera une liste de posts qui parlent de M. Moustaches.

Ajoute un autre élément `section` avec un `id` mis à `"posts"`.

À l'intérieur de l'élément `section`, ajoute un élément `h2` avec le texte `Posts`.

# --hints--

Ton code devrait avoir une balise ouvrante `section`.

```js
assert.match(code, /<section.*>/i);
```

Ton code devrait avoir une balise fermante `section`.

```js
assert.match(code, /<\/section>/i);
```

Ton élément `section` devrait avoir un `id` mis à `"posts"`.

```js
assert.match(code, /<section.*id="posts".*>/i);
```

Ton élément `section` devrait être à l'intérieur de ton élément `main`.

```js
const sectionElement = document.getElementById('posts');
assert.strictEqual(sectionElement?.parentElement.tagName, 'MAIN');
```

Ton élément `h2` devrait être à l'intérieur de l'élément `section`.

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.parentElement.tagName, 'SECTION');
```

Ton élément `h2` devrait avoir le texte `Posts`.

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.innerText.trim(), 'Posts');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </main>
  </body>
</html>
```
