---
id: 669fde2081f65141ad703fe4
title: Étape 7
challengeType: 0
dashedName: step-7
---

# --description--

La première section sur la page sera la section « à propos ». La section présentera M. Moustaches et donnera aux utilisateurs une idée du sujet de ce blog. 

À l'intérieur de ton élément `main`, ajoute un élément `section` avec l'attribut `id` mis à `"about"`. 

À l'intérieur de l'élément `section`, ajoute un `h2` avec le texte `About`.

# --hints--

Ton élément `section` devrait avoir une balise ouvrante. 

```js
assert.match(code, /<section.*>/);
```

Ton élément `section` devrait avoir une balise fermante. 

```js
assert.match(code, /<\/section\>/);
```

Ton élément `section` devrait avoir un `id` mis à `"about"`.

```js
const sectionElement = document.querySelector('section');
assert.strictEqual(sectionElement?.getAttribute("id"), "about");
```

Ton élément `section` devrait être à l'intérieur de ton élément `main`.

```js
const sectionElement = document.getElementById('about');
assert.strictEqual(sectionElement?.parentElement.tagName, 'MAIN');
```

Ton élément `h2` devrait avoir une balise ouvrante. 

```js
assert.match(code, /<h2\>/);
```

Ton élément `h2` devrait avoir une balise fermante. 

```js
assert.match(code, /<\/h2\>/);
```

Ton élément `h2` devrait avoir le texte `About`. Vérifie bien l'orthographe.

```js
const h2Element = document.querySelector('body h2');
assert.strictEqual(h2Element?.innerText.trim(), "About");
```

Ton élément `h2` devrait être à l'intérieur de ton élément `section`.

```js
const h2Element = document.querySelector('#about h2');
assert.strictEqual(h2Element?.parentElement.tagName, 'SECTION');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
