---
id: 669fcb5106eeb13ab7afcd09
title: Étape 1
challengeType: 0
dashedName: step-1
---

# --description--

Dans cet atelier, tu vas t'entraîner à utiliser le HTML sémantique en construisant une page de blog dédiée à M. Moustaches le chat. 

La première section que tu vas construire est l'<dfn>en-tête</dfn> de la page. 

L'élément `header` est utilisé pour représenter un contenu d'introduction comme la navigation de la page et d'autres informations introductives. 

Voici un exemple qui utilise l'élément `header` :

```html
<header>
  <h1>Main Page Title Goes Here</h1>
  <img src="example-logo.png" alt="Example logo" />
</header>
```

À l'intérieur de l'élément `body`, ajoute un élément `header`.

# --hints--

Tu devrais avoir une balise ouvrante `header`.

```js
assert.match(code, /<header>/i);
```

Tu devrais avoir une balise fermante `header`.

```js
assert.match(code, /<\/header>/i);
```

Ton élément `header` devrait être à l'intérieur de l'élément `body`.

```js
assert.exists(document.querySelector('body > header'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
  --fcc-editable-region--
    
  --fcc-editable-region--
  </body>
</html>
```
