---
id: 669fdc11c9b0263fe0814a7a
title: Étape 5
challengeType: 0
dashedName: step-5
---

# --description--

À l'intérieur de chacun des éléments `li`, tu auras besoin d'avoir un élément d'ancrage.

Pour le premier élément d'ancrage, le texte devrait être `About` et la valeur de l'attribut `href` devrait être `"#about"`. Le symbole dièse devant `about` représente un nom d'`id` qui sera ajouté plus tard dans le projet.

Pour le deuxième élément d'ancrage, le texte devrait être `Posts` et la valeur de l'attribut `href` devrait être `"#posts"`.

Pour le troisième élément d'ancrage, le texte devrait être `Contact` et la valeur de l'attribut `href` devrait être `"#contact"`.

# --hints--

Ton premier élément d'ancrage devrait avoir le texte `About` à l'intérieur de ton premier élément `li`.

```js
const anchorElement = document.querySelector("ul li:first-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "About");
```

Ton premier élément d'ancrage devrait avoir un attribut `href` mis à `"#about"`.

```js
const anchorElement = document.querySelector("ul li:first-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#about");
```

Ton deuxième élément d'ancrage devrait avoir le texte `Posts` à l'intérieur de ton deuxième élément `li`.

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
assert.strictEqual(anchorElement?.innerText.trim(), "Posts");
```

Ton deuxième élément d'ancrage devrait avoir un attribut `href` mis à `"#posts"`.

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#posts");
```

Ton troisième élément d'ancrage devrait avoir le texte `Contact` à l'intérieur de ton troisième élément `li`.

```js
const anchorElement = document.querySelector("ul li:last-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "Contact");
```

Ton troisième élément d'ancrage devrait avoir un attribut `href` mis à `"#contact"`.

```js
const anchorElement = document.querySelector("ul li:last-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#contact");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
--fcc-editable-region--
          <li></li>
          <li></li>
          <li></li>
--fcc-editable-region--
        </ul>
      </nav>
    </header>
  </body>
</html>
```
