---
id: 66a49f685961e997e337cab1
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

Le `header` sera chargé d'afficher le titre principal, l'image et la navigation de la page du blog.

À l'intérieur de l'élément `header`, ajoute un `h1` avec le texte `Welcome to Mr. Whiskers' Blog Page!`.

# --hints--

Tu devrais avoir une balise ouvrante `h1`.

```js
assert.match(code, /<h1>/i);
```

Tu devrais avoir une balise fermante `h1`.

```js
assert.match(code, /<\/h1>/i);
```

Ton élément `h1` devrait avoir le texte `Welcome to Mr. Whiskers' Blog Page!`. Vérifie bien l'orthographe.

```js
const h1Text = document.querySelector("h1")?.innerText.trim();
assert.strictEqual(h1Text, "Welcome to Mr. Whiskers' Blog Page!");
```

Ton élément `h1` devrait être imbriqué dans ton élément `header`.

```js
const h1Element = document.querySelector("h1");
assert.strictEqual(h1Element?.parentElement.tagName, 'HEADER');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
--fcc-editable-region--
      
--fcc-editable-region--
    </header>
  </body>
</html>
```
