---
id: 68e6a99a0fd7386fb66e156f
title: Étape 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Dans cet atelier, tu vas t'entraîner à styliser des cases à cocher en construisant un panneau de paramètres de jeu.

Tout le boilerplate HTML (`DOCTYPE`, `html`, `head` et `body`) a déjà été fourni.

Dans le `body`, crée un élément `div` avec un élément `h1` imbriqué à l'intérieur. Donne à l'élément `h1` le texte `Game Settings`

# --hints--

Tu devrais avoir un élément `div`.

```js
assert.lengthOf(document.querySelectorAll("div"), 1);
```

Tu devrais avoir un élément `h1`.

```js
assert.lengthOf(document.querySelectorAll("h1"), 1);
```

Ton élément `h1` devrait être à l'intérieur de ton élément `div`.

```js
assert.exists(document.querySelector("div")?.querySelector("h1"));
```

Ton élément `h1` devrait avoir le texte `Game Settings`.

```js
assert.match(document.querySelector("h1")?.textContent, /Game Settings/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
