---
id: 68e9a834dfef752dc35f6f90
title: Étape 14
challengeType: 0
dashedName: step-14
---

# --description--

Maintenant que tes cases à cocher changent de couleur quand l'utilisateur clique dessus, tu vas appliquer un peu de style pour rendre plus évident le fait que la case à cocher a été sélectionnée.

Les pseudo-classes CSS peuvent être chaînées. Tu peux ajouter une pseudo-classe ou un pseudo-élément après l'autre pour cibler les éléments qui répondent à plusieurs conditions.

Crée un nouveau sélecteur pour `input[type="checkbox"]`. Ajoute la pseudo-classe `:checked` suivie du pseudo-élément `::after`. Cela te permet d'ajouter un indicateur visuel quand la case à cocher est cochée.

Ensuite, ajoute `content` avec la valeur `"✓"`.

# --hints--

Tu devrais avoir un sélecteur de type combiné avec un sélecteur de pseudo-classe et de pseudo-élément `input[type="checkbox"]:checked::after`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after'));
```

Tu devrais ajouter une propriété `content` à ton sélecteur de type combiné avec un sélecteur de pseudo-classe et de pseudo-élément `input[type="checkbox"]:checked::after`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.content);
```

Tu devrais avoir une propriété `content` avec une valeur de `"✓"`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.content, '"✓"');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

--fcc-editable-region--

--fcc-editable-region--

```
