---
id: 68e9a6bf1a74721a4e43e217
title: Étape 10
challengeType: 0
dashedName: step-10
---

# --description--

Tu vas maintenant retirer la case à cocher par défaut appliquée par les navigateurs. Dans ton sélecteur `input[type='checkbox']`, ajoute `appearance` avec une valeur de `none`.

Définir la propriété `appearance` à `none` supprime l'apparence que le navigateur applique aux cases à cocher, ce qui te permet d'afficher le style que tu veux.

Après ça, comme la case à cocher ne sera plus visible, définis une `border` d'épaisseur `2px`, de style `solid`, et avec le code hexadécimal `#f1be32`.

# --hints--

Tu devrais ajouter une propriété `appearance` à ton sélecteur `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.appearance);
```

Tu devrais avoir une propriété `appearance` avec une valeur de `none`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.appearance, "none");
```

Tu devrais ajouter une propriété `border` à ton sélecteur `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.border);
```

Tu devrais avoir une propriété `border` avec une valeur de `2px solid #f1be32`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.border, "2px solid rgb(241, 190, 50)");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
