---
id: 68e9a7bd60c971278d2d4ffc
title: Étape 13
challengeType: 0
dashedName: step-13
---

# --description--

Dans une leçon précédente, tu as appris les pseudo-classes et pseudo-éléments en CSS. Tu vas mettre en pratique cette connaissance maintenant en créant un sélecteur de type combiné avec un sélecteur de pseudo-classe.

D'abord, commence par le format pour le sélecteur de type en le définissant comme `input[type="checkbox"]` et en ajoutant `:checked` à la fin.

Ensuite, donne-lui un `background-color` avec la valeur `#f1be32`.

Et pour finir, définis la `border-color` avec une valeur de `#e2a60d`.

# --hints--

Tu devrais avoir un sélecteur de type et de pseudo-classe combiné `input[type="checkbox"]:checked`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked'));
```

Tu devrais ajouter une propriété `background-color` à ton sélecteur combiné de type et de pseudo-classe `input[type="checkbox"]:checked`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.backgroundColor);
```

Tu devrais avoir une propriété `background-color` avec une valeur de `#f1be32`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.backgroundColor, "rgb(241, 190, 50)");
```

Tu devrais ajouter une propriété `border-color` à ton sélecteur combiné de type et de pseudo-classe `input[type="checkbox"]:checked`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.borderColor);
```

Tu devrais avoir une propriété `border-color` avec une valeur de `#e2a60d`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.borderColor, "rgb(226, 166, 13)");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

--fcc-editable-region--

--fcc-editable-region--

```
