---
id: 68e9a3194160f8f5620ddf21
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

Immédiatement après ton élément `h1`, crée quatre éléments `label`.

À l'intérieur de chacun des éléments `label`, crée un élément `input` avec l'attribut `type` défini sur `"checkbox"`.

Après chaque élément `input`, donne à tes éléments `label` les textes suivants, dans l'ordre : `Sound Effects`, `Background Music`, `Hard Mode`, et `Haptic Feedback`.

# --hints--

Tu devrais avoir quatre éléments `label`.

```js
assert.lengthOf(document.querySelectorAll("label"), 4);
```

Ton premier élément `label` devrait avoir un élément `input`.

```js
assert.exists(document.querySelectorAll("label")[0]?.querySelector("input"));
```

Ton premier élément `input` devrait avoir l'attribut `type` défini sur `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[0]?.querySelector("input")?.type, "checkbox");
```

Ton premier élément `label` devrait avoir le texte `Sound Effects`.

```js
assert.match(document.querySelectorAll("label")[0]?.textContent, /Sound Effects/i);
```

Ton texte `Sound Effects` devrait venir après l'élément `input`.

```js
assert.match(document.querySelectorAll("label")[0]?.innerHTML, /\<input[^>]*>\s*Sound Effects/i)
```

Ton deuxième élément `label` devrait avoir un élément `input`.

```js
assert.exists(document.querySelectorAll("label")[1]?.querySelector("input"));
```

Ton deuxième élément `input` devrait avoir l'attribut `type` défini sur `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[1]?.querySelector("input")?.type, "checkbox");
```

Ton deuxième élément `label` devrait avoir le texte `Background Music`.

```js
assert.match(document.querySelectorAll("label")[1]?.textContent, /Background Music/i)
```

Ton texte `Background Music` devrait venir après l'élément `input`.

```js
assert.match(document.querySelectorAll("label")[1]?.innerHTML, /\<input[^>]*>\s*Background Music/i)
```

Ton troisième élément `label` devrait avoir un élément `input`.

```js
assert.exists(document.querySelectorAll("label")[2]?.querySelector("input"));
```

Ton troisième élément `input` devrait avoir l'attribut `type` défini sur `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[2]?.querySelector("input")?.type, "checkbox");
```

Ton troisième élément `label` devrait avoir le texte `Hard Mode`.

```js
assert.match(document.querySelectorAll("label")[2]?.textContent, /Hard Mode/i)
```

Ton texte `Hard Mode` devrait venir après l'élément `input`.

```js
assert.match(document.querySelectorAll("label")[2]?.innerHTML, /\<input[^>]*>\s*Hard Mode/i)
```

Ton quatrième élément `label` devrait avoir un élément `input`.

```js
assert.exists(document.querySelectorAll("label")[3]?.querySelector("input"));
```

Ton quatrième élément `input` devrait avoir l'attribut `type` défini sur `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[3]?.querySelector("input")?.type, "checkbox");
```

Ton quatrième élément `label` devrait avoir le texte `Haptic Feedback`.

```js
assert.match(document.querySelectorAll("label")[3]?.textContent, /Haptic Feedback/i)
```

Ton texte `Haptic Feedback` devrait venir après l'élément `input`.

```js
assert.match(document.querySelectorAll("label")[3]?.innerHTML, /\<input[^>]*>\s*Haptic Feedback/i)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h1>Game Settings</h1>
--fcc-editable-region--
      
--fcc-editable-region--
    </div>
  </body>
</html>
```
