---
id: 68e9a6201e2efa142d62bfe1
title: Étape 9
challengeType: 0
dashedName: step-9
---

# --description--

Tu vas maintenant agrandir les cases à cocher pour une meilleure visibilité.

Commence par configurer un sélecteur pour `input`, en ciblant spécifiquement tes `[type="checkbox"]`.

Dans ton sélecteur, définis la `width` et la `height` à `20px`. Cela les rend plus grandes qu'avant.

Et pour rester cohérent avec ton paramétrage `cursor` qui a été défini sur les labels, ajoute `cursor` et attribue-lui la valeur `pointer`. Ainsi, quand tu survoles les cases à cocher, un pointeur s'affichera.

# --hints--

Tu devrais avoir un sélecteur dans ton CSS de `input[type="checkbox"]`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]'));
```

Tu devrais ajouter une propriété `width` à ton sélecteur `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.width);
```

Tu devrais avoir une propriété `width` avec une valeur de `20px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.width, "20px");
```

Tu devrais ajouter une propriété `height` à ton sélecteur `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.height);
```

Tu devrais avoir une propriété `height` avec une valeur de `20px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.height, "20px");
```

Tu devrais ajouter une propriété `cursor` à ton sélecteur `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.cursor);
```

Tu devrais avoir une propriété `cursor` avec une valeur de `pointer`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.cursor, "pointer");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

--fcc-editable-region--

--fcc-editable-region--
```
