---
id: 68ee86e3ee07b01626a392cc
title: Étape 6
challengeType: 0
dashedName: step-6
---

# --description--

Dans ton sélecteur `.settings-card`, définis la propriété `margin` à `auto`.

Définir la propriété `margin` à `auto` ajuste automatiquement les marges d'un élément pour répartir uniformément l'espace restant dans son conteneur, ce qui est couramment utilisé pour centrer horizontalement des éléments de type block.

Et pour finir, définis une propriété `text-align` avec la valeur `left`. Cela alignera le contenu inline, comme le texte, sur le côté gauche de son élément conteneur.

# --hints--

Tu devrais ajouter une propriété `margin` au sélecteur de classe `settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.margin);
```

Tu devrais avoir une propriété `margin` avec une valeur de `auto`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.margin, "auto");
```

Tu devrais ajouter une propriété `text-align` au sélecteur de classe `settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.textAlign);
```

Tu devrais avoir une propriété `text-align` avec une valeur de `left`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.textAlign, "left");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
