---
id: 68f15ef0446464325ce1aa45
title: Étape 8
challengeType: 0
dashedName: step-8
---

# --description--

Il te faut un peu d'espace entre les cases à cocher et les labels. Commence par utiliser le sélecteur `label` dans ton CSS et définis une propriété `display` à `block`.

Ensuite, définis une propriété `margin` à `8px auto` pour ajouter un espacement vertical entre les éléments. Cela créera un espace constant au-dessus et en dessous de chaque élément, ce qui aide à séparer le contenu et améliore la lisibilité.

Enfin, définis la propriété `cursor` à `pointer` sur les éléments `label`. Cela changera le curseur en icône de main quand un utilisateur survole un `label`, signalant que l'élément est cliquable et améliorant l'expérience utilisateur globale.

# --hints--

Tu devrais avoir un sélecteur de type dans ton CSS pour `label`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("label"));
```

Tu devrais ajouter une propriété `display` à ton sélecteur `label`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.display);
```

Tu devrais avoir une propriété `display` avec une valeur de `block`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.display, "block");
```

Tu devrais ajouter la propriété `margin` à ton sélecteur `label`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.margin);
```

Tu devrais avoir une propriété `margin` avec une valeur de `8px auto`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("label")?.margin, "8px auto");
```

Tu devrais ajouter une propriété `cursor` à ton sélecteur `label`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.cursor);
```

Tu devrais avoir une propriété `cursor` avec une valeur de `pointer`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.cursor, "pointer");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}
--fcc-editable-region--

--fcc-editable-region--
```
