---
id: 68f15cd397503429f0345262
title: Étape 15
challengeType: 0
dashedName: step-15
---

# --description--

Dans la déclaration pour le sélecteur `input[type="checkbox"]:checked::after`, définis le `display` à `block`. Définir ta propriété `display` à `block` rend l'élément de type block, ce qui signifie qu'il prend toute la largeur de son conteneur et commence sur une nouvelle ligne, te permettant de contrôler plus facilement sa largeur, sa hauteur et son espacement.

Ensuite, définis `text-align` à `center`. Cela centrera le contenu inline (comme le texte ou les éléments inline) horizontalement dans le block. Le block prendra toujours toute la largeur de son conteneur mais tout ce qui est à l'intérieur sera aligné au centre.

# --hints--

Tu devrais ajouter une propriété `display` à ton sélecteur de type combiné avec un sélecteur de pseudo-classe et de pseudo-élément `input[type="checkbox"]:checked::after`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.display);
```

Tu devrais avoir une propriété `display` avec une valeur de `block`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.display, "block");
```

Tu devrais ajouter une propriété `text-align` à ton sélecteur de type combiné avec un sélecteur de pseudo-classe et de pseudo-élément `input[type="checkbox"]:checked::after`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.textAlign);
```

Tu devrais avoir une propriété `text-align` avec une valeur de `center`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.textAlign, "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

input[type="checkbox"]:checked::after {
  content: "✓";
--fcc-editable-region--

--fcc-editable-region--
}
```
