---
id: 68f16192e8a66f3b948d3ebd
title: Étape 16
challengeType: 0
dashedName: step-16
---

# --description--

Pour finir la mise en place de l'effet visuel de ta case à cocher, définis la propriété `font-weight` avec une valeur de `bold`. Cela augmentera la visibilité de la coche.

Maintenant que c'est plus facile à voir, change la `color` à une valeur de `white`. Cela va changer la couleur de la coche dans la case à cocher quand elle est cochée.

Pour finir, ajuste la `line-height` à une valeur de `20px`.

Avec ça, tu as terminé le panneau de paramètres de jeu !

# --hints--

Tu devrais ajouter une propriété `font-weight` à ton sélecteur de type combiné avec un sélecteur de pseudo-classe et de pseudo-élément `input[type="checkbox"]:checked::after`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.fontWeight);
```

Tu devrais avoir une propriété `font-weight` avec une valeur de `bold`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.fontWeight, "bold");
```

Tu devrais ajouter une propriété `color` à ton sélecteur de type combiné avec un sélecteur de pseudo-classe et de pseudo-élément `input[type="checkbox"]:checked::after`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.color);
```

Tu devrais avoir une propriété `color` avec une valeur de `white`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.color, ["white", "#ffffff", "rgb(255,255,255)"]);
```

Tu devrais ajouter une propriété `line-height` à ton sélecteur de type combiné avec un sélecteur de pseudo-classe et de pseudo-élément `input[type="checkbox"]:checked::after`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.lineHeight);
```

Tu devrais avoir une propriété `line-height` avec une valeur de `20px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.lineHeight, "20px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

input[type="checkbox"]:checked::after {
  content: "✓";
  display: block;
  text-align: center;
--fcc-editable-region--
  
--fcc-editable-region--
}
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

input[type='checkbox']:checked:after {
  content: '✓';
  display: block;
  text-align: center;
  font-weight: bold;
  color: white;
  line-height: 20px;
}
```
