---
id: 68f1560aff1214103ab6b59b
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

Maintenant tu vas commencer à embellir la page avec un peu de style CSS. Commence par créer le sélecteur `body`.

Définis pour le `body` une propriété `height` avec une valeur de `100vh` et une propriété `background-color` avec une valeur de `#f0f0f0`.

La `height` de `100vh` fait que le `body` prend toute la hauteur de la fenêtre du navigateur, tandis que la couleur de fond gris clair fournit une base discrète pour la page.

Enfin, définis une propriété `text-align` avec la valeur `center`. Cela centrera tout le contenu inline contenu dans la page, sauf si un élément enfant la redéfinit avec son propre alignement.

# --hints--

Tu devrais utiliser le sélecteur de type `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("body"));
```

Tu devrais ajouter la propriété `height` dans le sélecteur de type `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.height);
```

Tu devrais avoir une propriété `height` avec la valeur de `100vh`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.height, "100vh");
```

Tu devrais ajouter la propriété `background-color` dans le sélecteur de type `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.backgroundColor);
```

Tu devrais avoir une propriété `background-color` avec la valeur de `#f0f0f0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("body")?.backgroundColor, "rgb(240, 240, 240)");
```

Tu devrais ajouter la propriété `text-align` dans le sélecteur de type `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.textAlign);
```

Tu devrais avoir une propriété `text-align` avec la valeur de `center`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.textAlign, "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
