---
id: 68f29dbf9125476b4808d6fb
title: Étape 5
challengeType: 0
dashedName: step-5
---

# --description--

C'est maintenant le moment de styliser le conteneur `settings-card`. Crée un sélecteur de classe pour `settings-card`.

C'est là que toute la mise en forme du conteneur va aller. Définis le `max-width` à `250px` pour fixer la taille globale de ton conteneur.

Ensuite, règle ton `padding` à `20px` pour que ton contenu ait de l'espace entre lui et la bordure du conteneur.

Après cela, crée un bord arrondi en réglant ton `border-radius` à `10px`.

Puis définis un `box-shadow` avec les valeurs `0 2px 6px rgba(0,0,0,0.2)`. Cela créera un subtil effet « surélevé » qui donnera de la profondeur au conteneur.

# --hints--

Tu devrais avoir un sélecteur de classe dans ton CSS pour `settings-card`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(".settings-card"));
```

Tu devrais ajouter la propriété `max-width` au sélecteur de classe `.settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth);
```

Tu devrais avoir une propriété `max-width` avec une valeur de `250px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth, "250px");
```

Tu devrais ajouter une propriété `padding` au sélecteur de classe `.settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding);
```

Tu devrais avoir une propriété `padding` avec une valeur de `20px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding, "20px");
```

Tu devrais ajouter une propriété `border-radius` à ton sélecteur de classe `.settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius);
```

Tu devrais avoir une propriété `border-radius` avec une valeur de `10px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius, "10px");
```

Tu devrais ajouter une propriété `box-shadow` à ton sélecteur de classe `.settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow);
```

Tu devrais avoir un `box-shadow` avec la valeur `0 2px 6px rgba(0,0,0,0.2)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow, "rgba(0, 0, 0, 0.2) 0px 2px 6px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--
```
