---
id: 6852db755c9c1274d75aae9a
title: Étape 1
challengeType: 20
dashedName: step-1
---

# --description--

Dans ce workshop, tu vas créer un extracteur de code PIN : les chiffres du code sont cachés dans chaque ligne d'un poème.

Commence par créer une fonction appelée `pin_extractor` avec un paramètre `poem`. Une fonction vide provoque une erreur, alors ajoute `pass` à l'intérieur de la fonction temporairement pour la rendre valide.

# --hints--

Tu devrais avoir une fonction nommée `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('pin_extractor')
    `)
})
```

La fonction `pin_extractor` devrait avoir un paramètre nommé `poem`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_args('poem')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
