---
id: 6852dd8486c7ee0171322950
title: Étape 2
challengeType: 20
dashedName: step-2
---

# --description--

Tu auras besoin d'une variable pour stocker le code secret pendant que tu le décodes, alors retire `pass` et crée une variable appelée `secret_code` à l'intérieur de la fonction, puis affecte-lui une chaîne vide.

# --hints--

Tu devrais avoir une variable `secret_code` à l'intérieur de la fonction `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('secret_code')
    `)
})
```

La variable `secret_code` devrait avoir pour valeur une chaîne vide.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('secret_code').is_equivalent('secret_code = ""')
    `)
})
```

Tu ne devrais plus avoir `pass` dans ton code.

```js
({
    test: () => runPython(`
    assert not _Node(_code).find_function('pin_extractor').has_pass()
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
