---
id: 68531fbb934d735415053274
title: Étape 4
challengeType: 20
dashedName: step-4
---

# --description--

Maintenant que tu as ajouté un appel de fonction, tu peux voir dans le terminal toute erreur créée à l'intérieur de la fonction, et tu peux toujours utiliser `print` pour voir la valeur des choses.

Il y a un chiffre du code PIN caché dans chaque ligne, alors à l'intérieur de la fonction, utilise la méthode `split` pour diviser la chaîne en une liste de lignes. Découpe les lignes sur le caractère de nouvelle ligne (`\n`), et affecte la liste résultante à une variable appelée `lines`.

# --hints--

Tu devrais avoir une variable nommée `lines` dans ta fonction `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('lines')
    `)
})
```

La variable `lines` devrait avoir pour valeur `poem.split('\n')`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('lines').is_equivalent('lines = poem.split("\\\\n")')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
--fcc-editable-region--

--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
```
