---
id: 68530231c05b93324e071104
title: Étape 3
challengeType: 20
dashedName: step-3
---

# --description--

Pendant que tu travailles sur la fonction, il serait utile d'avoir un appel de fonction pour le débogage, alors crée une variable `poem`, en dehors de la fonction, qui contient ce poème :

```md
Stars and the moon
shine in the sky
white and bright
until the end of the night
```

Tu peux utiliser une chaîne multi-lignes pour cela.

Ensuite, appelle la fonction avec la variable `poem` en argument.

# --hints--

Tu devrais avoir une variable globale nommée `poem`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_variable('poem')
    `)
})
```

La variable `poem` devrait avoir une valeur égale à `Stars and the moon\nshine in the sky\nwhite and bright\nuntil the end of the night`.

```js
({
    test: () => runPython(`
    pm = """Stars and the moon
    shine in the sky
    white and bright
    until the end of the night"""
    assert _Node(_code).find_variable('poem').is_equivalent(f'poem = """{pm}"""')
    `)
})
```

Tu devrais appeler la fonction `pin_extractor` avec `poem` en argument.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('pin_extractor(poem)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''

--fcc-editable-region--

--fcc-editable-region--
```
