---
id: 685578aae65e4106ad2b4300
title: Étape 7
challengeType: 20
dashedName: step-7
---

# --description--

Comme tu l'as appris dans une leçon précédente, la fonction `enumerate` te permet de suivre l'indice tout en parcourant un itérable :

```py
fruits = ['apple', 'plum', 'bananas']

for index, item in enumerate(fruits):
    print(index, item)

# 0 apple
# 1 plum
# 2 bananas
```

Le n-ième chiffre du code secret vient du n-ième mot de la n-ième ligne, donc tu dois savoir quel est le numéro de chaque ligne.

Modifie la boucle pour qu'elle itère sur `enumerate(lines)`, et ajoute une autre variable de boucle avant `line` nommée `line_index`.

Aussi, mets à jour l'appel à `print` pour qu'il soit `print(line_index, line)`.

# --hints--

La boucle devrait itérer sur `enumerate(lines)`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_iter().is_equivalent('enumerate(lines)')
    `)
})
```

La boucle devrait avoir `line_index, line` comme variable de boucle.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_vars().is_equivalent('line_index, line')
    `)
})
```

Tu devrais avoir `print(line_index, line)` à l'intérieur de la boucle.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(line_index, line)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
--fcc-editable-region--
    for line in lines:
        print(line)
--fcc-editable-region--
        words = line.split()
        print(words)

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
