---
id: 68553423d775cc60134ac8bb
title: Étape 5
challengeType: 20
dashedName: step-5
---

# --description--

Tu dois travailler sur chaque ligne indépendamment : crée une boucle sur `lines` qui utilise `line` comme variable de boucle. À l'intérieur de la boucle, affiche `line`.

# --hints--

Tu devrais avoir une boucle `for` dans ta fonction `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0]
    `)
})
```

La boucle devrait itérer sur `lines`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_iter().is_equivalent('lines')
    `)
})
```

La boucle devrait avoir `line` comme variable de boucle (`for line in lines`).

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_vars().is_equivalent('line')
    `)
})
```

Tu devrais avoir `print(line)` à l'intérieur de la boucle.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(line)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
--fcc-editable-region--

--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
