---
id: 6855918e87c33c6979d21d3f
title: Étape 14
challengeType: 20
dashedName: step-14
---

# --description--

Maintenant, la fonction extrait correctement le code du poème, mais il y a un cas limite à considérer : la ligne du poème pourrait être plus courte que prévu, et cela ferait échouer la fonction (tu peux essayer en retirant le mot `bright` du poème).

Place la ligne `secret_code += str(len(words[line_index]))` dans une instruction `if` qui vérifie qu'il y a assez de mots dans la liste `words`.

# --hints--

À l'intérieur de la boucle, tu devrais avoir une instruction `if`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0]
    `)
})
```

La condition de l'instruction if devrait être `len(words) > line_index`.


```js
({
    test: () => runPython(`
    _cond = _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]
    assert _cond.is_equivalent('len(words) > line_index') or _cond.is_equivalent('line_index < len(words)')
    `)
})
```

À l'intérieur de l'instruction `if`, tu devrais avoir `secret_code += str(len(words[line_index]))`.


```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt('secret_code += str(len(words[line_index]))')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
--fcc-editable-region--
        secret_code += str(len(words[line_index]))
--fcc-editable-region--
    return secret_code

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

print(pin_extractor(poem))

```
