---
id: 68559f3561d7df287586142e
title: Étape 16
challengeType: 20
dashedName: step-16
---

# --description--

Maintenant, l'extracteur de code PIN fonctionne, mais il pourrait être plus efficace ! Il pourrait extraire le code PIN de plusieurs poèmes en même temps !

Ajoute deux autres poèmes dans la portée globale, `poem2 = 'The grass is green\nhere and there\nhoping for rain\nbefore it turns yellow'` et `poem3 = 'There\nonce\nwas\na\ndragon'`, et pour l'instant, commente le `print` avec l'appel à `pin_extractor`.


# --hints--

Les variables `poem2` et `poem3` devraient exister et avoir la bonne valeur.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_variable('poem2')
        assert _Node(_code).has_variable('poem3')
        assert _Node(_code).find_variable('poem2').is_equivalent("poem2 = 'The grass is green\\\\nhere and there\\\\nhoping for rain\\\\nbefore it turns yellow'")
        assert _Node(_code).find_variable('poem3').is_equivalent("poem3 = 'There\\\\nonce\\\\nwas\\\\na\\\\ndragon'")
    `)
})
```

Tu devrais commenter l'appel `print(pin_extractor(poem))`.

```js
({
    test: () => runPython(`
        assert not _Node(_code).has_call('print(pin_extractor(poem))')`)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
        if len(words) > line_index:
            secret_code += str(len(words[line_index]))
        else:
            secret_code += '0'
    return secret_code

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

--fcc-editable-region--


print(pin_extractor(poem))
--fcc-editable-region--

```
