---
id: 685535bf527b2c77236b3a9a
title: Étape 6
challengeType: 20
dashedName: step-6
---

# --description--

Le n-ième chiffre du code PIN est caché sous la forme de la longueur du n-ième mot de la n-ième ligne.

Pour trouver la longueur du n-ième mot, l'étape suivante est de séparer la ligne du poème en une liste de mots : à l'intérieur de la boucle, crée une variable `words` et affecte-lui la valeur de `line` découpée en mots avec la méthode `split`.

Ensuite, ajoute un appel à `print` avec `words` comme argument.

# --hints--

Tu devrais avoir une variable `words` à l'intérieur de la boucle.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_variable('words')
    `)
})
```

La variable `words` devrait avoir pour valeur `line.split()`.

```js
({
    test: () => runPython(`
    variable = _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_variable('words')
    assert variable.is_equivalent('words = line.split()') or variable.is_equivalent('words = line.split(" ")')
    `)
})
```

Tu devrais avoir `print(words)`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(words)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line in lines:
        print(line)
--fcc-editable-region--

--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
