---
id: 6859a037b3d6721c7fbb8055
title: Étape 18
challengeType: 20
dashedName: step-18
---

# --description--

Avant la boucle, crée une nouvelle variable `secret_codes` qui a pour valeur initiale une liste vide.

Puis remplace la ligne `return secret_code` par une ligne qui ajoute `secret_code` à la liste `secret_codes`.

# --hints--

Tu devrais avoir une variable `secret_codes` à l'intérieur de la fonction `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('secret_codes')
    `)
})
```

La variable `secret_codes` devrait avoir pour valeur une liste vide.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('secret_codes').is_equivalent('secret_codes = []')
    `)
})
```

À la fin de la boucle, tu devrais avoir `secret_codes.append(secret_code)`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('secret_codes.append(secret_code)')
    `)
})
```

Tu ne devrais pas renvoyer `secret_code`.


```js
({
    test: () => runPython(`
    assert not _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_return('secret_code')
    `)
})
```


# --seed--

## --seed-contents--

```py
def pin_extractor(poems):
--fcc-editable-region--

    for poem in poems:
        secret_code = ''
        lines = poem.split('\n')
        for line_index, line in enumerate(lines):
            words = line.split()
            if len(words) > line_index:
                secret_code += str(len(words[line_index]))
            else:
                secret_code += '0'
        return secret_code
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

poem2 = 'The grass is green\nhere and there\nhoping for rain\nbefore it turns yellow'
poem3 = 'There\nonce\nwas\na\ndragon'

# print(pin_extractor(poem))

```
