---
id: 68592807861cfb658c90f70f
title: Étape 11
challengeType: 20
dashedName: step-11
---

# --description--

Maintenant, tu peux concaténer `str(len(words[line_index]))` à `secret_code`. Retire l'appel à `print` et utilise l'affectation augmentée `+=` pour la concaténation.

# --hints--

À l'intérieur de la boucle, tu devrais avoir `secret_code += str(len(words[line_index]))`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_aug_variable('secret_code').is_equivalent('secret_code += str(len(words[line_index]))')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        print(line_index, line)
        words = line.split()
--fcc-editable-region--
        print(str(len(words[line_index])))
--fcc-editable-region--


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
