---
id: 68592b42b9c7f27cbe123203
title: Étape 15
challengeType: 20
dashedName: step-15
---

# --description--

La troisième ligne du poème n'a pas de troisième mot, donc le code PIN est plus court que prévu.

Ajoute une clause `else` qui concatène un `'0'` à `secret_code` lorsqu'il n'y a pas assez de mots, afin que toutes les lignes du poème servent à trouver des chiffres.

# --hints--

Tu devrais avoir une clause `else`.


```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[1] == _Node()
    `)
})
```

La clause `else` devrait contenir `secret_code += '0'`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_if('len(words) > line_index').is_equivalent('if len(words) > line_index:\\n  secret_code += str(len(words[line_index]))\\nelse:\\n  secret_code += \\'0\\'')
    `)
})
```

`pin_extractor(poem)` devrait renvoyer `5202`.


```js
({
    test: () => runPython(`
    assert pin_extractor(poem) == '5202'
    `)
})
```

`secret_code('There\nonce\nwas\na\ndragon')` devrait renvoyer `50000`.

```js
({
    test: () => runPython(`
    assert pin_extractor('There\\nonce\\nwas\\na\\ndragon') == '50000'
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
        if len(words) > line_index:
            secret_code += str(len(words[line_index]))
--fcc-editable-region--

--fcc-editable-region--
    return secret_code

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

print(pin_extractor(poem))

```
