---
id: 6859290fc8e07e6b2a8d8523
title: Étape 12
challengeType: 20
dashedName: step-12
---

# --description--

Retire les deux appels à `print` et ajoute une instruction `return` qui renvoie `secret_code` à la dernière ligne de la fonction.

# --hints--

Tu ne devrais pas avoir d'appel à `print`.

```js
({
    test: () => runPython(`
    assert not _Node(_code).block_has_call('print')
    `)
})
```

La fonction `pin_extractor` devrait renvoyer `secret_code` à la fin.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_return('secret_code')
    `)
})
```

`pin_extractor(poem)` devrait renvoyer `5262`.


```js
({
    test: () => runPython(`
    assert pin_extractor(poem) == '5262'
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
--fcc-editable-region--
        print(line_index, line)
        words = line.split()
        print(str(len(words[line_index])))
        secret_code += str(len(words[line_index]))

--fcc-editable-region--


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
