---
id: 68592730aaf34b61ae279168
title: Étape 10
challengeType: 20
dashedName: step-10
---

# --description--

La longueur renvoyée par `len` est un entier ; tu dois la convertir en chaîne avec la fonction `str` pour pouvoir l'ajouter à la chaîne `secret_code`.

Mets à jour l'appel à `print` pour afficher la longueur sous forme de chaîne.

# --hints--

Tu devrais avoir `print(str(len(words[line_index])))` à la fin de la boucle.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(str(len(words[line_index])))')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        print(line_index, line)
        words = line.split()
--fcc-editable-region--
        print(len(words[line_index]))
--fcc-editable-region--


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
