---
id: 68e44c305f1f935d949d76bb
title: Étape 11
challengeType: 1
dashedName: step-11
---

# --description--

Comme mentionné à l'étape précédente, il y a un espace en trop après le dernier `"love"` dans la phrase `"I love love love  learning."`. Pour supprimer cet espace en trop, tu peux utiliser la méthode `trimEnd()`.

À la fin de la méthode `"love ".repeat(3)`, chaîne la méthode `trimEnd()`. Tu peux chaîner des méthodes comme ceci :

```js
.firstMethod().secondMethod()
```

Maintenant, quand tu regardes la console, tu devrais voir que l'espace en trop a été supprimé.

Et avec ce dernier changement, ton workshop est terminé !

# --hints--

Tu devrais chaîner la méthode `trimEnd()` à la méthode `"love ".repeat(3)`.

```js
assert.equal(repeatedLove, "love ".repeat(3).trimEnd());
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

--fcc-editable-region--
const repeatedLove = "love ".repeat(3)
--fcc-editable-region--
console.log(repeatedLove);

const newSentence = `I ${repeatedLove} learning.`;
console.log(newSentence);
```

# --solutions--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

const repeatedLove = "love ".repeat(3).trimEnd();
console.log(repeatedLove);

const newSentence = `I ${repeatedLove} learning.`;
console.log(newSentence);
```
