---
id: 68e4417b7317835208145323
title: Étape 6
challengeType: 1
dashedName: step-6
---

# --description--

Ensuite, crée une nouvelle variable appelée `dogsOnlySentence` et affecte-lui `exampleSentence.replace("cats", "dogs")`.

Puis sous cette variable, affiche dans la console `"Replacing all occurrences of cats with dogs:"`. Sous ce `console.log()`, ajoute un autre `console.log()` pour la variable `dogsOnlySentence`.

Tu devrais maintenant voir le texte `"I love dogs and cats are so much fun!"` affiché dans la console. Mais ce n'est pas le résultat souhaité. À l'étape suivante, tu apprendras à le corriger pour que le texte affiche `"I love dogs and dogs are so much fun!"`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Tu devrais avoir une variable `dogsOnlySentence`.

```js
assert.exists(dogsOnlySentence);
```

Ta variable `dogsOnlySentence` devrait être une chaîne.

```js
assert.isString(dogsOnlySentence);
```

Tu devrais affecter `exampleSentence.replace("cats", "dogs")` à ta variable `dogsOnlySentence`.

```js
assert.equal(dogsOnlySentence, "I love dogs and cats are so much fun!");
```

Tu devrais afficher `"Replacing all occurrences of cats with dogs:"` dans la console.

```js
assert.equal(getLogs()[6], "Replacing all occurrences of cats with dogs:");
```

Tu devrais afficher la variable `dogsOnlySentence` dans la console.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingDogsOnlySentence = codeWithoutComments.match(/console\.log\(\s*dogsOnlySentence\s*\)/g)

assert.isAtLeast(loggingDogsOnlySentence.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

--fcc-editable-region--

--fcc-editable-region--
```
