---
id: 68e43e884ef09d4d2ab05631
title: Étape 4
challengeType: 1
dashedName: step-4
---

# --description--

Il est maintenant temps de voir le résultat de l'utilisation de la méthode `replace()`.

Commence par ajouter un `console.log()` avec la valeur `"After using the replace() method:"`.

Puis sous ce `console.log()`, ajoute un autre `console.log()` avec la variable `replacedString`.

Regarde la console : tu devrais voir que la nouvelle phrase dit `"I love dogs."` au lieu de l'originale `"I love cats."`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Tu devrais afficher `"After using the replace() method:"` dans la console.

```js
assert.equal(getLogs()[2], "After using the replace() method:");
```

Tu devrais afficher la variable `replacedString` dans la console.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingReplacedStr = codeWithoutComments.match(/console\.log\(\s*replacedString\s*\)/g)

assert.isAtLeast(loggingReplacedStr.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
--fcc-editable-region--

--fcc-editable-region--
```
