---
id: 68e4465fba0bd256868b7deb
title: Étape 8
challengeType: 1
dashedName: step-8
---

# --description--

Il est maintenant temps de travailler avec une nouvelle méthode de chaîne.

Commence par créer une variable appelée `learningSentence` et affecte-lui la chaîne `"I love learning!"`.

Puis sous cette variable, affiche dans la console `"Original learning sentence:"`. Sous ce `console.log()`, ajoute un autre `console.log()` pour la variable `learningSentence`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Tu devrais avoir une variable `learningSentence`.

```js
assert.exists(learningSentence);
```

Ta variable `learningSentence` devrait être une chaîne.

```js
assert.isString(learningSentence);
```

Tu devrais affecter la chaîne `"I love learning!"` à la variable `learningSentence`.

```js
assert.equal(learningSentence, "I love learning!");
```

Tu devrais afficher `"Original learning sentence:"` dans la console.

```js
assert.equal(getLogs()[8], "Original learning sentence:");
```

Tu devrais afficher la variable `learningSentence` dans la console.

```js
assert.equal(getLogs()[9], learningSentence);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

--fcc-editable-region--

--fcc-editable-region--
```
