---
id: 68e44abea678a65b877864dc
title: Étape 10
challengeType: 1
dashedName: step-10
---

# --description--

Ce serait bien de voir ce mot répété dans une phrase. Crée une nouvelle variable appelée `newSentence` et affecte-lui le résultat de la concaténation du mot `"I "` avec la variable `repeatedLove`, suivie de la concaténation de la chaîne `" learning."`. Puis affiche la variable `newSentence` dans la console. Tu peux choisir d'utiliser soit l'opérateur `+`, soit les littéraux de modèle pour la concaténation de chaînes.

Tu devrais maintenant voir le texte `"I love love love  learning."`.

*REMARQUE* : Il y a un espace en trop après le dernier `"love"`, ce qui est normal. À la dernière étape, tu supprimeras cet espace en trop.

# --hints--

Tu devrais avoir une variable appelée `newSentence`.

```js
assert.exists(newSentence);
```

Ta variable `newSentence` devrait être une chaîne.

```js
assert.isString(newSentence);
```

Tu devrais concaténer le mot `"I "` avec la variable `repeatedLove`, suivie de la concaténation de la chaîne `" learning."`. Ce résultat devrait être affecté à la variable `newSentence`. Fais attention aux espaces.

```js
assert.equal(newSentence, `I ${repeatedLove} learning.`);
```

Tu ne devrais PAS affecter la chaîne littérale `"I love love love  learning."` à la variable `newSentence`. Relis les consignes pour la bonne approche.

```js
assert.notMatch(code, /(["'])I\s+love\s+love\s+love\s+\s+learning.(["'])/);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

const repeatedLove = "love ".repeat(3);
console.log(repeatedLove);

--fcc-editable-region--

--fcc-editable-region--
```
