---
id: 68e41084bb275539f72de9fe
title: Étape 1
challengeType: 1
dashedName: step-1
---

# --description--

Dans ce workshop, tu vas t'entraîner à utiliser les méthodes `replace()`, `replaceAll()` et `repeat()` en construisant une application de transformation de chaînes.

Pour commencer, crée une variable appelée `originalString` et affecte-lui la valeur `"I love cats."`.

# --hints--

Tu devrais avoir une variable appelée `originalString`.

```js
assert.exists(originalString);
```

Ta variable `originalString` devrait être une chaîne.

```js
assert.isString(originalString);
```

Tu devrais affecter la chaîne `"I love cats."` à la variable `originalString`.

```js
assert.equal(originalString, "I love cats.");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
