---
id: 68e43f94365a964f26c88f78
title: Étape 5
challengeType: 1
dashedName: step-5
---

# --description--

Il est maintenant temps de travailler avec une nouvelle méthode de chaîne.

Commence par créer une nouvelle variable appelée `exampleSentence` et affecte-lui la chaîne `"I love cats and cats are so much fun!"`.

Puis sous cette variable, affiche dans la console `"Original sentence:"`. Ensuite, sous ce `console.log()`, ajoute un autre `console.log()` avec la variable `exampleSentence`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Tu devrais avoir une variable `exampleSentence`.

```js
assert.exists(exampleSentence);
```

Ta variable `exampleSentence` devrait être une chaîne.

```js
assert.isString(exampleSentence);
```

Tu devrais affecter la chaîne `"I love cats and cats are so much fun!"` à la variable `exampleSentence`.

```js
assert.equal(exampleSentence, "I love cats and cats are so much fun!");
```

Tu devrais afficher `"Original sentence:"` dans la console.

```js
assert.equal(getLogs()[4], "Original sentence:");
```

Tu devrais afficher la variable `exampleSentence` dans la console.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingExSentence = codeWithoutComments.match(/console\.log\(\s*exampleSentence\s*\)/g)

assert.isAtLeast(loggingExSentence.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

--fcc-editable-region--

--fcc-editable-region--
```
