---
id: 68e42b62d7093345c6975f22
title: Étape 2
challengeType: 1
dashedName: step-2
---

# --description--

Sous ta variable `originalString`, ajoute un `console.log()` avec la chaîne `"Original string:"`.

Puis sous ce `console.log()`, ajoute un autre `console.log()` avec la variable `originalString`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Tu devrais afficher la chaîne `"Original string:"` dans la console.

```js
assert.equal(getLogs()[0], "Original string:");
```

Tu devrais afficher la variable `originalString` dans la console.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingOriginalStr = codeWithoutComments.match(/console\.log\(\s*originalString\s*\)/g)

assert.isAtLeast(loggingOriginalStr.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
--fcc-editable-region--

--fcc-editable-region--
```
