---
id: 68e44519c769bf545bb3f22a
title: Étape 7
challengeType: 1
dashedName: step-7
---

# --description--

Pour obtenir le résultat souhaité `"I love dogs and dogs are so much fun!"`, tu devras utiliser la méthode `replaceAll()` au lieu de la méthode `replace()`. Cette méthode renvoie une nouvelle chaîne pour toutes les correspondances de la sous-chaîne.

Mets à jour `exampleSentence.replace("cats", "dogs")` pour utiliser plutôt la méthode `replaceAll()`. Tu devrais maintenant voir le bon texte dans la console.

# --hints--

Tu devrais mettre à jour la ligne de code `exampleSentence.replace("cats", "dogs")` pour utiliser la méthode `replaceAll()`.

```js
assert.equal(dogsOnlySentence, "I love dogs and dogs are so much fun!");
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

--fcc-editable-region--
const dogsOnlySentence = exampleSentence.replace("cats", "dogs");
--fcc-editable-region--
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);
```
