---
id: 68e42de6396f5c49d90d1951
title: Étape 3
challengeType: 1
dashedName: step-3
---

# --description--

Dans une leçon précédente, tu as appris la méthode `replace()` qui sert à localiser une sous-chaîne dans une chaîne et à la remplacer par une autre valeur.

Voici un exemple :

```js
const text = "I love JavaScript!";
const newText = text.replace("JavaScript", "coding");
// Result: "I love coding!"
```

Souviens-toi que les chaînes sont immutables, ce qui signifie que la chaîne `text` d'origine n'est pas modifiée dans cet exemple. Une nouvelle chaîne est créée à la place.

Crée une nouvelle variable appelée `replacedString` et affecte-lui `originalString.replace("cats", "dogs")`.

# --hints--

Tu devrais avoir une variable appelée `replacedString`.

```js
assert.exists(replacedString);
```

Ta variable `replacedString` devrait être une chaîne.

```js
assert.isString(replacedString);
```

Tu devrais affecter `originalString.replace("cats", "dogs")` à la variable `replacedString`.

```js
assert.equal(replacedString, "I love dogs.");
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

--fcc-editable-region--

--fcc-editable-region--
```
