---
id: 68e448ec4ac4205968c0a805
title: Étape 9
challengeType: 1
dashedName: step-9
---

# --description--

Ce serait bien que le mot `"love"` soit répété trois fois dans la phrase `"I love learning!"`.

Comme tu t'en souviens d'une leçon précédente, tu peux répéter une chaîne un nombre précis de fois à l'aide de la méthode `repeat()`. Voici un exemple :

```js
const word = "Hello!";
const repeatedWord = word.repeat(3);

console.log(repeatedWord);  // "Hello!Hello!Hello!"
```

Comme les chaînes sont immutables, cette méthode ne modifie pas la chaîne d'origine. Elle renvoie une nouvelle chaîne avec le contenu répété.

Crée une variable appelée `repeatedLove` et affecte-lui `"love ".repeat(3)`. Puis affiche la variable `repeatedLove` dans la console pour voir le résultat.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Tu devrais avoir une variable appelée `repeatedLove`.

```js
assert.exists(repeatedLove);
```

Ta variable `repeatedLove` devrait être une chaîne.

```js
assert.isString(repeatedLove);
```

Tu devrais affecter `"love ".repeat(3)` à la variable `repeatedLove`.

```js
assert.equal(repeatedLove, "love ".repeat(3));
```

Tu devrais afficher la variable `repeatedLove` dans la console.

```js
assert.equal(getLogs()[10], repeatedLove);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

--fcc-editable-region--

--fcc-editable-region--
```
