---
id: 66fe4f33a2cc9b33f4d5cd9b
title: Construire un générateur de permutations
challengeType: 26
dashedName: build-a-permutation-generator
---

# --description--

Dans ce lab, tu vas construire un générateur de permutations qui prendra une chaîne et renverra toutes les permutations possibles des caractères de la chaîne. Par exemple, les permutations possibles de la chaîne `cat` sont `cat`, `cta`, `act`, `atc`, `tac` et `tca`.

La méthode récursive pour créer les permutations d'une chaîne consiste à stocker la partie initiale fixe de la chaîne (préfixe), et à créer les permutations du reste.

Par exemple, considérons le mot `machine`. Le premier tour de création des permutations se ferait en fixant le `m` comme préfixe de la chaîne, puis en créant les permutations du reste de la chaîne, `achine`.

Pour le reste de la chaîne, les permutations continuent de la même façon. Une lettre est ajoutée au préfixe, par exemple le `c`, donc le préfixe devient `mc`. Ensuite, chacune des permutations de `ahine` est concaténée au préfixe.

Cela continue jusqu'à ce que le préfixe contienne toutes les lettres, et que le reste de la chaîne soit vide : cela signifie qu'une permutation a été créée.

**Objectif :** Remplir les user stories ci-dessous et faire passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu dois créer une fonction nommée `permuteString`.
2. La fonction `permuteString` doit prendre un paramètre, une chaîne, puis deux paramètres avec une valeur par défaut : une valeur de préfixe et un tableau vide pour stocker et renvoyer les résultats. La valeur de préfixe doit être une chaîne utilisée pour accumuler les caractères afin de former une permutation. La fonction sera appelée avec un seul argument, comme `permuteString("cat")`.
3. À l'intérieur de la fonction, tu dois vérifier si la longueur de la chaîne passée est `0`. Si c'est le cas, pousse le préfixe courant dans les résultats et renvoie les résultats.
4. Itère sur chaque caractère de la chaîne d'entrée et, pour chaque itération, retire le caractère courant de la chaîne et appelle la fonction `permuteString` de façon récursive avec des arguments mis à jour pour construire les permutations restantes.
5. Tu dois renvoyer le tableau final des résultats.
6. Tu dois t'assurer que les permutations sont uniques en supprimant les doublons.

# --hints--

Tu dois avoir une fonction `permuteString`.

```js
assert.isFunction(permuteString);
```

Tu dois utiliser la récursion dans ta fonction `permuteString`.

```js
assert.match(permuteString.toString(), /permuteString\s*\(/);
```

`permuteString("far")` devrait renvoyer `[ "far", "fra", "afr", "arf", "rfa", "raf" ]`.

```js
assert.sameMembers(permuteString("far"), [ "far", "fra", "afr", "arf", "rfa", "raf" ]);
```

`permuteString("fcc")` devrait renvoyer `[ "fcc", "cfc", "ccf" ]`.

```js
assert.sameMembers(permuteString("fcc"), [ "fcc", "cfc", "ccf" ]);
```

`permuteString("p")` devrait renvoyer `[ "p" ]`.

```js
assert.deepStrictEqual(permuteString("p"), ["p"]);
```

`permuteString("")` devrait renvoyer `[""]`.

```js
let emptyArr= permuteString("");

// 1 because the empty string is being pushed and is a permutation of itself
assert.lengthOf(emptyArr, 1);

```

`permuteString("walk")` devrait renvoyer `["walk", "wakl", "wlak", "wlka", "wkla", "wkal", "awlk", "awkl", "alwk", "alkw", "aklw", "akwl", "lawk", "lakw", "lwak", "lwka", "lkaw", "lkwa", "kawl", "kalw", "kwal", "kwla", "klaw", "klwa"]`.
`.

```js
assert.sameMembers(permuteString("walk"), ["walk", "wakl", "wlak", "wlka", "wkla", "wkal", "awlk", "awkl", "alwk", "alkw", "aklw", "akwl", "lawk", "lakw", "lwak", "lwka", "lkaw", "lkwa", "kawl", "kalw", "kwal", "kwla", "klaw", "klwa"]);
```

`permuteString` devrait renvoyer les résultats corrects.

```js
// tests to prevent hard coding
assert.sameMembers(permuteString("road"), [ "road",
  "roda",
  "raod",
  "rado",
  "rdoa",
  "rdao",
  "orad",
  "orda",
  "oard",
  "oadr",
  "odra",
  "odar",
  "arod",
  "ardo",
  "aord",
  "aodr",
  "adro",
  "ador",
  "droa",
  "drao",
  "dora",
  "doar",
  "daro",
  "daor" ]);

assert.sameMembers(permuteString("fog"), [ "fog",
  "fgo",
  "ofg",
  "ogf",
  "gfo",
  "gof" ]);

assert.sameMembers(permuteString("bird"), [ "bird",
  "bidr",
  "brid",
  "brdi",
  "bdir",
  "bdri",
  "ibrd",
  "ibdr",
  "irbd",
  "irdb",
  "idbr",
  "idrb",
  "rbid",
  "rbdi",
  "ribd",
  "ridb",
  "rdbi",
  "rdib",
  "dbir",
  "dbri",
  "dibr",
  "dirb",
  "drbi",
  "drib" ]);

```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const permuteString = (str, prefix = "", result = []) => {
  if (str.length === 0) {
    result.push(prefix);
    return result;
  }

  // Sort the string initially to group duplicates
  const sortedStr = str.split("").sort().join("");

  for (let i = 0; i < sortedStr.length; i++) {
    if (i > 0 && sortedStr[i] === sortedStr[i - 1]) {
      // Skip duplicate characters
      continue;
    }

    const rem = sortedStr.slice(0, i) + sortedStr.slice(i + 1);
    permuteString(rem, prefix + sortedStr[i], result);
  }

  return result;
};
```
