---
id: 68e97fe79367ad7b5dd6c9cd
title: Étape 4
challengeType: 0
dashedName: step-4
---

# --description--

Dans cette étape, ajoute un élément `h1` avec le texte `XYZ Bookstore`.

# --hints--

Tu dois avoir un élément `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Le texte de ton élément `h1` doit être `XYZ Bookstore`. Vérifie bien l'orthographe et les majuscules.

```js
assert.equal(document.querySelector('h1')?.innerText.trim(), 'XYZ Bookstore');
```

Ton élément `h1` doit être dans l'élément `body`.

```js
assert.match(code, /<body>[\s\S]*<h1>[\s\S]*<\/h1>[\s\S]*<\/body>/i);
```
# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
