---
id: 68ebdbacdd3fa474132cc975
title: Étape 5
challengeType: 0
dashedName: step-5
---

# --description--

Sous l'élément `h1`, ajoute un élément `p` avec ce texte : `Browse our collection of amazing books!`.

# --hints--

Tu dois avoir un élément `p`.

```js
assert.exists(document.querySelector('p'));
```

Le texte de ton élément `p` doit être : `Browse our collection of amazing books!`.

```js
assert.equal(document.querySelector('p')?.innerText.trim(), 'Browse our collection of amazing books!');
```

Ton élément `p` doit se trouver sous ton élément `h1`.

```js
assert.exists(document.querySelector('h1 + p'));
```
# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
