---
id: 68eca5c0065c82d256d29ca3
title: Étape 17
challengeType: 0
dashedName: step-17
---

# --description--

Dans la deuxième carte, ajoute un élément `button` avec l'attribut `class` défini sur `btn` et le texte `Buy Now`.

Les deux éléments `button` partagent maintenant la même `class`, ce qui signifie qu'ils peuvent être stylisés ensemble de manière cohérente.

# --hints--

Tu dois avoir un élément `button` imbriqué dans le deuxième élément qui a une classe `card`.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('button'));
```

L'élément `button` de ta deuxième carte doit avoir un attribut `class`.

```js
const cards = document.querySelectorAll('.card');
assert.isTrue(cards[1]?.querySelector('button')?.hasAttribute('class'));
```

L'élément `button` de ta deuxième carte doit avoir une `class` avec la valeur `btn`.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('button')?.className, 'btn');
```

Le texte de l'élément `button` de ta deuxième carte doit être `Buy Now`.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('button')?.innerText.trim(), 'Buy Now');
```
# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
