---
id: 68ec99e478211dc578699944
title: Étape 11
challengeType: 0
dashedName: step-11
---

# --description--

Sous l'élément `h2` dans le premier élément ayant une classe `card`, ajoute un élément `p` avec le texte suivant :

```md
This is an epic story of Sally and her dog Rex as they navigate through other worlds.
```

# --hints--

Tu dois avoir un élément `p` imbriqué dans l'élément ayant une classe `card`.

```js
assert.exists(document.querySelector('.card p'));
```

Le texte de ton élément `p` doit être `This is an epic story of Sally and her dog Rex as they navigate through other worlds.`

```js
assert.equal(document.querySelector('.card p')?.innerText.trim(), "This is an epic story of Sally and her dog Rex as they navigate through other worlds.");
```

Ton élément `p` doit se trouver sous ton élément `h2`.

```js
assert.exists(document.querySelector('.card h2 + p'));
```
# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
