---
id: 68eca2a795b333ca5fee30a8
title: Étape 13
challengeType: 0
dashedName: step-13
---

# --description--

Maintenant, crée une deuxième carte de livre. Ajoute un autre élément `div` avec l'attribut `class` défini sur `card`. Remarque que tu peux réutiliser le même nom de classe pour plusieurs éléments afin d'appliquer un style cohérent.

# --hints--

Tu dois avoir deux éléments `div` avec un attribut `class` ayant la valeur `card`.

```js
assert.lengthOf(document.querySelectorAll('.card'), 2);
```

Le deuxième élément avec une classe `card` doit être dans l'élément ayant une classe `card-container`.

```js
assert.lengthOf(document.querySelectorAll('.card-container .card'), 2);
```

Les deux éléments avec une classe `card` doivent être frères.

```js
assert.exists(document.querySelector('.card + .card'));
```
# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
  --fcc-editable-region--
    
  --fcc-editable-region--
  </div>
</body>

</html>
```
