---
id: 68eca6364c5616d393389a30
title: Étape 18
challengeType: 0
dashedName: step-18
---

# --description--

Souviens-toi, un élément HTML ressemble à ceci :

```html
<element attribute="value">
    inner text
</element>
```

Sous l'élément avec la classe `card-container`, ajoute un nouvel élément `p` avec ce texte :

```md
Review your selections and continue to checkout.
```

Sous l'élément `p`, crée un élément `div` avec l'attribut `class` défini sur `btn-container`. Ce conteneur regroupera tes éléments de bouton de navigation.

# --hints--

Le nouvel élément `p` ajouté doit se trouver sous l'élément avec une classe `card-container`.

```js
assert.exists(document.querySelector('.card-container + p'));
```

Le texte de ton nouvel élément `p` doit être `Review your selections and continue to checkout.`

```js
const cardContainer = document.querySelector('.card-container');
assert.equal(cardContainer?.nextElementSibling?.innerText.trim(), 'Review your selections and continue to checkout.');
```

Tu dois avoir un élément `div` avec une `class` de `btn-container`.

```js
assert.exists(document.querySelector('div.btn-container'));
```

Ton élément `div` ayant une `class` de `btn-container` doit venir après le nouvel élément `p` ajouté.

```js
assert.exists(document.querySelector('p + .btn-container'));
```
# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
      <button class="btn">Buy Now</button>
    </div>
  </div>
  
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
