---
id: 68ec6cd5b7e8f5a8f7319e32
title: Étape 6
challengeType: 0
dashedName: step-6
---

# --description--

L'élément `div` sert de conteneur pour regrouper d'autres éléments HTML. Tu utiliseras surtout l'élément `div` quand tu veux regrouper des éléments HTML qui partageront un même ensemble de styles CSS.

Sous l'élément `p`, ajoute un élément `div`. Ce `div` sera un conteneur pour tes cartes de livres.

**Remarque :** cet atelier n'applique pas de CSS. Les classes et les éléments regroupés sont utiles pour le style CSS, mais dans cet atelier ils servent seulement à structurer et regrouper le contenu. Tu apprendras comment fonctionne le style dans un module suivant.

# --hints--

Tu dois avoir un élément `div`.

```js
assert.exists(document.querySelector('div'));
```

Ton élément `div` doit se trouver sous ton élément `p`.

```js
assert.exists(document.querySelector('p + div'));
```
# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
