---
id: 68ec6d9a315221aa31e54816
title: Étape 7
challengeType: 0
dashedName: step-7
---

# --description--

L'attribut `class` sert à identifier un ou plusieurs éléments pour le style. Contrairement à l'attribut `id`, les noms de classe n'ont pas besoin d'être uniques : plusieurs éléments peuvent partager la même classe.

Voici un exemple :

```html
<p class="example">example paragraph</p>
```

Ajoute un attribut `class` à ton élément `div` et définis sa valeur sur `card-container`.

# --hints--

Ton élément `div` doit avoir un attribut `class`.

```js
assert.isTrue(document.querySelector('div')?.hasAttribute('class'));
```

Ton élément `div` doit avoir un attribut `class` avec la valeur `card-container`.

```js
assert.equal(document.querySelector('div')?.className, 'card-container');
```
# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
--fcc-editable-region--
  <div>
--fcc-editable-region--
    
  </div>
</body>

</html>
```
