---
id: 68eca4c4b0f952cf09fabe09
title: Étape 15
challengeType: 0
dashedName: step-15
---

# --description--

Dans le deuxième élément ayant une classe `card`, ajoute un élément `h2` avec le texte `Dave's Cooking Adventure`.

# --hints--

Tu dois avoir un élément `h2` imbriqué dans le deuxième élément qui a une classe `card`.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('h2'));
```

L'élément `h2` dans le deuxième élément ayant une classe `card` doit avoir le texte `Dave's Cooking Adventure`.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('h2')?.innerText.trim(), "Dave's Cooking Adventure");
```
# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
