---
id: 68eca6e5a6759ed4ea0034dc
title: Étape 19
challengeType: 0
dashedName: step-19
---

# --description--

Dans l'élément avec une classe `btn-container`, ajoute deux éléments `button` :

Premier bouton :

- Id : `view-cart-btn`
- Class : `btn`
- Texte : `View Cart`

Deuxième bouton :

- Id : `checkout-btn`
- Class : `btn`
- Texte : `Checkout`


Bravo ! Tu as réussi à construire la structure d'une page de librairie avec des divs, des classes et des ids pour organiser ton contenu.

# --hints--

Tu dois avoir deux éléments `button` imbriqués dans l'élément `.btn-container`.

```js
assert.lengthOf(document.querySelectorAll('.btn-container button'), 2);
```

Ton premier bouton dans l'élément avec la classe `btn-container` doit avoir un attribut `id` avec la valeur `view-cart-btn`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.id, 'view-cart-btn');
```

Ton premier bouton dans l'élément ayant une classe `btn-container` doit avoir un attribut `class` avec la valeur `btn`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.className, 'btn');
```

Le texte de ton premier bouton doit être `View Cart`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.innerText.trim(), 'View Cart');
```

Ton deuxième bouton dans l'élément avec une classe `btn-container` doit avoir un attribut `id` avec la valeur `checkout-btn`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.id, 'checkout-btn');
```

Ton deuxième bouton dans l'élément avec `class="btn-container"` doit avoir un attribut `class` avec la valeur `btn`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.className, 'btn');
```

Le texte de ton deuxième bouton doit être `Checkout`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.innerText.trim(), 'Checkout');
```
# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
      <button class="btn">Buy Now</button>
    </div>
  </div>
  
  <p>Review your selections and continue to checkout.</p>
  
  <div class="btn-container">
  --fcc-editable-region--
    
  --fcc-editable-region--
  </div>
</body>

</html>
```

# --solutions--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>

  <div class="card-container">
     <div class="card" id="sally-adventure-book">
     <h2>Sally's SciFi Adventure</h2>
     <p>
          This is an epic story of Sally and her dog Rex as they navigate
          through other worlds.
     </p>
     <button class="btn">Buy Now</button>
     </div>

     <div class="card" id="dave-cooking-book">
     <h2>Dave's Cooking Adventure</h2>
     <p>
          This is the story of Dave as he learns to cook everything from
          pancakes to pasta, one recipe at a time.
     </p>
     <button class="btn">Buy Now</button>
     </div>
  </div>

  <p>Review your selections and continue to checkout.</p>

  <div class="btn-container">
     <button id="view-cart-btn" class="btn">View Cart</button>
     <button id="checkout-btn" class="btn">Checkout</button>
  </div>
</body>

</html>
```
