---
id: 68ec9332a9b5b2b32487bd00
title: Étape 9
challengeType: 0
dashedName: step-9
---

# --description--

L'attribut `id` ajoute un identifiant unique à un élément HTML. Chaque `id` doit être unique dans une page et ne doit être utilisé qu'une seule fois.

Les valeurs de `id` ne peuvent pas contenir d'espaces et doivent seulement contenir des lettres, des chiffres, des underscores et des tirets.

Voici un exemple :

```html
<p id="para">example paragraph</p>
```

Ajoute un attribut `id` à ton élément qui a une classe `card` et définis sa valeur sur `sally-adventure-book`.

# --hints--

Ton élément avec une classe `card` doit avoir un attribut `id`.

```js
assert.isTrue(document.querySelector('.card')?.hasAttribute('id'));
```

Ton élément ayant une classe `card` doit avoir un `id` avec la valeur `sally-adventure-book`.

```js
assert.equal(document.querySelector('.card')?.id, 'sally-adventure-book');
```
# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
  --fcc-editable-region--
    <div class="card">
  --fcc-editable-region--
      
    </div>
  </div>
</body>

</html>
```
