---
id: 68eca11e99e3c5c894ca9d69
title: Étape 12
challengeType: 0
dashedName: step-12
---

# --description--

L'élément `button` sert à créer des boutons cliquables sur une page web. Les boutons sont des éléments interactifs sur lesquels les utilisateurs peuvent cliquer pour effectuer des actions.

Ajoute un élément `button` dans l'élément qui a une `class` de `card`, donne au bouton un attribut `class` défini sur `btn`, et le texte `Buy Now`.

# --hints--

Tu dois avoir un élément `button` imbriqué dans l'élément qui a une classe `card`.

```js
assert.exists(document.querySelector('.card button'));
```

Ton élément `button` doit avoir un attribut `class`.

```js
assert.isTrue(document.querySelector('.card button')?.hasAttribute('class'));
```

Ton élément `button` doit avoir une `class` avec la valeur `btn`.

```js
assert.equal(document.querySelector('.card button')?.className, 'btn');
```

Le texte de ton élément `button` doit être `Buy Now`.

```js
assert.equal(document.querySelector('.card button')?.innerText.trim(), 'Buy Now');
```
# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
