---
id: 69134d14150d51f8066a3246
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

Maintenant, améliore la structure de ton document HTML pour que ta page soit encodée correctement.

Dans l'élément `head`, ajoute l'élément `<meta charset="UTF-8">`.

Enfin, ajoute un élément `body` sous la section `head`. C'est là que tout le contenu visible de ta page ira.

# --hints--

Ton code doit inclure un élément `<meta charset="UTF-8">`.

```js
assert.exists(document.querySelector('head > meta[charset="utf-8"]'));
```

Ton code doit inclure des balises ouvrante et fermante `<body>`.

```js
assert.match(code,/<body>[\s\S]*<\/body>/i);
```

Ton élément `body` doit se trouver sous l'élément `head`.

```js
assert.match(code, /<\/head>[\s\S]*<body>[\s\S]*<\/body>[\s\S]*<\/html>/i);
```
# --seed--
## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
--fcc-editable-region--
  
  <title>XYZ Bookstore Page</title>
</head>

--fcc-editable-region--
</html>
```
