---
id: 69134b6cb43aa5f254950a10
title: Étape 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Dans cet atelier, tu vas créer une page de librairie avec des cartes qui affichent des informations sur différents livres. Tu vas t'entraîner à organiser du contenu avec des éléments `div`, des classes et des ID.

Commence ta page de librairie en créant le code de base HTML.

Ajoute la déclaration `<!DOCTYPE html>`, puis les éléments `html` et `head`.

Ajoute un attribut `lang` à l'élément `html` et définis-le sur `"en"`.

# --hints--

Ton code doit inclure une déclaration `<!DOCTYPE html>`.

```js
assert.match(code, /<!DOCTYPE html>/i);
```

Ton code doit inclure des balises ouvrante et fermante `<html>`.

```js
assert.match(code,/<html\b[^>]*>[\s\S]*<\/html>/i);
```

Ton élément `html` doit avoir un attribut `lang` défini sur `"en"`.

```js
assert.match(code, /<html[^>]*\blang\s*=\s*["']en["'][^>]*>/i);
```

Ton code doit inclure des balises ouvrante et fermante `<head>`.

```js
assert.match(code,/<head>[\s\S]*<\/head>/i);
```

Ton élément `head` doit être imbriqué dans l'élément `html`.

```js
assert.match(code, /<\/html>\s*$/);
```
# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
