---
id: 691f85085787a8c4ca589fbb
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

Ajoute l'élément `title` dans l'élément `head`.

Définis le titre de la page sur `XYZ Bookstore Page`.

# --hints--

Ton code doit inclure des balises ouvrante et fermante `<title>`.

```js
assert.match(code,/<title>[\s\S]*<\/title>/i);
```

Ton titre doit être `XYZ Bookstore Page`.

```js
assert.equal(document.querySelector('title')?.innerText.trim(),'XYZ Bookstore Page');
```

Ton élément `title` doit être imbriqué dans l'élément `head`.

```js
assert.match(code, /<head>[\s\S]*<title>[\s\S]*<\/title>[\s\S]*<\/head>/i);
```
# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
--fcc-editable-region--
  
--fcc-editable-region-- 
</head>

</html>
```
