---
id: 67c562286b29447da020d407
title: Construire un générateur de mot de passe à usage unique
challengeType: 25
dashedName: build-a-one-time-password-generator
demoType: onClick
---

# --description--

Dans ce labo, tu vas générer un OTP (One-Time Password) à 6 chiffres et l'afficher à l'utilisateur. L'OTP expirera après 5 secondes, et un nouvel OTP sera généré quand l'utilisateur cliquera sur le bouton `"Generate New OTP"`.

**Objectif :** Remplis les user stories ci-dessous et fais passer tous les tests pour terminer le labo.

**User Stories :**

1. Tu dois utiliser le hook `useEffect` pour gérer le compte à rebours.
2. Ton composant `OTPGenerator` doit renvoyer un élément `div` avec le nom de classe `container`.
3. Le `div` ayant la classe `container` doit inclure les éléments suivants :

    - Un élément `h1` avec l'ID `otp-title` et le texte `"OTP Generator"`.
    - Un élément `h2` avec l'ID `otp-display` qui affiche soit le message `"Click 'Generate OTP' to get a code"`, soit l'OTP généré s'il en existe un.
    - Un élément `p` avec l'ID `otp-timer` et un attribut `aria-live` défini à une valeur valide qui :
        - Commence vide.
        - Affiche `"Expires in: X seconds"` après le clic sur le bouton, où `X` représente le temps restant avant l'expiration de l'OTP.
        - Affiche le message `"OTP expired. Click the button to generate a new OTP."` une fois que le compte à rebours atteint `0`.
    - Un élément `button` avec l'ID `generate-otp-button` libellé `"Generate OTP"`. Au clic, il doit générer un nouvel OTP et démarrer un compte à rebours de 5 secondes.
    - Le bouton `"Generate OTP"` doit être désactivé pendant que le compte à rebours est actif.

4. Tu dois t'assurer que le compte à rebours s'arrête automatiquement une fois qu'il atteint `0` seconde pour éviter des mises à jour inutiles.
5. L'OTP généré doit faire 6 chiffres.

# --before-all--

```js
let clock = __FakeTimers.install();
```

# --after-all--

```js
clock.uninstall();
```

# --hints--

Tu ne dois pas retirer l'assignation existante `const { useState, useEffect, useRef } = React;` du code.

```js
assert.match(code, /{\s*(useState\s*,\s*useEffect\s*,\s*useRef|useState\s*,\s*useRef\s*,\s*useEffect|useEffect\s*,\s*useState\s*,\s*useRef|useEffect\s*,\s*useRef\s*,\s*useState|useRef\s*,\s*useState\s*,\s*useEffect|useRef\s*,\s*useEffect\s*,\s*useState|useState|useEffect|useRef)\s*}\s*=\s*React;/);
```

Tu dois afficher un `div` avec la classe `.container`.

```js
const containerDiv = document.querySelector('.container');
assert.exists(containerDiv);
```

Le `.container` doit contenir un élément `h1` avec l'ID `otp-title` et le texte `"OTP Generator"`.

```js
const h1 = document.querySelector(".container h1#otp-title");
assert.exists(h1);
assert.equal(h1.textContent, "OTP Generator");
```

Le `.container` doit contenir un élément `h2` avec l'ID `otp-display` qui affiche l'OTP une fois généré.

```js
const h2 = document.querySelector(".container h2#otp-display");
assert.exists(h2);
```

Au départ, le `h2` à l'intérieur de `.container` doit afficher le message `"Click 'Generate OTP' to get a code"`.

```js
assert.strictEqual(document.querySelector('.container h2#otp-display').textContent.trim(), "Click 'Generate OTP' to get a code");
```

L'élément `.container` doit contenir un élément `p` avec l'ID `otp-timer`.

```js
const container = document.querySelector('.container');
assert.exists(container, "The container element should exist");
const pElements = container.querySelectorAll('p');
assert.strictEqual(pElements.length, 1);
assert.strictEqual(pElements[0].id, "otp-timer", "The p element should have the id otp-timer");
```

L'élément `p` doit avoir un attribut `aria-live` défini à `assertive` ou `polite`.

```js
const pElement = document.querySelector('.container p');
assert.exists(pElement?.ariaLive);
assert.oneOf(pElement?.ariaLive, ["assertive", "polite"]);
```

Au départ, l'élément `p` doit être vide.

```js
const pElement = document.querySelector('.container p');
assert.strictEqual(pElement?.innerText, "", "The `p` element should be empty");
```

Le `.container` doit contenir un élément `button` avec l'ID `generate-otp-button` et le texte `"Generate OTP"`.

```js
const button = document.querySelector('.container button#generate-otp-button').textContent;
assert.exists(button);
assert.strictEqual(button, 'Generate OTP');
```

Quand on clique sur le bouton, un nouvel OTP doit s'afficher dans l'élément `h2` avec l'id `otp-display`.

```js
try {
// allow the page to render
clock.tick(1)
const button = document.querySelector(".container button#generate-otp-button");
const h2 = document.querySelector(".container h2#otp-display");

button.click();

clock.tick(300)

assert.match(h2.textContent?.trim(), /[0-9]/, "OTP should be a number");
} finally {
// we need to wait for the countdown to finish for the next test to work otherwise the button will be disabled
await clock.runAllAsync();
}
```

L'OTP généré doit faire 6 chiffres.

```js
  try {
const button = document.querySelector(".container button#generate-otp-button");
button.click();
clock.tick(1);
const h2 = document.querySelector(".container h2#otp-display");

assert.match(h2.textContent.trim(), /^[0-9]{6}$/, "OTP should be a 6-digit number");

 } finally {
// we need to wait for the countdown to finish for the next test to work otherwise the button will be disabled
await clock.runAllAsync();
}
```


Quand on clique sur le bouton, l'élément `p` avec l'id `otp-timer` doit afficher un compte à rebours de 5 secondes.

```js
  try {
// allow the page to render
clock.tick(1)
const button = document.querySelector(".container button#generate-otp-button");
const h2 = document.querySelector(".container h2#otp-display");

button.click();

clock.tick(300)
const p = document.querySelector(".container p#otp-timer");
assert.match(p.textContent, /Expires in: 5 seconds/, "Countdown should start at 5 seconds");

  } finally {
// we need to wait for the countdown to finish for the next test to work otherwise the button will be disabled
await clock.runAllAsync();
}
```

Le message dans l'élément `p` avec l'id `otp-timer` doit se mettre à jour chaque seconde pour montrer le temps restant.

```js
const button = document.querySelector(".container button#generate-otp-button");
button.click();
clock.tick(1)
let p = document.querySelector(".container p#otp-timer");
assert.match(p.textContent, /Expires in: 5 seconds/);
await clock.tickAsync(1000)
p = document.querySelector(".container p#otp-timer");
assert.match(p.textContent, /Expires in: 4 seconds/, "Countdown should update to 4 seconds after 1 second");
```

Le bouton `"Generate OTP"` doit être désactivé pendant que le compte à rebours tourne.

```js
try {
  const button = document.querySelector(".container button#generate-otp-button");
  button.click();
  clock.tick(1);

  assert.isTrue(button.disabled);
} finally {
  await clock.runAllAsync();
}
```

Le bouton `"Generate OTP"` doit être réactivé une fois que le compte à rebours atteint 0.

```js

const button = document.querySelector(".container button#generate-otp-button");
button.click();
await clock.tickAsync(6000);

assert.isFalse(button.disabled);
```

Quand le compte à rebours atteint `0`, tu dois afficher le message `OTP expired. Click the button to generate a new OTP.`.

```js
  const button = document.querySelector(".container button#generate-otp-button");
  button.click();

  await clock.tickAsync(6000);

  const p = document.querySelector(".container p#otp-timer");
  assert.equal(p.textContent.trim(), "OTP expired. Click the button to generate a new OTP.");
```

Tu dois exporter le composant `OTPGenerator`.

```js
assert.isFunction(window.index.OTPGenerator);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>OTP Generator</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { OTPGenerator } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<OTPGenerator />);
    </script>
</body>

</html>
```

```css

```

```jsx
const { useState, useEffect, useRef } = React;

export const OTPGenerator = () => {};
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>OTP Generator</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { OTPGenerator } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<OTPGenerator />);
    </script>
</body>

</html>
```

```css
body {
    font-family: 'Castoro Titling', cursive;
}

.container {
    text-align: center;
    margin-top: 50px;
}

h1 {
    font-size: 30px;
    color: #333;
}

h2 {
    font-size: 25px;
    color: #555;
}

p {
    font-size: 16px;
    color: #888;
}

button {
    padding: 10px 20px;
    margin-top: 20px;
    font-size: 16px;
    border: none;
    background-color: #383f59;
    color: white;
    border-radius: 5px;
}

button:disabled {
  background-color:#57585c;
  cursor: not-allowed;
}

button:hover {
    background-color: #444;
}
```

```jsx
const { useState, useEffect, useRef } = React;

export const OTPGenerator = () => {
  const intervalId = useRef(null);
  const [otp, setOtp] = useState(null);
  const [secondsLeft, setSecondsLeft] = useState(5);
  const [isCounting, setIsCounting] = useState(false);
  const [hasGeneratedOtp, setHasGeneratedOtp] = useState(false);

  const generateOTP = () => Math.floor(100000 + Math.random() * 900000);

  const handleGenerateOtp = () => {
    setOtp(generateOTP());
    setSecondsLeft(5);
    setIsCounting(true);
    setHasGeneratedOtp(true);
    intervalId.current = setInterval(() => {
      setSecondsLeft(prevSeconds => prevSeconds - 1);
    }, 1000);
  };

  useEffect(() => {
    if (secondsLeft === 0) {
      setIsCounting(false);
      clearInterval(intervalId.current);
    }
  }, [secondsLeft]);

  useEffect(
    () => () => {
      clearInterval(intervalId.current);
    },
    []
  );

  return (
    <div className='container'>
      <h1 id='otp-title'>OTP Generator</h1>
      <h2 id='otp-display'>
        {otp ? otp : "Click 'Generate OTP' to get a code"}
      </h2>
      <p id='otp-timer' aria-live='assertive'>
        {isCounting
          ? `Expires in: ${secondsLeft} seconds`
          : hasGeneratedOtp &&
            "OTP expired. Click the button to generate a new OTP."}
      </p>
      <button
        id='generate-otp-button'
        onClick={handleGenerateOtp}
        disabled={isCounting}
      >
        Generate OTP
      </button>
    </div>
  );
};
```
