---
id: 673b3d6b7ef7318eef926d5a
title: Créer un mood board
challengeType: 25
dashedName: build-a-mood-board
demoType: onClick
---

# --description--

Un mood board est un collage d'images et de texte qui véhicule une idée générale, un objectif ou un sentiment sur un sujet particulier.

Dans ce lab, tu vas créer un mood board avec des composants React réutilisables. Le CSS a déjà été fourni pour toi.

**Objectif :** Remplis les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu dois créer et exporter un composant `MoodBoardItem` qui accepte trois props : `color`, `image` et `description`.
2. Ton composant `MoodBoardItem` doit renvoyer un `div` avec la classe `mood-board-item` comme élément de plus haut niveau.
3. Tu dois définir la couleur de fond de l'élément `.mood-board-item` à la valeur de la prop `color` avec des styles en ligne.
4. Tu dois afficher un élément `img`, avec une classe `mood-board-image` et son attribut `src` défini à la valeur de la prop `image`, dans l'élément `.mood-board-item`.
5. Tu dois afficher un élément `h3`, avec une classe `mood-board-text` et pour texte la valeur de la prop `description`, dans l'élément `.mood-board-item`.
6. Tu dois créer et exporter un `MoodBoard`.
7. Ton composant `MoodBoard` doit renvoyer un `div` comme élément de plus haut niveau.
8. Ton composant `MoodBoard` doit afficher un élément `h1` avec une classe `mood-board-heading` et le texte `Destination Mood Board`.
9. Ton composant `MoodBoard` doit afficher un `div` avec une classe `mood-board`.
10. Ton composant `MoodBoard` doit afficher au moins trois composants `MoodBoardItem` dans l'élément `.mood-board`, chacun devant passer les props `color`, `image` et `description` avec des valeurs valides.

Tu peux utiliser les images suivantes dans ton Mood Board si tu le souhaites :

- `https://cdn.freecodecamp.org/curriculum/labs/pathway.jpg`
- `https://cdn.freecodecamp.org/curriculum/labs/shore.jpg`
- `https://cdn.freecodecamp.org/curriculum/labs/grass.jpg`
- `https://cdn.freecodecamp.org/curriculum/labs/ship.jpg`
- `https://cdn.freecodecamp.org/curriculum/labs/santorini.jpg`
- `https://cdn.freecodecamp.org/curriculum/labs/pigeon.jpg`

# --hints--

Tu dois exporter un composant `MoodBoardItem`.

```js
assert.isFunction(window.index.MoodBoardItem, 1);
```

Ton composant `MoodBoardItem` doit renvoyer un `div` avec une classe `mood-board-item` comme élément de plus haut niveau.

```js
const item = document.getElementsByClassName('mood-board-item')[0];
assert.exists(item);
assert.equal(item.tagName, 'DIV');
```

La couleur de fond de l'élément `.mood-board-item` doit être définie à la valeur de la prop `color` avec des styles en ligne.

```js
  const container = await __helpers.prepTestComponent(window.index.MoodBoardItem, { color: "red" });
  const moodBoardItem = container.querySelector(".mood-board-item");
  assert.equal(moodBoardItem.style.backgroundColor, "red");
```

Ton composant `MoodBoardItem` doit afficher un élément `img` avec une classe `mood-board-image` et son `src` défini à la valeur de la prop `image`.

```js
  const container = await __helpers.prepTestComponent(window.index.MoodBoardItem, { image: "https://cdn.freecodecamp.org/curriculum/labs/pathway.jpg" });
  const img = container.querySelector(".mood-board-image");
  assert.exists(img);
  assert.equal(img.tagName, "IMG");
  assert.equal(img.src, "https://cdn.freecodecamp.org/curriculum/labs/pathway.jpg");
```

Ton composant `MoodBoardItem` doit afficher un élément `h3` avec une classe `mood-board-text` et son texte défini à la valeur de la prop `description`.

```js
  const container = await __helpers.prepTestComponent(window.index.MoodBoardItem, { description: "Caribbean" });
  const text = container.querySelector(".mood-board-text");
  assert.exists(text);
  assert.equal(text.tagName, "H3");
  assert.equal(text.textContent, "Caribbean");
```

Tu dois exporter un composant `MoodBoard`.

```js
assert.isFunction(window.index.MoodBoard, 1);
```

Ton composant `MoodBoard` doit renvoyer un `div` comme élément de plus haut niveau.

```js
const root = document.getElementById('root');
assert.equal(root.children[0].tagName, 'DIV');
```

Ton composant `MoodBoard` doit afficher un élément `h1` avec une classe `mood-board-heading` et le texte `Destination Mood Board`.

```js
const heading = document.getElementsByClassName('mood-board-heading')[0];
assert.exists(heading);
assert.equal(heading.tagName, 'H1');
assert.equal(heading.textContent, 'Destination Mood Board');
```

Ton composant `MoodBoard` doit afficher au moins trois composants `MoodBoardItem`, chacun devant passer les props `color`, `image` et `description` avec des valeurs valides.

```js
const moodboard = document.getElementsByClassName('mood-board')[0];
const items = moodboard.querySelectorAll(".mood-board-item");
assert.isAtLeast(items.length, 3);

items.forEach((item) => {
  const color = item.style.backgroundColor;
  const img = item.querySelector("img");
  const text = item.querySelector("h3");

  assert.isAtLeast(color.length, 1);
  assert.isAtLeast(img.src.length, 1);
  assert.isAtLeast(text.textContent.length, 1);
});
```

Ton composant `MoodBoard` doit être affiché dans l'élément `#root` de la page.

```js
const root = document.getElementById('root');
const moodBoard = root.getElementsByClassName('mood-board')[0];
assert.exists(moodBoard);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Mood Board</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { MoodBoard } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <MoodBoard />
      );
    </script>
  </body>
</html>
```

```css
body {
  background-color: #ffffcc;
}

.mood-board-heading {
  text-align: center;
  font-size: 2.5em;
  color: #333;
  margin-top: 20px;
}

.mood-board {
  display: grid;
  grid-template-columns: repeat(3, 1fr);
  gap: 20px;
  padding: 20px;
  max-width: 900px;
  margin: 0 auto;
}

.mood-board-item {
  border-radius: 10px;
  padding: 10px;
  display: flex;
  flex-direction: column;
  align-items: center;
  justify-content: center;
  color: #fff;
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.2);
  text-align: center;
  height: 250px;
}

.mood-board-image {
  border-radius: 5px;
  width: 180px;
  height: 150px;
  object-fit: cover;
}

.mood-board-text {
  margin-top: 20px;
  font-size: 1.2em;
}
```

```jsx

```

# --solutions--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Mood Board</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { MoodBoard } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <MoodBoard />
      );
    </script>
  </body>
</html>
```

```css
body {
  background-color: #ffffcc;
}

.mood-board-heading {
  text-align: center;
  font-size: 2.5em;
  color: #333;
  margin-top: 20px;
}

.mood-board {
  display: grid;
  grid-template-columns: repeat(3, 1fr);
  gap: 20px;
  padding: 20px;
  max-width: 900px;
  margin: 0 auto;
}

.mood-board-item {
  border-radius: 10px;
  padding: 10px;
  display: flex;
  flex-direction: column;
  align-items: center;
  justify-content: center;
  color: #fff;
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.2);
  text-align: center;
  height: 250px;
}

.mood-board-image {
  border-radius: 5px;
  width: 180px;
  height: 150px;
  object-fit: cover;
}

.mood-board-text {
  margin-top: 20px;
  font-size: 1.2em;
}
```

```jsx
export function MoodBoardItem(props) {
  return (
    <div className='mood-board-item' style={{ backgroundColor: props.color }}>
      <img src={props.image} alt='Mood' className='mood-board-image' />
      <h3 className='mood-board-text'>{props.description}</h3>
    </div>
  );
}

export function MoodBoard() {
  return (
    <div>
      <h1 className='mood-board-heading'>Destination Mood Board</h1>
      <div className='mood-board'>
        <MoodBoardItem
          color='#2da64f'
          image='https://cdn.freecodecamp.org/curriculum/labs/pathway.jpg'
          description='Caribbean'
        />
        <MoodBoardItem
          color='#8e44ad'
          image='https://cdn.freecodecamp.org/curriculum/labs/shore.jpg'
          description='Gawadar Beach'
        />
        <MoodBoardItem
          color='#3498db'
          image='https://cdn.freecodecamp.org/curriculum/labs/grass.jpg'
          description='Cape Town'
        />
        <MoodBoardItem
          color='#bf3d7e'
          image='https://cdn.freecodecamp.org/curriculum/labs/ship.jpg'
          description='Suez Canal'
        />
        <MoodBoardItem
          color='#e74c3c'
          image='https://cdn.freecodecamp.org/curriculum/labs/santorini.jpg'
          description='Santorini'
        />
        <MoodBoardItem
          color='#95a5a6'
          image='https://cdn.freecodecamp.org/curriculum/labs/pigeon.jpg'
          description='Istanbul'
        />
      </div>
    </div>
  );
}
```
