---
id: 66ba762af611169359d9d369
title: Construire un article de journal
challengeType: 25
dashedName: build-a-newspaper-article
demoType: onClick
---

# --description--

**Objectif :** Réponds aux scénarios utilisateur ci-dessous et fais passer tous les tests pour terminer le lab.

**Scénarios utilisateur :**

1. Tu dois régler la `font-size` racine de ton document HTML sur `24px`.
1. Tu dois avoir un élément avec la classe `newspaper` qui contient tous tes autres éléments.
1. Ton élément `.newspaper` doit avoir une `font-size` de `16px` et une police `Open Sans` avec une police de secours `sans-serif`.
1. Dans ton élément `.newspaper`, tu dois avoir au moins sept autres éléments : un pour le nom du journal avec la classe `name`, un pour la date de l'article avec la classe `date`, un pour le titre avec la classe `headline`, un pour le sous-titre avec la classe `sub-headline`, un pour l'auteur avec la classe `author`, et deux paragraphes ayant chacun la classe `text`. Tous ces éléments doivent être remplis avec les informations de ton article.
1. Ton élément `.name` doit avoir une `font-size` deux fois plus grande que la `font-size` de l'élément racine et doit utiliser la police `Times New Roman`, avec une police de secours `serif`.
1. Tes éléments `.name` et `.author` doivent utiliser CSS pour mettre tous leurs caractères en majuscules.
1. Ton élément `.headline` doit avoir une `font-size` deux fois plus grande que la `font-size` de son élément parent et doit être en gras.
1. Ton élément `.sub-headline` doit avoir un `font-weight` de `100`, une `font-size` égale à `1.5` fois la `font-size` de son élément parent et doit être en italique.
1. Ton élément `.author` doit utiliser CSS pour le mettre en gras.
1. Tes éléments `.text` doivent avoir un `text-indent` de `20px`.
1. Tes éléments `.text` doivent avoir une `line-height` deux fois plus grande que la `font-size` de leur élément parent.
1. La première lettre de tes éléments `.text` doit être en gras et deux fois plus grande que la `font-size` de leur élément parent. Utilise le sélecteur `::first-letter` pour cela.

# --hints--

Tu dois avoir un sélecteur `html` qui règle sa `font-size` sur `24px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.fontSize, '24px');
```

Tu dois avoir un élément avec la classe `newspaper`.

```js
assert.exists(document.querySelector('.newspaper'));
```

Tu dois avoir un sélecteur `.newspaper` qui règle la `font-size` de ses éléments sur `16px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.newspaper')?.fontSize, '16px');
```

Tu dois avoir un sélecteur `.newspaper` qui règle la `font-family` de ses éléments sur `'Open Sans', sans-serif`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.newspaper')?.fontFamily, ['"Open Sans", sans-serif', 'Open Sans, sans-serif']);
```

Tu dois avoir un élément avec la classe `name` dans ton élément `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .name'));
```

Ton élément `.name` doit contenir le nom de ton journal.

```js
assert.isAtLeast(document.querySelector('.newspaper > .name')?.innerText.length, 1);
```

Tu dois avoir un élément avec la classe `date` dans ton élément `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .date'));
```

Ton élément `.date` doit contenir la date.

```js
assert.isAtLeast(document.querySelector('.newspaper > .date')?.innerText.length, 1);
```

Tu dois avoir un élément avec la classe `headline` dans ton élément `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .headline'));
```

Ton élément `.headline` doit contenir le titre de ton article.

```js
assert.isAtLeast(document.querySelector('.newspaper > .headline')?.innerText.length, 1);
```

Tu dois avoir un élément avec la classe `sub-headline` dans ton élément `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .sub-headline'));
```

Ton élément `.sub-headline` doit contenir le sous-titre de ton article.

```js
assert.isAtLeast(document.querySelector('.newspaper > .sub-headline')?.innerText.length, 1);
```

Tu dois avoir un élément avec la classe `author` dans ton élément `.newspaper`.

```js
assert.isAtLeast(document.querySelectorAll('.newspaper > .author').length, 1);
```

Ton élément `.author` doit contenir l'auteur de ton article.

```js
assert.isAtLeast(document.querySelector('.newspaper > .author')?.innerText.length, 1);
```

Tu dois avoir au moins deux éléments paragraphe, chacun avec la classe `text`, dans ton élément `.newspaper`.

```js
assert.isAtLeast(document.querySelectorAll('.newspaper > .text').length, 2);
```

Tes éléments `.text` doivent contenir le texte de ton article.

```js
const textEls = document.querySelectorAll('.newspaper > .text');
assert.isAtLeast(textEls.length, 1);
textEls.forEach(el => assert.isAtLeast(el.innerText.length, 1))
```

Tu dois avoir un sélecteur `.name` qui règle la `font-size` de ses éléments sur `2rem`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.name')?.fontSize, '2rem');
```

Tu dois avoir un sélecteur `.name` qui règle la `font-family` de ses éléments sur `'Times New Roman', serif`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.name')?.fontFamily, ['"Times New Roman", serif', 'Times New Roman, serif']);
```

L'élément `.name` doit avoir la propriété `text-transform` réglée sur `uppercase`.

```js
const condition1 = new __helpers.CSSHelp(document).getStyle('.name')?.textTransform === 'uppercase';
const condition2 = new __helpers.CSSHelp(document).getStyle('.author, .name')?.textTransform === 'uppercase';
const condition3 = new __helpers.CSSHelp(document).getStyle('.name, .author')?.textTransform === 'uppercase';
assert.isTrue(condition1 || condition2 || condition3);
```

L'élément `.author` doit avoir la propriété `text-transform` réglée sur `uppercase`.

```js
const condition1 = new __helpers.CSSHelp(document).getStyle('.author')?.textTransform === 'uppercase';
const condition2 = new __helpers.CSSHelp(document).getStyle('.author, .name')?.textTransform === 'uppercase';
const condition3 = new __helpers.CSSHelp(document).getStyle('.name, .author')?.textTransform === 'uppercase';
assert.isTrue(condition1 || condition2 || condition3);
```

Tu dois avoir un sélecteur `.headline` qui règle la `font-size` de ses éléments sur `2em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.headline')?.fontSize, '2em');
```

Tu dois avoir un sélecteur `.headline` qui règle le `font-weight` de ses éléments sur `bold`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.headline')?.fontWeight, ['bold', '700']);
```

Tu dois avoir un sélecteur `.sub-headline` qui règle le `font-weight` de ses éléments sur `100`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontWeight, '100');
```

Tu dois avoir un sélecteur `.sub-headline` qui règle la `font-size` de ses éléments sur `1.5em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontSize, '1.5em');
```

Tu dois avoir un sélecteur `.sub-headline` qui règle le `font-style` de ses éléments sur `italic`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontStyle, 'italic');
```

Tu dois avoir un sélecteur `.author` qui règle le `font-weight` de ses éléments sur `bold`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.author')?.fontWeight, ['bold', '700']);
```

Tu dois avoir un sélecteur `.text` qui règle le `text-indent` de ses éléments sur `20px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text')?.textIndent, '20px');
```

Tu dois avoir un sélecteur `.text` qui règle la `line-height` de ses éléments sur `2em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text')?.lineHeight, '2em');
```

Tu dois avoir un sélecteur `.text::first-letter` qui règle le `font-weight` de ses éléments sur `bold`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.text::first-letter')?.fontWeight, ['bold', '700']);
```

Tu dois avoir un sélecteur `.text::first-letter` qui règle la `font-size` de ses éléments sur `2em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text::first-letter')?.fontSize, '2em');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Newspaper Article</title>
  </head>
  <body>

  </body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Daily Chuckles</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main class="newspaper">
    <h1 class="name">Daily Chuckles</h1>
    <p class="date">July 5, 2024</p>
    <h2 class="headline">Breaking: Grandma Edna Saves Earth</h2>
    <h3 class="sub-headline">Alien Invasion Foiled by Tech-Savvy Grandma's Wi-Fi Password</h3>
    <p class="author">By Jane Doe</p>
    <p class="text">In an extraordinary twist of fate, Grandma Edna found herself at the forefront of a potential crisis
      when her clever Wi-Fi security measures unexpectedly thwarted an alien invasion. As alien spacecraft descended
      upon the town, panic spread until Grandma Edna calmly intervened, resetting her router with a cryptic passphrase
      known only to her.</p>
    <p class="text">The aliens, encountering unforeseen technological barriers, struggled to breach Grandma Edna's
      fortified network. Frustrated and bewildered, they eventually retreated to reassess their invasion strategy,
      leaving residents both relieved and amazed at Grandma Edna's resourcefulness.</p>
  </main>
</body>

</html>
```

```css
html {
  font-size: 24px;
  background-color: MintCream;
  padding-top: 100px;
}

.newspaper {
  font-family: 'Open Sans', sans-serif;
  max-width: 800px;
  margin: 20px auto;
  padding: 20px;
  background-color: white;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
  font-size: 16px;
}

.name {
  font-family: 'Times New Roman', serif;
  font-size: 2rem;
  text-transform: uppercase;
  text-align: center;
  margin-bottom: 10px;
}

.date {
  font-size: 1em;
  text-align: center;
  margin-bottom: 20px;
}

.headline {
  font-size: 2em;
  font-weight: bold;
  line-height: 1.2;
  margin-bottom: 10px;
}

.sub-headline {
  font-size: 1.5em;
  font-style: italic;
  font-weight: 100;
  line-height: 1.4;
  margin-bottom: 15px;
}

.author {
  font-size: 1em;
  text-transform: uppercase;
  font-weight: bold;
  margin-bottom: 20px;
}

.text {
  text-indent: 20px;
  font-size: 1em;
  line-height: 2em;
  margin-bottom: 15px;
}

.text::first-letter {
  font-size: 2em;
  font-weight: bold;
}
```
