---
id: 67cd525b0e575b1e5abf4c03
title: Étape 9
challengeType: 0
dashedName: step-9
---

# --description--

Comme mentionné plus tôt, la fonction set est responsable de la mise à jour du state. Une approche courante quand on bascule entre des valeurs d'état booléennes est de mettre à jour le state de cette façon :

```js
const [isUserLoggedIn, setIsUserLoggedIn] = useState(false);

// updating state
setIsUserLoggedIn(!isUserLoggedIn);
```

L'exemple suivant utilise la fonction set appelée `setIsUserLoggedIn` pour mettre à jour le statut de connexion de l'utilisateur. Il est courant d'utiliser l'opérateur logique NOT (`!`) pour basculer entre `true` et `false`. Comme la valeur initiale est `false`, appliquer l'opérateur `!` la changera en `true`.

À l'intérieur de la fonction `handleToggleVisibility`, utilise `setIsVisible` pour mettre à jour l'état de visibilité avec l'opérateur logique NOT (`!`).

# --hints--

Tu devrais utiliser `setIsVisible` pour mettre à jour la visibilité du texte. Reporte-toi à l'exemple si tu as besoin d'aide.

```js
assert.match(code, /setIsVisible\s*\(\s*!isVisible\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
<head>
  <meta charset="UTF-8" />
  <title>Toggle Visibility</title>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
  <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    src="index.jsx"
  ></script>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <div id="root"></div>
  <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    data-presets="react"
    data-type="module"
  >
    import { ToggleApp } from "./index.jsx";
    ReactDOM.createRoot(document.getElementById("root")).render(<ToggleApp />);
  </script>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f4f4f4;
}

#toggle-container {
  text-align: center;
  background: white;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

#toggle-button {
  padding: 10px 20px;
  border: none;
  background: #007BFF;
  color: white;
  border-radius: 5px;
  cursor: pointer;
  transition: background 0.3s ease;
}

#toggle-button:hover {
  background: #0056b3;
}

#message {
  margin-top: 20px;
  font-size: 1.125rem;
  color: #333;
}
```

```jsx
const { useState } = React;

export const ToggleApp = () => {
  const [isVisible, setIsVisible] = useState(false);

  --fcc-editable-region--
  const handleToggleVisibility = () => {

  }
  --fcc-editable-region--

  return (
    <div id="toggle-container">
      <button id="toggle-button">Message</button>
      {isVisible && <p id="message">I love freeCodeCamp!</p>}
    </div>
  );
};
```
