---
id: 67cd4cf8b5fe391b324ad11f
title: Étape 5
challengeType: 0
dashedName: step-5
---

# --description--

Maintenant que tu as créé la variable d'état `isVisible`, il est temps de commencer à l'utiliser.

Dans les leçons précédentes, tu as appris le rendu conditionnel en ligne qui te permet d'afficher un contenu différent selon une certaine condition. C'est un motif courant d'utiliser l'opérateur logique AND (`&&`) pour rendre conditionnellement un texte comme ceci :

```jsx
function Notification({ message }) {
  return (
    <div>
      {message && <p>{message}</p>}
    </div>
  );
}
```

Dans cet exemple, l'élément paragraphe ne s'affichera que si le message est truthy. Souviens-toi que « truthy » désigne toute valeur qui s'évalue à `true` dans un contexte booléen.

Pour cette étape, tu devras rendre conditionnellement l'élément paragraphe. Quand ce sera correctement fait, tu devrais voir que le paragraphe n'est plus visible sur la page.

# --hints--

Tu devrais rendre conditionnellement le paragraphe pour qu'il ne s'affiche que quand `isVisible` est `true`. Reporte-toi à l'exemple si tu as besoin d'aide.

```js
  {
    React.useState = (arg) => ([false, () => {}]);
    const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;
    const exports = {};
    const _a = eval(script);
    const div = await __helpers.prepTestComponent(exports.ToggleApp);
    const pEl = div.querySelector("p#message");
    assert.notExists(pEl);
  }
  {
    React.useState = (arg) => ([true, () => {}]);
    const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;
    const exports = {};
    const _a = eval(script);
    const div = await __helpers.prepTestComponent(exports.ToggleApp);
    const pEl = div.querySelector("p#message");
    assert.exists(pEl);
  }
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
<head>
  <meta charset="UTF-8" />
  <title>Toggle Visibility</title>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
  <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    src="index.jsx"
  ></script>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <div id="root"></div>
  <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    data-presets="react"
    data-type="module"
  >
    import { ToggleApp } from "./index.jsx";
    ReactDOM.createRoot(document.getElementById("root")).render(<ToggleApp />);
  </script>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f4f4f4;
}

#toggle-container {
  text-align: center;
  background: white;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

#toggle-button {
  padding: 10px 20px;
  border: none;
  background: #007BFF;
  color: white;
  border-radius: 5px;
  cursor: pointer;
  transition: background 0.3s ease;
}

#toggle-button:hover {
  background: #0056b3;
}

#message {
  margin-top: 20px;
  font-size: 1.125rem;
  color: #333;
}
```

```jsx
const { useState } = React;

export const ToggleApp = () => {
  const [isVisible, setIsVisible] = useState(false);

  return (
    <div id="toggle-container">
      <button id="toggle-button">Message</button>
      --fcc-editable-region--
      <p id="message">I love freeCodeCamp!</p>
      --fcc-editable-region--
    </div>
  );
};
```
