---
id: 68599633407374c0d574587b
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

L'étape suivante est d'ajouter les boutons qui serviront à basculer entre les différentes planètes et leurs faits.

À l'intérieur de l'élément `tablist`, ajoute trois éléments `button` chacun avec un attribut `role` défini à `"tab"`. Le premier élément `button` devrait avoir le texte `🌍 Earth`. Le deuxième `button` devrait avoir le texte `🪐 Saturn`. Le troisième élément `button` devrait avoir le texte `🔴 Mars`.

# --hints--

Tu devrais avoir trois éléments `button` à l'intérieur de l'élément `tablist`.

```js
const btns = document.querySelectorAll("div[role='tablist'] button");
assert.lengthOf(btns, 3);
```

Chacun de tes éléments `button` devrait avoir un attribut `role` défini à `"tab"`.

```js
const btns = document.querySelectorAll("div[role='tablist'] button[role='tab']");
assert.lengthOf(btns, 3);
```

Ton premier élément `button` devrait avoir le texte `🌍 Earth`.

```js
const firstBtn = document.querySelector("[role='tab']");
assert.equal(firstBtn.textContent, "🌍 Earth");
```

Ton deuxième élément `button` devrait avoir le texte `🪐 Saturn`.

```js
const secondBtn = document.querySelectorAll("[role='tab']")[1];
assert.equal(secondBtn.textContent, "🪐 Saturn")
```

Ton troisième élément `button` devrait avoir le texte `🔴 Mars`.

```js
const thirdBtn = document.querySelectorAll("[role='tab']")[2]
assert.equal(thirdBtn.textContent, "🔴 Mars");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./style.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        --fcc-editable-region--

        --fcc-editable-region--
      </div>
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
