---
id: 68599af28ead38c22590a89d
title: Étape 4
challengeType: 0
dashedName: step-4
---

# --description--

Il faut un moyen de relier les boutons au contenu qu'ils représentent. On le fait avec l'attribut `aria-controls` sur chaque bouton. Plus tard, tu ajouteras le contenu et les attributs `id` correspondants sur les éléments de contenu.

Commence par ajouter l'attribut `aria-controls` à chacun des éléments `button`. La valeur de l'attribut `aria-controls` du premier bouton devrait être `"panel-earth"`. L'attribut `aria-controls` du deuxième bouton devrait être `"panel-saturn"`. L'attribut `aria-controls` du troisième bouton devrait être `"panel-mars"`.

# --hints--

Chacun de tes éléments `button` devrait avoir un attribut `aria-controls`.

```js
const btns = document.querySelectorAll("div[role='tablist'] button[aria-controls]");
assert.lengthOf(btns, 3);
```

Le premier élément `button` devrait avoir un attribut `aria-controls` défini à `"panel-earth"`.

```js
const firstBtn = document.querySelectorAll("[role='tab']")[0];
assert.equal(firstBtn.getAttribute("aria-controls"), "panel-earth");
```

Le deuxième élément `button` devrait avoir un attribut `aria-controls` défini à `"panel-saturn"`.

```js
const secondBtn = document.querySelectorAll("[role='tab']")[1];
assert.equal(secondBtn.getAttribute("aria-controls"), "panel-saturn");
```

Le troisième élément `button` devrait avoir un attribut `aria-controls` défini à `"panel-mars"`.

```js
const thirdBtn = document.querySelectorAll("[role='tab']")[2];
assert.equal(thirdBtn.getAttribute("aria-controls"), "panel-mars");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        --fcc-editable-region--
        <button role="tab">🌍 Earth</button>
        <button role="tab">🪐 Saturn</button>
        <button role="tab">🔴 Mars</button>
        --fcc-editable-region--
      </div>
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
