---
id: 685976a92ffab8b2dc744b72
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

À l'intérieur de ton élément `.tabs`, ajoute un élément `div` avec un attribut `role` défini à `"tablist"` et un attribut `aria-labelledby` défini à `"tabs-title"`.

L'attribut `aria-labelledby` est nécessaire ici pour associer la tablist à un libellé visible, afin que les technologies d'assistance comme les lecteurs d'écran annoncent le rôle de la tablist aux utilisateurs.

# --hints--

Tu devrais avoir un élément `div` à l'intérieur de ton élément `.tabs`.

```js
assert.exists(document.querySelector(".tabs div"));
```

Ton élément `div` devrait avoir un attribut `role` défini à `"tablist"`.

```js
assert.equal(document.querySelector(".tabs div").getAttribute("role"), "tablist");
```

Ton élément `div` devrait avoir un attribut `aria-labelledby` défini à `"tabs-title"`.

```js
assert.equal(document.querySelector(".tabs div").getAttribute("aria-labelledby"), "tabs-title");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>

      --fcc-editable-region--

      --fcc-editable-region--
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
