---
id: 685db0bd1280154ae85bcf9f
title: Étape 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Dans ce workshop, tu vas construire une tablist accessible sur les planètes avec HTML, CSS et JavaScript. Le code de base HTML et le CSS t'ont été fournis.

Dans ces premières étapes, tu vas compléter le HTML restant de la tablist.

Commence par créer un élément `div` avec une classe définie à `"tabs"`. À l'intérieur de ce `div`, ajoute un élément `h2` avec un attribut `id` défini à `"tabs-title"` et le contenu textuel `Planets`.

# --hints--

Tu devrais avoir un élément `div`.

```js
assert.exists(document.querySelector("div"));
```

Ton élément `div` devrait avoir une classe définie à `"tabs"`.

```js
assert.exists(document.querySelector("div.tabs"));
```

Ton élément `div` devrait contenir un élément `h2`.

```js
assert.exists(document.querySelector(".tabs h2"));
```

Ton élément `h2` devrait avoir un attribut `id` défini à `"tabs-title"`.

```js
assert.equal(document.querySelector(".tabs h2").getAttribute("id"), "tabs-title");
```

Ton élément `h2` devrait avoir le contenu textuel `Planets`.

```js
assert.equal(document.querySelector(".tabs h2").textContent, "Planets");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    --fcc-editable-region--

    --fcc-editable-region--

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
