---
id: 68e45a2b4ab9b2f48b1cc6df
title: Concevoir une page de mise en page de formules tarifaires
challengeType: 25
dashedName: design-a-pricing-plans-layout-page
demoType: onClick
---

# --description--

**Objectif :** Remplis les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User Stories**

1. Ta page devrait avoir un élément `h1` avec le texte `Pricing Plans`.
2. Ta page devrait avoir un élément `div` avec la classe `pricing-container` sous l'élément `h1`.
   - Le sélecteur `.pricing-container` devrait avoir une propriété `display` avec la valeur `flex` et une propriété `flex-wrap` avec la valeur `wrap`.
3. Dans l'élément `.pricing-container`, tu devrais avoir trois éléments `div` avec la classe `pricing-card` pour représenter les formules tarifaires.
   - L'un des éléments `.pricing-card` devrait avoir la classe `basic-plan` en plus de la classe `pricing-card`.
   - L'un des éléments `.pricing-card` devrait avoir la classe `pro-plan` en plus de la classe `pricing-card`.
   - L'un des éléments `.pricing-card` devrait avoir la classe `premium-plan` en plus de la classe `pricing-card`.
4. Ton élément `.basic-plan` devrait avoir un élément `h2` avec le texte `Basic`.
5. Ton élément `.basic-plan` devrait avoir un élément `p` avec le texte `$9/month`.
6. Ton élément `.pro-plan` devrait avoir un élément `h2` avec le texte `Pro`.
7. Ton élément `.pro-plan` devrait avoir un élément `p` avec le texte `$19/month`.
8. Ton élément `.premium-plan` devrait avoir un élément `h2` avec le texte `Premium`.
9. Ton élément `.premium-plan` devrait avoir un élément `p` avec le texte `$29/month`.
10. Chacun de tes éléments `.pricing-card` devrait :
    - Utiliser Flexbox pour empiler son contenu dans une colonne et répartir l'espace entre les enfants avec `space-between`.
    - Définir la propriété `flex` à `0 0 200px` pour lui donner une largeur cohérente et l'empêcher de grandir ou de rétrécir dans la mise en page.
    - Définir la propriété `border` à `2px solid black` pour voir comment les différentes cartes occupent l'espace.
11. L'élément `.basic-plan` devrait apparaître en premier dans la mise en page. Tu devrais utiliser la propriété `order` pour cela.
12. L'élément `.pro-plan` devrait apparaître en deuxième dans la mise en page. Tu devrais utiliser la propriété `order` et définir sa propriété `flex-grow` à `2` pour qu'il occupe plus d'espace que les autres formules.
13. L'élément `.premium-plan` devrait venir en dernier dans la mise en page. Tu devrais utiliser la propriété `order` pour cela.

**Note :** Assure-toi de lier ta feuille de style dans ton HTML et d'appliquer ton CSS.

# --hints--

Tu devrais avoir un élément `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Ta page devrait avoir un élément `h1` avec le texte `Pricing Plans`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent.trim(), 'Pricing Plans');
```

Tu devrais avoir un élément `div` sous l'élément `h1`.

```js
assert.exists(document.querySelector('h1 + div'));
```

Ton élément `div` devrait avoir la classe `pricing-container`.

```js
assert.exists(document.querySelector('h1 + div.pricing-container'));
```

Tu devrais avoir un sélecteur `.pricing-container`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.pricing-container'));
```

Ton sélecteur `pricing-container` devrait avoir la propriété `display` avec une valeur de `flex`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-container')?.display, 'flex');
```

Ton sélecteur `.pricing-container` devrait avoir la propriété `flex-wrap` avec une valeur de `wrap`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-container')?.flexWrap, 'wrap');
```

Tu devrais avoir trois éléments `div` avec la classe `pricing-card` dans l'élément `.pricing-container`.

```js
const divEls = document.querySelectorAll('.pricing-container .pricing-card');
assert.lengthOf(divEls, 3);
```

L'un de tes éléments `.pricing-card` devrait avoir la classe `basic-plan`.

```js
assert.exists(document.querySelector('div.pricing-card.basic-plan'));
```

L'un de tes éléments `.pricing-card` devrait avoir la classe `premium-plan`.

```js
assert.exists(document.querySelector('div.pricing-card.premium-plan'));
```

L'un de tes éléments `.pricing-card` devrait avoir la classe `pro-plan`.

```js
assert.exists(document.querySelector('div.pricing-card.pro-plan'));
```

Tu devrais avoir un sélecteur `.pricing-card`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.pricing-card'));
```

Ton sélecteur `.pricing-card` devrait avoir la propriété `display` avec une valeur de `flex`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.display, 'flex');
```

Ton sélecteur `.pricing-card` devrait avoir la propriété `flex-direction` avec une valeur de `column`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.flexDirection, 'column');
```

Ton sélecteur `.pricing-card` devrait avoir la propriété `justify-content` avec une valeur de `space-between`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.justifyContent, 'space-between');
```

Ton sélecteur `.pricing-card` devrait avoir la propriété `flex` avec une valeur de `0 0 200px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.flex, '0 0 200px');
```

Ton sélecteur `.pricing-card` devrait avoir la propriété `border` avec une valeur de `2px solid black`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.border, '2px solid black');
```

Ton sélecteur `.basic-plan` devrait avoir la propriété `order` avec une valeur de `0`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.basic-plan')?.order, '0');
```

Ton élément `.basic-plan` devrait avoir un élément `h2` avec le texte `Basic`.

```js
const h2 = document.querySelector('.basic-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Basic');
```

Ton élément `.basic-plan` devrait avoir un élément `p` avec le texte `$9/month`.

```js
const p = document.querySelector('.basic-plan > p');
assert.strictEqual(p?.textContent.trim(), '$9/month');
```

Ton élément `.pro-plan` devrait avoir un élément `h2` avec le texte `Pro`.

```js
const h2 = document.querySelector('.pro-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Pro');
```

Ton élément `.pro-plan` devrait avoir un élément `p` avec le texte `$19/month`.

```js
const p = document.querySelector('.pro-plan > p');
assert.strictEqual(p?.textContent.trim(), '$19/month');
```

Ton élément `.pro-plan` devrait avoir la propriété `order` avec une valeur de `1`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pro-plan')?.order, '1');
```

Ton élément `.pro-plan` devrait avoir la propriété `flex-grow` avec une valeur de `2`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pro-plan')?.flexGrow, '2');
```

Ton élément `.premium-plan` devrait avoir un élément `h2` avec le texte `Premium`.

```js
const h2 = document.querySelector('.premium-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Premium');
```

Ton élément `.premium-plan` devrait avoir un élément `p` avec le texte `$29/month`.

```js
const p = document.querySelector('.premium-plan > p');
assert.strictEqual(p?.textContent.trim(), '$29/month');
```

Ton élément `premium-plan` devrait être le dernier élément de la mise en page et avoir la propriété `order` avec la valeur `2`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.premium-plan')?.order, '2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Pricing Plans Layout Page</title>
  </head>
  <body></body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Pricing Plans Layout</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Pricing Plans</h1>
    <div class="pricing-container">
      <div class="pricing-card pro-plan">
        <h2>Pro</h2>
        <p class="price">$19/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="check">✓</span> Advanced analytics</li>
          <li><span class="check">✓</span> Custom integrations</li>
          <li><span class="cross">✗</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>

      <div class="pricing-card basic-plan">
        <h2>Basic</h2>
        <p class="price">$9/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="cross">✗</span> Advanced analytics</li>
          <li><span class="cross">✗</span> Custom integrations</li>
          <li><span class="cross">✗</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>

      <div class="pricing-card premium-plan">
        <h2>Premium</h2>
        <p class="price">$29/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="check">✓</span> Advanced analytics</li>
          <li><span class="check">✓</span> Custom integrations</li>
          <li><span class="check">✓</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>
    </div>
  </body>
</html>
```

```css
body {
  display: flex;
  align-items: center;
  justify-content: center;
  flex-direction: column;
  min-height: 100vh;
}

h1 {
  margin-bottom: 20px;
}

h2,
p {
  margin: 0;
}

.pricing-container {
  display: flex;
  flex-wrap: wrap;
  align-content: center;
  border: 2px solid #000;
  gap: 20px;
  width: 80%;
  max-width: 900px;
  height: 70vh;
  padding: 0 20px;
}

.pricing-card {
  background: rgb(247, 247, 247);
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
  padding: 20px;
  flex: 0 0 200px;
  text-align: center;
  transition: transform 0.3s ease;
  display: flex;
  flex-direction: column;
  justify-content: space-between;
  border: 2px solid black;
}

.pricing-card:hover {
  transform: scale(1.05);
}

.price {
  font-size: 1.5rem;
  font-weight: bold;
  color: #3e3e3e;
  margin: 10px;
  border-bottom: 1px solid #433d28;
}

.features {
  list-style-type: none;
  padding: 0;
  text-align: left;
}

.features li {
  padding: 5px 0;
  font-size: 0.95rem;
}

.check {
  color: #219c54;
  font-weight: bold;
  margin-right: 8px;
}

.cross {
  color: crimson;
  font-weight: bold;
  margin-right: 8px;
}

a {
  text-decoration: none;
  display: inline-block;
  margin-top: 15px;
  background: #f1be32;
  color: #000;
  font-weight: bold;
  border: 3px solid #eebe3a;
  padding: 5px;
  transition: background 0.3s;
}

a:hover {
  background: #f7c436;
}

.basic-plan {
  order: 0;
}
.pro-plan {
  order: 1;
  flex-grow: 2;
  background: #f0f8ff;
}
.premium-plan {
  order: 2;
}

@media screen and (max-width: 1000px) {
  .pricing-container {
    height: auto;
    padding-bottom: 20px;
  }

  h1 {
    display: block;
    font-size: 1.5rem;
    text-align: center;
  }

  .pricing-card {
    flex: 0 0 175px;
  }
}

@media screen and (max-width: 650px) {
  .pricing-container {
    flex-direction: column;
    height: auto;
  }

  .pricing-card {
    flex: 1 1 auto;
    width: 100%;
  }
}
```
