---
id: 66a84111965a0c46df6bbd0a
title: Étape 6
challengeType: 0
dashedName: step-6
---

# --description--

Dans les leçons précédentes, tu as appris à associer un élément `label` à un `input` comme ceci :

```html
<label for="name">Name:</label>
<input type="text" id="name" name="name">
```

L'attribut `for` de l'élément `label` doit correspondre a l'attribut `id` de l'élément `input`. C'est ce qu on appelle une association explicite.

Sous ton élément `legend`, ajoute un élément `label` avec le texte `Name (required):`. Définis son attribut `for` avec la valeur `"full-name"`.

Ensuite, sous ton élément `label`, ajoute un élément `input` sans attribut. Dans les prochaines étapes, tu ajouteras les attributs nécessaires.

# --hints--

Tu dois avoir une balise ouvrante `label`.

```js
assert.match(code, /<label.*>/);
```

Tu dois avoir une balise fermante `label`.

```js
assert.match(code, /<\/label>/);
```

Ton élément `label` doit être place sous ton `legend`, dans l'élément `fieldset`.

```js
assert.isNotNull(document.querySelector('fieldset legend + label'));
```

Ton élément `label` doit avoir un attribut `for` avec la valeur `"full-name"`.

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.getAttribute('for'), 'full-name');
```

Ton élément `label` doit avoir le texte `Name (required):`.

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.innerText.trim(), 'Name (required):');
```

Tu dois avoir un élément `input` sous ton élément `label`.

```js
assert.isNotNull(document.querySelector('fieldset label + input'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
