---
id: 66a83bdcf425e7446900b7c4
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

Dans les leçons précédentes, tu as appris à travailler avec l'élément `form` comme ceci :

```html
<form method="value-goes-here" action="url-goes-here">
  <!-- inputs go inside here -->
</form>
```

L'attribut `action` sert à indiquer où les données du formulaire doivent être envoyées quand le formulaire est soumis.

L'attribut `method` sert à indiquer la méthode HTTP à utiliser pour envoyer les données du formulaire. Les méthodes les plus courantes sont `GET` et `POST`.

**NOTE** : Tu apprendras comment fonctionnent les méthodes HTTP dans des modules plus avancés.

Dans ton élément `main`, ajoute un élément `form` avec un attribut `action` défini sur `"https://hotel-feedback.freecodecamp.org"` et un attribut `method` défini sur `"POST"`.

# --hints--

Tu dois avoir une balise ouvrante `form`.

```js
assert.match(code, /<form.*>/i);
```

Tu dois avoir une balise fermante `form`.

```js
assert.match(code, /<\/form>/i);
```

Ton élément `form` doit avoir un attribut `action` défini sur `"https://hotel-feedback.freecodecamp.org"`.

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('action'), 'https://hotel-feedback.freecodecamp.org');
```

Ton élément `form` doit avoir un attribut `method` défini sur `"POST"`.

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('method'), 'POST');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </main>
  </body>
</html>
```
