---
id: 66a83601cd819e37f0dccd14
title: Étape 1
challengeType: 0
dashedName: step-1
---

# --description--

Dans cet atelier, tu vas t'entraîner à travailler avec les formulaires HTML en créant un formulaire de retour pour hôtel.

Pour le texte d'introduction, tu dois afficher le titre principal suivi d'une courte note invitant l'utilisateur à laisser un retour.

Dans ton élément `body`, ajoute un élément `header`.

Dans l'élément `header`, ajoute un élément `h1`. Le texte de l'élément `h1` doit être `Hotel Feedback Form`.

Sous ton élément `h1`, ajoute un élément `p`. Le texte de l'élément `p` doit être `Thank you for staying with us. Please provide feedback on your recent stay.`

# --hints--

Tu dois avoir une balise ouvrante `header`.

```js
assert.match(code, /<header>/i);
```

Tu dois avoir une balise fermante `header`.

```js
assert.match(code, /<\/header>/i);
```

Tu dois avoir une balise ouvrante `h1`.

```js
assert.match(code, /<h1>/i);
```

Tu dois avoir une balise fermante `h1`.

```js
assert.match(code, /<\/h1>/i);
```

Ton élément `h1` doit avoir le texte `Hotel Feedback Form`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.innerText.trim(), 'Hotel Feedback Form');
```

Ton élément `h1` doit être dans l'élément `header`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.parentElement.tagName, 'HEADER');
```

Tu dois avoir une balise ouvrante `p`.

```js
assert.match(code, /<p>/i);
```

Tu dois avoir une balise fermante `p`.

```js
assert.match(code, /<\/p>/i);
```

Ton élément de paragraphe doit avoir le texte `Thank you for staying with us. Please provide feedback on your recent stay.`

```js
const pElement = document.querySelector('header p');
assert.strictEqual(pElement?.innerText.trim(), 'Thank you for staying with us. Please provide feedback on your recent stay.');
```

Ton élément de paragraphe doit être dans ton `header`.

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement?.parentElement.tagName, 'HEADER');
```

Ton élément de paragraphe doit venir après ton `h1`.

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement?.previousElementSibling.tagName, 'H1');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
