---
id: 66a83fec026a7a4631e084d2
title: Étape 5
challengeType: 0
dashedName: step-5
---

# --description--

Quand tu travailles avec des éléments `fieldset`, il est courant d'utiliser une légende pour décrire le groupe de champs. Tu peux utiliser l'élément `legend` pour cela.

Voici un exemple d'utilisation d'un élément `legend` :

```html
<form action="/example-url">
  <fieldset>
    <legend>Personal Information</legend>
    <!-- inputs go inside here-->
  </fieldset>
</form>
```

Dans ton élément `fieldset`, ajoute un élément `legend` avec le texte `Personal Information`.

# --hints--

Tu dois avoir une balise ouvrante `legend`.

```js
assert.match(code, /<legend>/);
```

Tu dois avoir une balise fermante `legend`.

```js
assert.match(code, /<\/legend>/);
```

Ton élément `legend` doit être dans l'élément `fieldset`.

```js
assert.isNotNull(document.querySelector('fieldset legend'));
```

Ton élément `legend` doit contenir le texte `Personal Information`.

```js
const legend = document.querySelector('fieldset legend');
assert.strictEqual(legend?.innerText.trim(), 'Personal Information');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
