---
id: 66a8380d911e3f4270d5cadc
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

Maintenant, il est temps d'ajouter l'élément `main`, qui représente le contenu principal de la page.

# --hints--

Tu dois avoir une balise ouvrante `main`.

```js
assert.match(code, /<main>/i);
```

Tu dois avoir une balise fermante `main`.

```js
assert.match(code, /<\/main>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>
```
