---
id: 66a83e5e491625454b6f62c3
title: Étape 4
challengeType: 0
dashedName: step-4
---

# --description--

Les formulaires sont composés d'`inputs` dans lesquels les utilisateurs peuvent saisir leurs données. Tu peux regrouper des champs liés entre eux avec l'élément `fieldset`.

Voici un exemple d'utilisation d'un élément `fieldset` :

```html
<form action="/example-url">
  <fieldset>
  <!-- inputs go inside here-->
  </fieldset>
</form>
```

Dans ton élément `form`, ajoute un élément `fieldset`.

# --hints--

Tu dois avoir une balise ouvrante `fieldset`.

```js
assert.match(code, /<fieldset>/i);
```

Tu dois avoir une balise fermante `fieldset`.

```js
assert.match(code, /<\/fieldset>/i);
```

Ton élément `fieldset` doit être dans l'élément `form`.

```js
assert.isNotNull(document.querySelector('form fieldset'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">        
--fcc-editable-region--
        
--fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
