---
id: 66a9419e2d18476c645ce693
title: Étape 14
challengeType: 0
dashedName: step-14
---

# --description--

L'`input` de type nombre sert à créer un champ de saisie numérique.

Voici un exemple de champ de saisie numérique :

```html
<input type="number" id="age" name="age" min="18" max="100">
```

Les attributs `min` et `max` servent à définir les valeurs minimale et maximale qui peuvent être saisies dans le champ.

Sous ton élément `label`, ajoute un input avec l'attribut `type` défini sur `"number"` et un `id` de `"age"`.

L'attribut `name` doit être défini sur `"age"`, l'attribut `min` doit être défini sur `"3"` et l'attribut `max` doit être défini sur `"100"`.

# --hints--

Tu dois avoir un `input` de type nombre sous ton élément `label`.

```js
assert.isNotNull(document.querySelector('label + input[type="number"]'));
```

Ton `input` de type nombre doit avoir un `id` de `age`.

```js
assert.isNotNull(document.querySelector('label + input#age'));
```

Ton `input` de type nombre doit avoir un attribut `name` défini sur `age`.

```js
assert.isNotNull(document.querySelector('label + input[name="age"]'));
```

Ton `input` de type nombre doit avoir un attribut `min` défini sur `"3"`.

```js
assert.isNotNull(document.querySelector('label + input[min="3"]'));
```

Ton `input` de type nombre doit avoir un attribut `max` défini sur `"100"`.

```js
assert.isNotNull(document.querySelector('label + input[max="100"]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
      --fcc-editable-region--
          
      --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
