---
id: 66a937e74920ba68ebe5e86d
title: Étape 9
challengeType: 0
dashedName: step-9
---

# --description--

Dans les leçons précédentes, tu as appris à travailler avec les attributs `placeholder` et `required` comme ceci :

```html
<input type="text" placeholder="e.g., John Doe" required>
```

Pour ton élément `input` existant, ajoute un attribut `placeholder` avec la valeur `"e.g., John Doe"`.

Ajoute aussi l'attribut `required` à l'élément `input`.

# --hints--

Tu ne dois pas supprimer l'élément `input` existant.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

Ton `input` doit avoir un attribut `placeholder` avec la valeur `"e.g., John Doe"`.

```js
assert.strictEqual(document.querySelector('fieldset input')?.getAttribute('placeholder'), 'e.g., John Doe');
```

Ton `input` doit avoir un attribut `required`.

```js
assert.isNotNull(document.querySelector('fieldset input[required]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name" name="name">
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
