---
id: 66a95d0eff8168747805f1f3
title: Étape 19
challengeType: 0
dashedName: step-19
---

# --description--

Quand tu veux qu'un utilisateur sélectionne plusieurs options dans une liste, tu peux utiliser des cases à cocher.

Voici un exemple d'utilisation de cases à cocher pour des options de nourriture :

```html
<fieldset>
  <legend>Food Options</legend>
  <input type="checkbox" id="pizza" name="food" value="pizza">
  <label for="pizza">Pizza</label>
  <input type="checkbox" id="burger" name="food" value="burger">
  <label for="burger">Burger</label>
</fieldset>
```

L'attribut `value` sert à indiquer la valeur qui sera envoyée au serveur quand le formulaire est soumis.

Sous ton élément `legend`, ajoute un `input` de type case à cocher avec les attributs `id` et `value` définis sur `"ads"`, et l'attribut `name` défini sur `"choice"`.

Sous ton `input` de type case à cocher, ajoute un élément `label` avec le texte `Social Media Ads`. L'attribut `for` doit être défini sur `"ads"`.

# --hints--

Tu dois avoir un `input` de type case à cocher sous ton élément `legend`.

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']"));
```

Ton `input` de type case à cocher doit avoir un attribut `id` défini sur `"ads"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']")?.id, "ads");
```

Ton `input` de type case à cocher doit avoir un attribut `name` défini sur `"choice"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']")?.name, "choice");
```

Ton `input` de type case à cocher doit avoir un attribut `value` défini sur `"ads"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']")?.value, "ads");
```

Tu dois avoir un élément `label` sous ton `input` de type case à cocher.

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label"));
```

Ton élément `label` doit avoir le texte `Social Media Ads`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label")?.textContent?.trim(), "Social Media Ads");
```

Ton élément `label` doit avoir un attribut `for` défini sur `"ads"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label")?.getAttribute("for"), "ads");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
