---
id: 66a9401c9d660d6bb15993e2
title: Étape 13
challengeType: 0
dashedName: step-13
---

# --description--

Ton formulaire de retour pour hôtel doit proposer une option pour que les utilisateurs puissent ajouter leur age.

Commence par ajouter au formulaire un élément `label` avec le texte `Age (optional):`.

L'attribut `for` doit être défini sur `"age"`.

# --hints--

Tu dois avoir un élément `label`.

```js
assert.lengthOf(document.querySelectorAll('fieldset label'), 3);
```

Ton élément `label` doit avoir un attribut `for` défini sur `"age"`.

```js
assert.isNotNull(document.querySelector('label[for="age"]'));
```

Ton élément `label` doit avoir le texte `Age (optional):`.

```js
assert.strictEqual(document.querySelector('label[for="age"]')?.textContent.trim(), "Age (optional):");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
