---
id: 66a930b20f589b6664c51cb0
title: Étape 7
challengeType: 0
dashedName: step-7
---

# --description--

Quand un utilisateur fournit son nom complet, l'`input` accepte du texte brut.

Dans les leçons précédentes, tu as appris à travailler avec l'attribut `type` comme ceci :

```html
<input type="text">
```

Pour ton élément `input` existant, ajoute un attribut `type` défini sur `"text"`.

Ajoute aussi à l'élément `input` un attribut `id` avec la valeur `"full-name"`.

# --hints--

Tu ne dois pas supprimer l'élément `input` existant.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

Tu ne dois pas créer un nouvel élément `input`.

```js
assert.lengthOf(document.querySelectorAll('fieldset input'), 1);
```

Ton élément `input` doit avoir un attribut `type` avec la valeur `"text"`.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('type'), 'text');
```

Ton élément `input` doit avoir un attribut `id` avec la valeur `"full-name"`.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('id'), 'full-name');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
        --fcc-editable-region--
          <input>
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
