---
id: 66a972137acd1179fa3fe8a0
title: Étape 26
challengeType: 0
dashedName: step-26
---

# --description--

Dans ton élément `select`, ajoute les cinq éléments `option` suivants avec ces valeurs correspondantes pour le texte de l'`option` et l'attribut `value` :

**Attributs value :**

- poor
- satisfactory
- good
- very-good
- excellent

**Texte des éléments option :**

- Poor
- Satisfactory
- Good
- Very Good
- Excellent

# --hints--

Tu dois avoir un élément `option` avec la valeur définie sur `"poor"`.

```js
assert.exists(document.querySelector('option[value="poor"]'));
```

Ton élément `option` avec la `value` de `"poor"` doit avoir le texte `Poor`.

```js
assert.strictEqual(document.querySelector('option[value="poor"]')?.textContent.trim(), 'Poor');
```

Tu dois avoir un élément `option` avec la `value` définie sur `"satisfactory"`.

```js
assert.exists(document.querySelector('option[value="satisfactory"]'));
```

Ton élément `option` avec la `value` de `"satisfactory"` doit avoir le texte `Satisfactory`.

```js
assert.strictEqual(document.querySelector('option[value="satisfactory"]')?.textContent.trim(), 'Satisfactory');
```

Tu dois avoir un élément `option` avec la `value` définie sur `"good"`.

```js
assert.exists(document.querySelector('option[value="good"]'));
```

Ton élément `option` avec la `value` de `"good"` doit avoir le texte `Good`.

```js

assert.strictEqual(document.querySelector('option[value="good"]')?.textContent.trim(), 'Good');
```

Tu dois avoir un élément `option` avec la valeur définie sur `"very-good"`.

```js
assert.exists(document.querySelector('option[value="very-good"]'));
```

Ton élément `option` avec la `value` de `"very-good"` doit avoir le texte `Very Good`.

```js
assert.strictEqual(document.querySelector('option[value="very-good"]')?.textContent.trim(), 'Very Good');
```

Tu dois avoir un élément `option` avec la valeur définie sur `"excellent"`.

```js
assert.exists(document.querySelector('option[value="excellent"]'));
```

Ton élément `option` avec la `value` de `"excellent"` doit avoir le texte `Excellent`.

```js
assert.strictEqual(document.querySelector('option[value="excellent"]')?.textContent.trim(), 'Excellent');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="choice" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="choice"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>
          
          <input type="checkbox" id="location" name="choice" value="location" />
          <label for="location">Location</label>

          <input
            checked
            type="checkbox"
            id="reputation"
            name="choice"
            value="reputation"
          />
          <label for="reputation">Reputation</label>
          
          <input type="checkbox" id="price" name="choice" value="price" />
          <label for="price">Price</label>
        </fieldset>

        <fieldset>
          <legend>Ratings</legend>

          <label for="service">How was the service?</label>

          <select name="service" id="service">
          --fcc-editable-region--
            
          --fcc-editable-region--
          </select>
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
