---
id: 66a93bbe65a26169dbf3bc39
title: Étape 10
challengeType: 0
dashedName: step-10
---

# --description--

Ton formulaire de retour pour hôtel doit aussi collecter une adresse e-mail de l'utilisateur.

Commence par ajouter au formulaire un nouvel élément `label` avec le texte `Email address (required):`. Ton élément `label` doit avoir un attribut `for` défini sur la valeur `"email"`.

# --hints--

Tu dois avoir un élément `label`.

```js
assert.lengthOf(document.querySelectorAll('fieldset label'), 2);
```

Ton élément `label` doit avoir le texte `Email address (required):`.

```js
assert.strictEqual(document.querySelector('fieldset label:nth-of-type(2)')?.innerText.trim(), 'Email address (required):');
```

Ton élément `label` doit avoir un attribut `for` défini sur la valeur `"email"`.

```js
assert.strictEqual(document.querySelector('fieldset label[for="email"]')?.getAttribute('for'), 'email');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
