---
id: 66a93730719e1f68410cce54
title: Étape 8
challengeType: 0
dashedName: step-8
---

# --description--

L'attribut `name` sert à identifier les données du formulaire après leur envoi au serveur.

Voici un exemple d'utilisation de l'attribut `name` :

```html
<input type="email" name="email">
```

Ajoute un attribut `name` à l'élément `input` avec la valeur `"name"`.

# --hints--

Tu ne dois pas supprimer l'élément `input` existant.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

Tu ne dois pas créer un nouvel élément `input`.

```js
assert.lengthOf(document.querySelectorAll('fieldset input'), 1);
```

Ton élément `input` doit avoir un attribut `name` avec la valeur `"name"`.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('name'), 'name');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name">
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
