---
id: 66a9521bc70162712caf118d
title: Étape 16
challengeType: 0
dashedName: step-16
---

# --description--

Si tu veux que les utilisateurs sélectionnent une seule option dans une liste, tu peux utiliser un groupe de boutons radio.

Voici un exemple avec deux boutons radio :

```html
<input type="radio" id="yes" name="first-time">
<label for="yes">Yes</label>
<input type="radio" id="no" name="first-time">
<label for="no">No</label>
```

Dans cet exemple, les boutons radio sont regroupes en utilisant la même valeur d'attribut `name`. Cela signifie qu'un seul bouton radio peut être sélectionne à la fois.

Sous ton élément `legend`, ajoute un bouton `radio` avec l'`id` défini sur `"yes-option"`, l'attribut `name` défini sur `"hotel-stay"` et un attribut `value` défini sur `"yes"`.

Sous ton bouton `radio`, ajoute un élément `label` avec le texte `Yes` et un attribut `for` défini sur `"yes-option"`.

# --hints--

Tu dois avoir un bouton `radio` sous ton élément `legend`.

```js
assert.exists(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]'));
```

Ton bouton `radio` doit avoir un attribut `id` défini sur `"yes-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.id, 'yes-option');
```

Ton bouton `radio` doit avoir un attribut `name` défini sur `"hotel-stay"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.name, 'hotel-stay');
```

Ton bouton radio doit avoir un attribut `value` défini sur `"yes"`.

```js
assert.strictEqual(
  document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.value,
  'yes'
);
```

Tu dois avoir un élément `label` sous ton bouton `radio`.

```js
assert.exists(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label'));
```

Ton élément `label` doit avoir un attribut `for` défini sur `"yes-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.htmlFor, 'yes-option');
```

Ton élément `label` doit avoir le texte `Yes`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.innerText?.trim(), 'Yes');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
