---
id: 66a93c95bc58e26a8fe95818
title: Étape 11
challengeType: 0
dashedName: step-11
---

# --description--

Ensuite, ajoute un `input` avec le type `"email"` sous ton `label` d'e-mail. Cet `input` doit avoir un attribut `id` défini sur la valeur `"email"`. Définis aussi l'attribut `name` sur la valeur `"email"`.

Cet `input` est aussi obligatoire, donc assure-toi d'ajouter l'attribut `required`.

Enfin, ajoute un attribut `placeholder` défini sur `"example@email.com"`.

# --hints--

Tu dois avoir un élément `input` sous ton `label` d'e-mail.

```js
assert.isNotNull(document.querySelector('label[for="email"] + input'));
```

Ton élément `input` doit avoir un `type` de `"email"`.

```js
assert.isNotNull(document.querySelector('input[type="email"]'));
```

Ton élément `input` doit avoir un attribut `id` défini sur `"email"`.

```js
assert.isNotNull(document.querySelector('input#email'));
```

Ton élément `input` doit avoir un attribut `name` défini sur `"email"`.

```js
assert.isNotNull(document.querySelector('input[name="email"]'));
```

Ton élément `input` doit avoir un attribut `required`.

```js
assert.strictEqual(document.querySelector('input#email')?.getAttribute('required'), "");
```

Ton élément `input` doit avoir un attribut `placeholder` défini sur `"example@email.com"`.
  
```js
assert.strictEqual(document.querySelector('input#email')?.getAttribute('placeholder'), "example@email.com");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
