---
id: 67a51d3a8a6fe123b77b0c6e
title: Étape 12
challengeType: 0
dashedName: step-12
---

# --description--

Les éléments `input` peuvent avoir un attribut `size`. Cet attribut définit le nombre de caractères visibles pendant que l'utilisateur saisit du texte dans le champ. La valeur de `size` doit être un entier positif supérieur à zéro. Si `size` n'est pas indiqué, ou si sa valeur est invalide, le champ aura la largeur par défaut définie par le navigateur.

```html
<label for="lastName">Last Name:</label>
<input id="lastName" name="lastName" type="text" size="10" />
```

Donne aux champs de nom et d'e-mail un attribut `size` avec la valeur `"20"`.

# --hints--

Tu dois donner à l'`input` du nom un attribut `size` avec la valeur `"20"`.

```js
const label1 = document.querySelector('label:nth-of-type(1) + input');
assert.strictEqual(label1?.getAttribute('size'), '20');
```

Tu dois donner à l'`input` d'e-mail un attribut `size` avec la valeur `"20"`.

```js
const label2 = document.querySelector('label:nth-of-type(2) + input');
assert.strictEqual(label2?.getAttribute('size'), '20');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
          />         
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
