---
id: af7588ade1100bde429baf20
title: Crée un détecteur de lettre manquante
challengeType: 26
dashedName: build-a-missing-letter-detector
---

# --description--

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois avoir une fonction nommée `fearNotLetter`.
1. La fonction `fearNotLetter` doit accepter un argument : une chaîne représentant une plage de lettres dans l'ordre alphabétique à laquelle il peut manquer une lettre.
1. La fonction doit trouver la lettre manquante dans la plage de lettres fournie et la retourner.
1. Si toutes les lettres sont présentes dans la plage, la fonction doit retourner `undefined`.

# --hints--

Tu dois avoir une fonction `fearNotLetter`.

```js
assert.isFunction(fearNotLetter);
```

`fearNotLetter("abce")` doit retourner la chaîne `d`.

```js
assert.deepEqual(fearNotLetter('abce'), 'd');
```

`fearNotLetter("abcdefghjklmno")` doit retourner la chaîne `i`.

```js
assert.deepEqual(fearNotLetter('abcdefghjklmno'), 'i');
```

`fearNotLetter("stvwx")` doit retourner la chaîne `u`.

```js
assert.deepEqual(fearNotLetter('stvwx'), 'u');
```

`fearNotLetter("bcdf")` doit retourner la chaîne `e`.

```js
assert.deepEqual(fearNotLetter('bcdf'), 'e');
```

`fearNotLetter("abcdefghijklmnopqrstuvwxyz")` doit retourner `undefined`.

```js
assert.isUndefined(fearNotLetter('abcdefghijklmnopqrstuvwxyz'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function fearNotLetter (str) {
  for (var i = str.charCodeAt(0); i <= str.charCodeAt(str.length - 1); i++) {
    let letter = String.fromCharCode(i);
    if (str.indexOf(letter) === -1) {
      return letter;
    }
  }

  return undefined;
}
```
