---
id: 68e6aff5325410118e76e7be
title: Étape 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Dans cet atelier, tu vas t'entraîner à utiliser ARIA en construisant un contrôleur audio accessible.

Pour commencer, ajoute un élément `h1` avec le texte `Audio Controls`.

# --hints--

Tu devrais avoir un élément `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Ton élément `h1` devrait être à l'intérieur de ton élément `body`.

```js
assert.exists(document.querySelector('body > h1'));
```

Ton élément `h1` devrait avoir le texte `Audio Controls`.

```js
const h1Text = document.querySelector('h1')?.innerText.trim();
assert.strictEqual(h1Text, 'Audio Controls');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
