---
id: 68eac57fadcc662ce917e75e
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

Ensuite, ajoute un élément `button` avec l'attribut `type` défini sur la valeur `button`.

Le texte de l'élément `button` devrait être `Play`.

# --hints--

Tu devrais avoir un élément `button`.

```js
assert.exists(document.querySelector('button'));
```

Ton élément `button` devrait avoir un attribut `type`.

```js
const button = document.querySelector('button');
assert.isTrue(button?.hasAttribute('type'));
```

Ton élément `button` devrait avoir l'attribut `type` défini sur `button`.

```js
const button = document.querySelector('button');
assert.strictEqual(button?.getAttribute('type'), 'button');
```

Ton élément `button` devrait avoir le texte `Play`.

```js
const button = document.querySelector('button');
assert.strictEqual(button?.innerText.trim(), 'Play');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
