---
id: 68eb7f2397da251a40e04aa2
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

Ensuite, ajoute un élément `div` pour contenir tous les éléments du contrôle du volume.

À l'intérieur de l'élément `div`, imbrique un élément `span` avec un attribut `id` défini sur `volume-label`.
Le texte de l'élément `span` devrait aussi être `Volume`.

L'`id` sera ensuite référencé par des attributs ARIA afin que les technologies d'assistance puissent identifier ce texte comme le libellé du contrôle du volume.

# --hints--

Tu devrais avoir un élément `div` sous ton élément `button`.

```js
assert.exists(document.querySelector('button + div'));
```

Tu devrais avoir un élément `span` à l'intérieur de ton élément `div`.

```js
assert.exists(document.querySelector('button + div > span'));
```

Ton élément `span` devrait avoir un attribut `id`.

```js
const spanEl = document.querySelector('button + div > span');
assert.isTrue(spanEl?.hasAttribute('id'));
```

Ton élément `span` devrait avoir un attribut `id` dont la valeur est `volume-label`.

```js
const spanEl = document.querySelector('button + div > span');
assert.equal(spanEl?.getAttribute('id'), 'volume-label');
```

Ton élément `span` devrait avoir le texte `Volume`.

```js
const spanEl = document.querySelector('button + div > span');
assert.equal(spanEl?.innerText.trim(), 'Volume');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

--fcc-editable-region--

--fcc-editable-region--

  </body>
</html>
```
