---
id: 68ec0a48c9d6db3594d768c3
title: Étape 4
challengeType: 0
dashedName: step-4
---

# --description--

Il est maintenant temps d'ajouter un autre élément `span`.

Cet élément `span` devrait avoir un attribut `id` défini sur `volume-description`.
Le texte de cet élément `span` devrait aussi être `Adjust the sound level`.

Comme pour l'élément `span` précédent, l'`id` de ce `span` sera ensuite référencé par des attributs ARIA afin que les technologies d'assistance puissent identifier ce texte comme la description du contrôle du volume.

# --hints--

Tu devrais avoir deux éléments `span` à l'intérieur de l'élément `div`.

```js
const spans = document.querySelectorAll('div > span');
assert.lengthOf(spans, 2);
```

Ton deuxième élément `span` devrait avoir un attribut `id`.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.isTrue(secondSpanEl?.hasAttribute('id'));
```

Ton deuxième élément `span` devrait avoir un attribut `id` avec la valeur `volume-description`.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.equal(secondSpanEl?.getAttribute('id'), 'volume-description');
```

Ton deuxième élément `span` devrait avoir le texte `Adjust the sound level`.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.equal(secondSpanEl?.innerText.trim(), 'Adjust the sound level');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      --fcc-editable-region--

      --fcc-editable-region--
    </div>

  </body>
</html>
```
