---
id: 68ed3a96e0ebcc36565b39e5
title: Étape 5
challengeType: 0
dashedName: step-5
---

# --description--

Ensuite, tu dois ajouter un élément `input` avec l'attribut `type` défini sur `range` pour créer le curseur de volume. Définis l'attribut `min` sur `0`, l'attribut `max` sur `100` et l'attribut `value` sur `50` pour définir le niveau de volume par défaut.

# --hints--

Tu devrais avoir un élément `input` sous ton deuxième élément `span`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl);
```

Ton élément `input` devrait avoir un attribut `type`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('type'));
```

Ton élément `input` devrait avoir un attribut `type` avec la valeur `range`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('type'), 'range');
```

Ton élément `input` devrait avoir un attribut `min`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl?.getAttribute('min'));
```

Ton élément `input` devrait avoir un attribut `min` avec la valeur `0`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('min'), '0');
```

Ton élément `input` devrait avoir un attribut `max`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('max'));
```

Ton élément `input` devrait avoir un attribut `max` avec la valeur `100`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('max'), '100');
```

Ton élément `input` devrait avoir un attribut `value`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('value'));
```

Ton élément `input` devrait avoir un attribut `value` avec la valeur `50`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('value'), '50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
    --fcc-editable-region--

    --fcc-editable-region--
    </div>

  </body>
</html>
```
