---
id: 68ed4204ab7edc66b6cb2383
title: Étape 6
challengeType: 0
dashedName: step-6
---

# --description--

Comme tu t'en souviens des leçons précédentes, l'attribut `aria-labelledby` est utilisé quand il existe déjà du texte sur la page pouvant servir de libellé. Dans ce cas, tu as ajouté l'attribut `id` à tes éléments `span`, qui vont maintenant être utilisés ici.

Ajoute un attribut `aria-labelledby` à l'élément `input` et définis-le sur `volume-label volume-description`.

# --hints--

Ton élément `input` devrait avoir un attribut `aria-labelledby`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('aria-labelledby'));
```

Ton élément `input` devrait avoir un attribut `aria-labelledby` avec la valeur `volume-label volume-description`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
const ariaLabelledbyAttr = inputEl?.getAttribute('aria-labelledby');
assert.oneOf(ariaLabelledbyAttr, ['volume-label volume-description', 'volume-description volume-label']);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
    --fcc-editable-region--
    <input type="range" min="0" max="100" value="50">
    --fcc-editable-region--
    </div>

  </body>
</html>
```
