---
id: 68ed488821845990c911a223
title: Étape 7
challengeType: 0
dashedName: step-7
---

# --description--

Ton contrôleur audio accessible est presque terminé, mais il reste une dernière chose à ajouter.

Ajoute un élément `button` avec l'attribut `type` défini sur `button` et le texte `Mute` sous l'élément `div`.

Avec ce changement, tu as terminé cet atelier.

# --hints--

Tu devrais avoir un élément `button` sous l'élément `div`.

```js
const buttonEl = document.querySelector('div + button');
assert.exists(buttonEl);
```

Ton élément `button` devrait avoir un attribut `type`.

```js
const buttonEl = document.querySelector('div + button');
assert.isTrue(buttonEl?.hasAttribute('type'));
```

Ton élément `button` devrait avoir un attribut `type` avec la valeur `button`.

```js
const buttonEl = document.querySelector('div + button');
assert.equal(buttonEl?.getAttribute('type'), 'button');
```

Ton élément `button` devrait avoir le texte `Mute`.

```js
const buttonEl = document.querySelector('div + button');
assert.equal(buttonEl?.innerText.trim(), 'Mute');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
      <input type="range" min="0" max="100" value="50"
           aria-labelledby="volume-label volume-description">
    </div>
  --fcc-editable-region--

  --fcc-editable-region--

  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
      <input type="range" min="0" max="100" value="50"
             aria-labelledby="volume-label volume-description">
    </div>

    <button type="button">Mute</button>

  </body>
</html>
```
