---
id: 6976db7ecfa770bf21307b20
title: Crée un moteur de remix de playlists
challengeType: 26
dashedName: build-a-playlist-remix-engine
---

# --description--

Dans ce lab, tu vas créer un programme qui génère une seule playlist remix à partir de plusieurs playlists soumises par les auditeurs.

Chaque auditeur fournit une liste de chansons qu'il souhaite entendre. Certaines chansons peuvent apparaître plus d'une fois, et certains artistes peuvent être trop présents. Ton travail consiste à traiter ces playlists étape par étape : les combiner en une seule liste, noter chaque chanson, supprimer les doublons, limiter la fréquence d'apparition d'un même artiste, puis créer un ordre de lecture final.

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois créer une fonction nommée `flattenPlaylists` qui accepte un tableau de playlists, chaque playlist étant un tableau d'objets avec les propriétés suivantes : `trackId`, `artist`, `title`, `votes`, `bpm`. Si l'entrée n'est pas un tableau, `flattenPlaylists` doit renvoyer un tableau vide. Un exemple de playlist t'a été fourni. Tu peux utiliser cet exemple pour tester ta fonction.

2. `flattenPlaylists` doit renvoyer un tableau plat d'objets piste, où chaque objet inclut toutes les propriétés originales de la piste plus une propriété `source` définie comme un tableau contenant l'indice de la playlist et l'indice de la piste indiquant l'origine de la piste.

3. Tu dois créer une fonction nommée `scoreTracks` qui accepte un tableau d'objets piste tels que renvoyés par `flattenPlaylists` (chacun avec les propriétés `trackId`, `artist`, `title`, `votes`, `bpm` et `source`) et renvoie un nouveau tableau d'objets piste, chacun avec une propriété `score` ajoutée selon la formule : `votes * 10 - Math.abs(bpm - 120)`.

4. Tu dois créer une fonction nommée `dedupeTracks` qui accepte un tableau d'objets piste tels que renvoyés par `scoreTracks` et renvoie un nouveau tableau sans les entrées `trackId` en double, en ne conservant que la première occurrence de chacune.

5. Tu dois créer une fonction nommée `enforceArtistQuota` qui accepte un tableau d'objets piste tels que renvoyés par `dedupeTracks` et un nombre représentant le nombre maximal d'occurrences autorisées par artiste. La fonction doit renvoyer un nouveau tableau où aucun artiste n'apparaît plus de fois que le nombre donné, en conservant les occurrences les plus anciennes.

6. Tu dois créer une fonction nommée `buildSchedule` qui accepte un tableau d'objets piste tels que renvoyés par `enforceArtistQuota` et renvoie un nouveau tableau d'objets `{ slot, trackId }`, où `slot` est un indice commençant à 1 représentant la position de chaque piste dans l'ordre de diffusion.

7. Tu dois créer une fonction nommée `remixPlaylist` qui accepte un tableau de playlists et le nombre maximal d'occurrences autorisées par artiste. La fonction doit renvoyer le planning de diffusion final sous forme de tableau d'objets `{ slot, trackId }`, en appelant `flattenPlaylists`, `scoreTracks`, `dedupeTracks`, `enforceArtistQuota` et `buildSchedule` dans cet ordre.

# --hints--

Tu dois avoir une fonction nommée `flattenPlaylists`.

```js
assert.isFunction(flattenPlaylists);
```

Tu dois renvoyer un tableau vide depuis `flattenPlaylists` lorsque l'entrée n'est pas un tableau.

```js
assert.deepEqual(flattenPlaylists(null), []);
assert.deepEqual(flattenPlaylists(undefined), []);
assert.deepEqual(flattenPlaylists({}), []);
assert.deepEqual(flattenPlaylists("not an array"), []);
assert.deepEqual(flattenPlaylists(123), []);
```

Chaque piste renvoyée par `flattenPlaylists` doit inclure un champ `source` qui est un tableau contenant l'indice de la playlist et l'indice de la piste.

```js
const playlists = [
  [
    { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120 },
    { trackId: "t2", artist: "B", title: "Song 2", votes: 1, bpm: 110 }
  ],
  [{ trackId: "t3", artist: "C", title: "Song 3", votes: 5, bpm: 130 }]
];

const flat = flattenPlaylists(playlists);

assert.isArray(flat);
assert.lengthOf(flat, 3);

flat.forEach((t) => {
  assert.property(t, "source");
  assert.isArray(t.source);
  assert.lengthOf(t.source, 2);
  assert.isNumber(t.source[0]);
  assert.isNumber(t.source[1]);
});

assert.deepEqual(flat[0].source, [0, 0]);
assert.deepEqual(flat[1].source, [0, 1]);
assert.deepEqual(flat[2].source, [1, 0]);
```

Tu dois avoir une fonction nommée `scoreTracks`.

```js
assert.isFunction(scoreTracks)
```

Chaque piste renvoyée par `scoreTracks` doit inclure un champ numérique `score`.

```js
const tracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120, source: [0, 0] }
];

const scored = scoreTracks(tracks);

assert.isArray(scored);
assert.lengthOf(scored, 1);

assert.property(scored[0], "score");
assert.isNumber(scored[0].score);
```

Tu dois calculer `score` en utilisant un BPM cible de `120` et cette formule : `votes * 10 - Math.abs(bpm - 120)`.

```js
const tracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 3, bpm: 120, source: [0, 0] }, // 30
  { trackId: "t2", artist: "B", title: "Song 2", votes: 1, bpm: 100, source: [0, 1] },  // -10
  { trackId: "t3", artist: "C", title: "Song 3", votes: 2, bpm: 135, source: [1, 0] }  // 5
];

const formulaScored = scoreTracks(tracks);

const s1 = formulaScored.find(t => t.trackId === "t1").score;
const s2 = formulaScored.find(t => t.trackId === "t2").score;
const s3 = formulaScored.find(t => t.trackId === "t3").score;

assert.equal(s1, 30);
assert.equal(s2, -10);
assert.equal(s3, 5);
```

Tu dois avoir une fonction nommée `dedupeTracks`.

```js
assert.isFunction(dedupeTracks)
```

Lorsque des valeurs `trackId` en double existent, `dedupeTracks` doit ne conserver que la première occurrence de la piste.

```js
const dupes = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 1, bpm: 120, source: [0, 0], score: 10 },
  { trackId: "t2", artist: "B", title: "Song 2", votes: 2, bpm: 119, source: [0, 1], score: 19 },
  { trackId: "t1", artist: "A", title: "Song 3", votes: 9, bpm: 140, source: [1, 0], score: 70 }
];

const deduped = dedupeTracks(dupes);

assert.isArray(deduped);
assert.lengthOf(deduped, 2);

assert.equal(deduped[0].trackId, "t1");
assert.equal(deduped[0].title, "Song 1");
assert.equal(deduped[1].trackId, "t2");
assert.equal(deduped[1].title, "Song 2");

```

Tu dois avoir une fonction nommée `enforceArtistQuota`.

```js
assert.isFunction(enforceArtistQuota)
```

`enforceArtistQuota` doit s'assurer qu'aucun artiste n'apparaît plus de `maxPerArtist` fois en supprimant les pistes en trop tout en conservant les plus anciennes.

```js
const quotaTracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 1, bpm: 120, source: [0, 0], score: 10 },
  { trackId: "t2", artist: "A", title: "Song 2", votes: 1, bpm: 121, source: [0, 1], score: 9 },
  { trackId: "t3", artist: "B", title: "Song 3", votes: 1, bpm: 118, source: [0, 2], score: 8 },
  { trackId: "t4", artist: "A", title: "Song 4", votes: 1, bpm: 110, source: [1, 0], score: 0 }
];

const limited = enforceArtistQuota(quotaTracks, 2);

let artistACount = 0;

for (const track of limited) {
  if (track.artist === "A") {
    artistACount++;
  }
}
assert.isAtMost(artistACount, 2);
assert.notInclude(limited.map(t => t.trackId), "t4");
```

Tu dois avoir une fonction nommée `buildSchedule`.

```js
assert.isFunction(buildSchedule);
```

`buildSchedule` doit renvoyer un tableau d'objets de la forme `{ slot, trackId }`, où `slot` commence à `1`.

```js
const scheduleInput = [{ trackId: "t1" }, { trackId: "t2" }, { trackId: "t3" }];
const schedule = buildSchedule(scheduleInput);

assert.lengthOf(schedule, 3);

schedule.forEach((item, index) => {
  assert.hasAllKeys(item, ["slot", "trackId"]);
  assert.equal(item.slot, index + 1);
  assert.equal(item.trackId, scheduleInput[index].trackId);
});
```

Tu dois avoir une fonction nommée `remixPlaylist`.

```js
assert.isFunction(remixPlaylist)
```

`remixPlaylist` doit appeler les fonctions auxiliaires dans l'ordre pour produire le planning final.

```js
assert.lengthOf(remixPlaylist, 2);

const testPlaylists = [
  [
    { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120 },
    { trackId: "t1", artist: "A", title: "Duplicate", votes: 9, bpm: 140 }
  ],
  [
    { trackId: "t2", artist: "A", title: "Song 2", votes: 1, bpm: 110 },
    { trackId: "t3", artist: "B", title: "Song 3", votes: 4, bpm: 118 }
  ]
];

const expected = buildSchedule(enforceArtistQuota(dedupeTracks(scoreTracks(flattenPlaylists(testPlaylists))),1));

const actual = remixPlaylist(testPlaylists, 1);

assert.deepEqual(actual, expected);
```

# --seed--

## --seed-contents--

```js
const playlists = [
  [
    {
      trackId: "trk101",
      artist: "Velvet Comet",
      title: "Crimson Afterglow",
      votes: 5,
      bpm: 122
    },
    {
      trackId: "trk102",
      artist: "Neon Harbor",
      title: "Static Horizon",
      votes: 2,
      bpm: 108
    },
    {
      trackId: "trk103",
      artist: "Lunar Arcade",
      title: "Midnight Frequency",
      votes: 4,
      bpm: 128
    }
  ],
  [
    {
      trackId: "trk201",
      artist: "Solar Echo",
      title: "Glass Skyline",
      votes: 3,
      bpm: 115
    },
    {
      trackId: "trk202",
      artist: "Velvet Comet",
      title: "Satellite Hearts",
      votes: 6,
      bpm: 124
    }
  ]
];
```

# --solutions--

```js
function flattenPlaylists(playlists) {
  if (!Array.isArray(playlists)) return [];

  const result = [];

  for (let i = 0; i < playlists.length; i++) {
    const playlist = playlists[i];
    if (!Array.isArray(playlist)) continue;

    for (let j = 0; j < playlist.length; j++) {
      const track = playlist[j];
      result.push({
        trackId: track.trackId,
        artist: track.artist,
        title: track.title,
        votes: track.votes,
        bpm: track.bpm,
        source: [i, j]
      });
    }
  }
  return result;
}

function scoreTracks(tracks) {
  const targetBpm = 120;
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    const track = tracks[i];

    result.push({
      trackId: track.trackId,
      artist: track.artist,
      title: track.title,
      votes: track.votes,
      bpm: track.bpm,
      source: track.source,
      score: (track.votes * 10) - Math.abs(track.bpm - targetBpm)
    });
  }
  return result;
}

function dedupeTracks(tracks) {
  const seenTrackIds = [];
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    if (!seenTrackIds.includes(tracks[i].trackId)) {
      seenTrackIds.push(tracks[i].trackId);
      result.push(tracks[i]);
    }
  }
  return result;
}

function enforceArtistQuota(tracks, maxPerArtist) {
  const artists = [];
  const artistCounts = [];
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    const artist = tracks[i].artist;
    const artistIndex = artists.indexOf(artist);

    if (artistIndex === -1) {
      artists.push(artist);
      artistCounts.push(1);
      result.push(tracks[i]);
    } else if (artistCounts[artistIndex] < maxPerArtist) {
      artistCounts[artistIndex]++;
      result.push(tracks[i]);
    }
  }

  return result;
}

function buildSchedule(tracks) {
  const schedule = [];

  for (let i = 0; i < tracks.length; i++) {
    schedule.push({
      slot: i + 1,
      trackId: tracks[i].trackId
    });
  }

  return schedule;
}

function remixPlaylist(playlists, maxPerArtist) {
  const flattened = flattenPlaylists(playlists);
  const scored = scoreTracks(flattened);
  const deduped = dedupeTracks(scored);
  const limited = enforceArtistQuota(deduped, maxPerArtist);
  return buildSchedule(limited);
}
```
