---
id: 698b42492a7f5f737895892e
title: Créer un convertisseur liste d'adjacence vers matrice
challengeType: 26
dashedName: build-an-adjacency-list-to-matrix-converter
---

# --description--

Dans ce labo, tu vas créer une fonction qui convertit une représentation en liste d'adjacence d'un graphe en une matrice d'adjacence. Une liste d'adjacence est un objet où chaque clé représente un nœud, et la valeur correspondante est un tableau de nœuds auxquels le nœud clé est connecté. Une matrice d'adjacence est un tableau 2D où l'entrée à la position `[i][j]` vaut `1` s'il y a une arête du nœud `i` vers le nœud `j`, et `0` sinon.

Par exemple, pour la liste d'adjacence suivante :

```js
{
  0: [1, 2],
  1: [2],
  2: [0, 3],
  3: [2]
}
```

La matrice d'adjacence correspondante serait :

```js
[
  [0, 1, 1, 0],
  [0, 0, 1, 0],
  [1, 0, 0, 1],
  [0, 0, 1, 0]
]
```

**Objectif :** Remplis les user stories ci-dessous et fais passer tous les tests pour terminer le labo.

**User Stories :**

1. Tu dois définir une fonction nommée `adjacencyListToMatrix` pour convertir une liste d'adjacence en matrice d'adjacence.
2. La fonction doit prendre un objet représentant la liste d'adjacence d'un graphe non pondéré (non orienté ou orienté) comme argument.
3. La fonction doit :
   - Convertir la liste d'adjacence en matrice d'adjacence.
   - Afficher chaque ligne de la matrice d'adjacence.
   - Renvoyer la matrice d'adjacence.

Par exemple, `adjacencyListToMatrix({0: [2], 1: [2, 3], 2: [0, 1, 3], 3: [1, 2]})` devrait afficher :

```md
[0, 0, 1, 0]
[0, 0, 1, 1]
[1, 1, 0, 1]
[0, 1, 1, 0]
```

et renvoyer `[[0, 0, 1, 0], [0, 0, 1, 1], [1, 1, 0, 1], [0, 1, 1, 0]]`.

# --hints--

Tu dois définir une fonction nommée `adjacencyListToMatrix`.

```js
assert.isFunction(adjacencyListToMatrix);
```

La fonction `adjacencyListToMatrix` doit avoir un paramètre.

```js
assert.lengthOf(adjacencyListToMatrix, 1);
```

La fonction doit déterminer correctement le nombre de nœuds à partir de la liste d'adjacence.

```js
const adjList = {0: [1], 1: [0], 2: []};
const result = adjacencyListToMatrix(adjList);
assert.lengthOf(result, 3);
result.forEach(arr => assert.lengthOf(arr, 3));
```

La fonction doit correctement fixer les valeurs de la matrice à `1` pour les arêtes existantes.

```js
const adjList = {0: [1], 1: [0]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1], [1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

La fonction doit afficher chaque ligne de la matrice.

```js
const spy = __helpers.spyOn(console, "log");
try {
  const adjList = {0: [1], 1: []}
  adjacencyListToMatrix(adjList);
  const args = spy.calls;
  assert.sameDeepOrderedMembers(args[0], [ [0, 1] ]);
  assert.sameDeepOrderedMembers(args[1], [ [0, 0] ]);
} catch (e) {
  assert.fail(e);
} finally {
  spy.restore();
}
```

La fonction doit renvoyer la matrice d'adjacence.

```js
const adjList = {0: [1], 1: []};
const result = adjacencyListToMatrix(adjList);
assert.sameDeepOrderedMembers(result, [[0, 1], [0, 0]]);
```

Quand on lui donne la liste d'adjacence `{0: [1, 2], 1: [2], 2: [0, 3], 3: [2]}`, la fonction doit renvoyer `[[0, 1, 1, 0], [0, 0, 1, 0], [1, 0, 0, 1], [0, 0, 1, 0]]`.

```js
const adjList = {0: [1, 2], 1: [2], 2: [0, 3], 3: [2]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1, 1, 0], [0, 0, 1, 0], [1, 0, 0, 1], [0, 0, 1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

Quand on lui donne la liste d'adjacence `{0: [1], 1: [0]}`, la fonction doit renvoyer `[[0, 1], [1, 0]]`.

```js
const adjList = {0: [1], 1: [0]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1], [1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

Quand on lui donne la liste d'adjacence `{0: [], 1: [], 2: []}`, la fonction doit renvoyer `[[0, 0, 0], [0, 0, 0], [0, 0, 0]]`.

```js
const adjList = {0: [], 1: [], 2: []};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 0, 0], [0, 0, 0], [0, 0, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function adjacencyListToMatrix(adjList) {
  const n = Object.keys(adjList).length;

  const adjMatrix = Array.from({ length: n }, () => Array(n).fill(0));

  for (const [srcStr, neighbors] of Object.entries(adjList)) {
    const srcNode = Number(srcStr);
    for (const destNode of neighbors) {
      adjMatrix[srcNode][destNode] = 1;
    }
  }

  for (const row of adjMatrix) {
    console.log(row);
  }

  return adjMatrix;
}

const adjList = {
  0: [1, 2],
  1: [2],
  2: [0, 3],
  3: [2]
};

const matrix = adjacencyListToMatrix(adjList);
```
