---
id: 587d8259367417b2b2512c85
title: Implémenter l'algorithme de tri par sélection
challengeType: 26
dashedName: implement-the-selection-sort-algorithm
---

# --description--

Dans ce lab, tu vas implémenter l'algorithme de tri par sélection. Le tri par sélection fonctionne en sélectionnant la valeur minimale d'une liste et en l'échangeant avec la première valeur de la liste. Il recommence ensuite à la deuxième position, sélectionne la plus petite valeur dans la liste restante, et l'échange avec le deuxième élément. Il continue à parcourir la liste et à échanger des éléments jusqu'à atteindre la fin de la liste. La liste est alors triée. Le tri par sélection a une complexité temporelle quadratique dans tous les cas.

**Objectif :** Remplir les user stories ci-dessous et faire passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu dois avoir une fonction `selectionSort`.
1. La fonction `selectionSort` doit prendre un tableau, et renvoyer un tableau avec les mêmes éléments mais ordonnés du plus petit au plus grand.

# --hints--

`selectionSort` doit être une fonction.

```js
assert.isFunction(selectionSort);
```

`selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` doit renvoyer un tableau inchangé sauf pour l'ordre.

```js
assert.sameMembers(
  selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`selectionSort` doit renvoyer un tableau trié (du plus petit au plus grand).

```js
assert.sameOrderedMembers(
  selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
)
```

`selectionSort` ne doit pas utiliser la méthode intégrée `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    selectionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function selectionSort(array) {
  for (let i = 0; i < array.length-1; i++) {
    let minimumIndex = i;
    for (let j = i+1; j < array.length; j++){
      if (array[j] < array[minimumIndex]) {
        minimumIndex = j;
      }
    }
    let value = array[minimumIndex];
    array[minimumIndex] = array[i];
    array[i] = value;
  }
    return array;
}
```
