---
id: 6899b593d232e07617f74055
title: Étape 4
challengeType: 0
dashedName: step-4
---

# --description--

À l'intérieur de l'élément `button`, crée un élément `span` qui contient `😊`

Donne à l'élément `span` un attribut `role` avec la valeur `img` et un attribut `aria-hidden` avec la valeur `true`.

Ensuite, crée un second élément `span` avec une `class` `count` contenant le texte `0/10`.

# --hints--

À l'intérieur de l'élément `button`, il devrait y avoir deux éléments `span`.

```js
assert.lengthOf(document.querySelectorAll("button > span"), 2);
```

Le premier élément `span` devrait avoir le contenu texte `😊`.

```js
assert.equal(document.querySelectorAll("button > span")[0]?.textContent.trim(), "😊");
```

Le premier élément `span` devrait avoir un attribut `role` avec la valeur `img`.

```js
assert.equal(document.querySelectorAll("button > span")[0]?.getAttribute("role"), "img");
```

Le premier élément `span` devrait avoir un attribut `aria-hidden` avec la valeur `true`.

```js
assert.equal(document.querySelectorAll("button > span")[0]?.getAttribute("aria-hidden"), "true");
```

Le second élément `span` devrait avoir un attribut `class` avec la valeur `count`.

```js
assert.equal(document.querySelectorAll("button > span")[1]?.className, "count");
```

Le second élément `span` devrait avoir le contenu texte `0/10`.

```js
assert.equal(document.querySelectorAll("button > span")[1]?.textContent.trim(), "0/10");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Emoji Reactor</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
<body>
  <main>

    <h1 class="title">How are you feeling today?</h1>

    <p class="description">
      Click on the buttons below to rate your emotions.
    </p>
    <div class="btn-container">
--fcc-editable-region--
      <button id="happy-btn" class="emoji-btn" aria-label="Happy face emoji">

      </button>
--fcc-editable-region--
    </div>
  </main>
</body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --light-grey: #efefef;
  --white: #ffffff;
  --very-dark-blue: #0a0a23;
  --light-purple: #a78bfa;
  --very-light-purple: #c4b5fd;
  --purple: #8b5cf6;
}

body {
  background-color: var(--very-dark-blue);
  color: var(--light-grey);
  font-family: sans-serif;
}

main {
  text-align: center;
  padding: 10px;
}

.btn-container,
button {
  display: flex;
  justify-content: space-evenly;
  align-items: center;
}

.btn-container {
  flex-direction: column;
}

.emoji-btn {
  width: 70%;
  cursor: pointer;
  color: var(--white);
  background-color: var(--light-purple);
  background-image: linear-gradient(
    to bottom,
    var(--very-light-purple),
    var(--light-purple)
  );
  border: 3px solid var(--purple);
  border-radius: 8px;
  padding: 10px;
  font-size: 1.5rem;
  margin: 10px 0;
  transition: background-color 0.3s ease, transform 0.2s ease;
}

@media (min-width: 768px) {
  .emoji-btn {
    width: 30%;
  }
}

.emoji-btn:hover {
  background-color: var(--purple);
  background-image: none;
}

.title {
  margin-top: 15px;
  font-size: 2rem;
}

.description {
  font-size: 1.4rem;
  margin: 20px 0;
}
```
