---
id: ab6137d4e35944e21037b769
title: Crée un convertisseur title case
challengeType: 26
dashedName: build-a-title-case-converter
---

# --description--

Dans ce lab, tu vas créer une fonction qui convertit une chaîne en title case. Le title case signifie que la première lettre de chaque mot est en majuscule et le reste du mot en minuscules.

`"Web Development Is Awesome"` est un exemple de chaîne au format title case.

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois avoir une fonction `titleCase` qui prend une chaîne en argument.
1. La fonction `titleCase` doit renvoyer une chaîne dont la première lettre de chaque mot est en majuscule et le reste du mot en minuscules.
1. `titleCase("I like to code")` doit renvoyer `"I Like To Code"`.
1. `titleCase("javaScript is fun")` doit renvoyer `"Javascript Is Fun"`.

# --hints--

Tu dois créer une fonction nommée `titleCase`.

```js
assert.isFunction(titleCase);
```

`titleCase` doit prendre un seul argument.

```js
assert.lengthOf(titleCase, 1);
```


`titleCase("I'm a little tea pot")` doit renvoyer une chaîne.

```js
assert.isString(titleCase("I'm a little tea pot"));
```

`titleCase("I'm a little tea pot")` doit renvoyer la chaîne `I'm A Little Tea Pot`.

```js
assert.strictEqual(titleCase("I'm a little tea pot"), "I'm A Little Tea Pot");
```

`titleCase("sHoRt AnD sToUt")` doit renvoyer la chaîne `Short And Stout`.

```js
assert.strictEqual(titleCase('sHoRt AnD sToUt'), 'Short And Stout');
```

`titleCase("HERE IS MY HANDLE HERE IS MY SPOUT")` doit renvoyer la chaîne `Here Is My Handle Here Is My Spout`.

```js
assert.strictEqual(
  titleCase('HERE IS MY HANDLE HERE IS MY SPOUT'),
  'Here Is My Handle Here Is My Spout'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function titleCase(str) {
  return str
    .split(' ')
    .map(word => word.charAt(0).toUpperCase() + word.substring(1).toLowerCase())
    .join(' ');
}
```
