---
id: 699f3f636aa840b7201aa37c
title: Étape 5
challengeType: 1
dashedName: step-5
---

# --description--

Si tu ouvres la console, tu devrais voir le message d’erreur TypeScript suivant :

```md
Argument of type 'string' is not assignable to parameter of type 'number'
```

Maintenant que le paramètre est explicitement typé comme un nombre, TypeScript peut détecter l’erreur qui consiste à passer un argument chaîne à la fonction `square`.

Change l’argument passé à la fonction `square` pour le nombre `5` afin de corriger l’erreur et de voir la bonne sortie dans la console.

# --hints--

Ton appel de fonction `square` doit avoir l’argument `5` au lieu de `"something"`.

```js
assert.strictEqual(result, 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square("something");
--fcc-editable-region--

console.log(result);
```
