---
id: 699f38937b053e4679878ff1
title: Étape 3
challengeType: 1
dashedName: step-3
---

# --description--

Quand tu travailles avec du JavaScript classique, rien ne t’empêche d’appeler la fonction `square` avec une chaîne ou un tableau, ce qui mènerait à des résultats inattendus.

Pour illustrer cela, change ton appel de fonction pour utiliser la chaîne `"something"` au lieu du nombre `5`.

Dans la prochaine leçon, tu apprendras à utiliser TypeScript pour ajouter la sûreté des types à ton code et prévenir ce genre d’erreurs.

# --hints--

Ton appel de fonction `square` doit avoir l’argument `"something"`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.variables.result.value.matches("square('something')"));
```

# --seed--

## --seed-contents--

```ts
function square(num) {
  return num * num;
}

--fcc-editable-region--
const result = square(5);
--fcc-editable-region--

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
