---
id: 699f41b8da20f816b5c7dece
title: Étape 7
challengeType: 1
dashedName: step-7
---

# --description--

Si tu ouvres la console, tu devrais voir le message d’erreur suivant :

```md
Expected 1 arguments, but got 2.
```

Pour résoudre l’erreur, mets à jour ton appel de fonction `square` pour qu’il ne prenne qu’un seul argument : `5`.

# --hints--

Ton appel de fonction `square` doit n’avoir qu’un seul argument : `5`.

```js
assert.strictEqual(result, 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square(5, 10);
--fcc-editable-region--

console.log(result);
```
