---
id: 699f52220de88c8091032076
title: Étape 18
challengeType: 1
dashedName: step-18
---

# --description--

Si tu ouvres la console, le résultat est `NaN` parce que le second argument pour `exponent` n’a jamais été fourni. Quelques mises à jour sont nécessaires pour cette fonction `raiseTo` afin de corriger cela.

Le premier changement est de renvoyer `base ** 2`.

Dans la prochaine étape, tu verras comment gérer le paramètre optionnel `exponent`.

# --hints--

Ta fonction `raiseTo` doit renvoyer un nombre.

```js
assert.isNumber(raiseTo(4));
```

Ta fonction `raiseTo` doit renvoyer la valeur de `base` élevée à la puissance de `2`.

```js
assert.strictEqual(raiseTo(4), 16);
assert.strictEqual(raiseTo(5), 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent?: number): number {
  return base ** exponent;
}
--fcc-editable-region--

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
