---
id: 699f434e294a9e11cf4f65f0
title: Étape 8
challengeType: 1
dashedName: step-8
---

# --description--

Il est maintenant temps d’ajouter un autre outil mathématique à ton projet de boîte à outils.

Commence par créer une fonction appelée `getAverage` qui prend un paramètre. Ta fonction doit renvoyer la moyenne d’un tableau de nombres.

Dans les prochaines étapes, tu ajouteras les annotations de type pour cette fonction.

# --hints--

Tu dois avoir une fonction appelée `getAverage`.

```js
assert.isFunction(getAverage);
```

Ta fonction doit prendre un paramètre.

```js
assert.strictEqual(getAverage.length, 1);
```

Ta fonction doit renvoyer un nombre.

```js
assert.isNumber(getAverage([1, 2, 3]));
```

Ta fonction doit renvoyer le résultat correct pour n’importe quel tableau de nombres.

```js
assert.strictEqual(getAverage([1, 2, 3, 4, 5]), 3);
assert.strictEqual(getAverage([10, 20, 30]), 20);
assert.strictEqual(getAverage([5, 10, 15, 20]), 12.5);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
