---
id: 699f4b903bf5427a599f791d
title: Étape 14
challengeType: 1
dashedName: step-14
---

# --description--

Dans les dernières étapes, tu vas construire une fonction de plus pour continuer à t’entraîner aux annotations de type avec les fonctions.

Commence par créer une fonction appelée `raiseTo` qui prend deux paramètres appelés `base` et `exponent`. Les deux paramètres doivent être de type `number`. La fonction doit renvoyer le résultat de l’élévation de `base` à la puissance de `exponent`. Le type de retour de la fonction doit aussi être annoté comme `number`.

# --hints--

Tu dois avoir une fonction `raiseTo`.

```js
assert.isFunction(raiseTo);
```

Ta fonction `raiseTo` doit avoir deux paramètres, `base` et `exponent`.

```js
assert.match(
  code,
   /function\s+raiseTo\s*\(\s*base\s*(?::\s*number)?\s*,\s*exponent\s*(?::\s*number)?\s*\)\s*|(?:const|let)\s+raiseTo\s*=\s*\(\s*base\s*(?::\s*number)?\s*,\s*exponent\s*(?::\s*number)?\s*\)\s*=>/
);
```

Tes paramètres `base` et `exponent` doivent tous les deux être de type `number`.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.raiseTo;
assert.isTrue(parameters[0].matches("base: number"));
assert.isTrue(parameters[1].matches("exponent: number"));
```

Ta fonction `raiseTo` doit renvoyer un nombre.

```js
assert.isNumber(raiseTo(2, 3));
```

Ta fonction `raiseTo` doit renvoyer le résultat de l’élévation de `base` à la puissance de `exponent`.

```js
assert.strictEqual(raiseTo(2, 3), 8);
```

Ta fonction `raiseTo` doit avoir une annotation de type de retour `number`.

```js
const explorer = await __helpers.Explorer(code);
const { raiseTo } = explorer.allFunctions;
assert.isTrue(raiseTo.hasReturnAnnotation("number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
