---
id: 699f52fac083214b6bf51de6
title: Étape 19
challengeType: 1
dashedName: step-19
---

# --description--

Dans cette dernière étape, tu devras ajouter une vérification conditionnelle pour le paramètre `exponent` dans la fonction `raiseTo`.

Si l’argument `exponent` est fourni, la fonction doit renvoyer `base` élevée à la puissance de `exponent`. Si l’argument `exponent` n’est pas fourni, elle doit renvoyer `base ** 2`.

Avec cette dernière étape, ta boîte à outils mathématique type-safe est terminée !

# --hints--

Ta fonction `raiseTo` doit renvoyer `base ** exponent` quand l’argument `exponent` est fourni.

```js
assert.strictEqual(raiseTo(3, 3), 27);
```

Ta fonction `raiseTo` doit renvoyer `base ** 2` quand l’argument `exponent` n’est pas fourni.

```js
assert.strictEqual(raiseTo(3), 9);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent?: number): number {

  return base ** 2;
}
--fcc-editable-region--

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent?: number): number {
  if (exponent) {
    return base ** exponent;
  }

  return base ** 2;
}

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
