---
id: 699f5063d2393c674b1f45dc
title: Étape 16
challengeType: 1
dashedName: step-16
---

# --description--

Comme tu te souviens d’une étape précédente, si tu appelles une fonction avec le mauvais nombre d’arguments, TypeScript lèvera une erreur.

Il existe un moyen de marquer des paramètres comme optionnels en ajoutant un `?` après le nom du paramètre. Cela te permet d’appeler la fonction avec moins d’arguments que de paramètres.

Voici un exemple :

```ts
function greet(name: string, greeting?: string) {
  if (greeting) {
    return `${greeting}, ${name}!`;
  }
  return `Hello, ${name}!`;
}
```

Dans cet exemple, si le paramètre `greeting` n’est pas fourni lors de l’appel de la fonction `greet`, il sera `undefined`, et la fonction renverra une salutation par défaut.

Mets à jour ta fonction `raiseTo` pour rendre le paramètre `exponent` optionnel.

# --hints--

Ta fonction `raiseTo` doit avoir le paramètre `exponent` marqué comme optionnel.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.raiseTo;
assert.isTrue(parameters[1].matches("exponent?: number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent: number): number {
  return base ** exponent;
}
--fcc-editable-region--

const raisedResult = raiseTo(2, 3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
