---
id: 699f513e01eb1768cc5df726
title: Étape 17
challengeType: 1
dashedName: step-17
---

# --description--

Maintenant que le paramètre `exponent` est marqué comme optionnel, tu peux appeler la fonction avec un seul argument.

Mets à jour ton appel de fonction `raiseTo` pour ne passer que le nombre `3`.

# --hints--

Ta fonction `raiseTo` doit être appelée avec un seul argument : `3`.

```js
assert.match(
  code,
  /raiseTo\s*\(\s*3\s*\)/
);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent?: number): number {
  return base ** exponent;
}

--fcc-editable-region--
const raisedResult = raiseTo(2, 3);
--fcc-editable-region--
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
