---
id: 699f49669ea51a31d2ba89d8
title: Étape 12
challengeType: 1
dashedName: step-12
---

# --description--

Comme tu te souviens des leçons précédentes, tu peux typer explicitement le type de retour d’une fonction comme ceci :

```ts
function circleArea(radius: number): number {
  return Math.PI * radius * radius;
}
```

Dans cet exemple, le type de retour de la fonction `circleArea` est `number`. Si tu essaies de renvoyer un type différent, TypeScript affichera une erreur.

Mets à jour ta fonction `getAverage` pour typer explicitement le type de retour comme `number`.

# --hints--

Ta fonction `getAverage` doit avoir un type de retour `number`.

```js
const explorer = await __helpers.Explorer(code);
const { getAverage } = explorer.allFunctions;
assert.isTrue(getAverage.hasReturnAnnotation("number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}
--fcc-editable-region--

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
