---
id: 699f3738921fb79a651ded56
title: Étape 2
challengeType: 1
dashedName: step-2
---

# --description--

Il est maintenant temps d’appeler ta fonction.

Crée une variable appelée `result` et assigne-lui l’appel de fonction `square` avec l’argument `5`.

En dessous, affiche la valeur de `result` dans la console.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Tu dois avoir une variable appelée `result`.

```js
assert.isDefined(result);
```

Ta variable `result` doit être un nombre.

```js
assert.isNumber(result);
```

Ta variable `result` doit être le carré de `5`.

```js
assert.strictEqual(result, 25);
```

Tu dois afficher la variable `result` dans la console.

```js
assert.equal(getLogs()[0], result);
```

# --seed--

## --seed-contents--

```ts
function square(num) {
  return num * num;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
