---
id: 699f412114c87221f16a99e3
title: Étape 6
challengeType: 1
dashedName: step-6
---

# --description--

Un autre avantage d’utiliser TypeScript est qu’il détectera les erreurs si tu passes trop ou trop peu d’arguments à un appel de fonction.

Pour illustrer cela, mets à jour ton appel de fonction `square` pour qu’il prenne deux arguments : `5` et `10`.

# --hints--

Ton appel de fonction `square` doit avoir deux arguments : `5` et `10`.

```js
const explorer = await __helpers.Explorer(code);
const { result } = explorer.variables;
assert.isTrue(result.value.matches("square(5, 10)"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square(5);
--fcc-editable-region--

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5, 10);

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
