---
id: 699f4861dd5b5f164e1bda4a
title: Étape 11
challengeType: 1
dashedName: step-11
---

# --description--

Un autre avantage d’utiliser TypeScript est de s’assurer que les types de retour sont imposés et correspondent à ce que tu attends. Parfois TypeScript peut inférer le type de retour d’une fonction, mais tu peux aussi le spécifier explicitement.

Pour illustrer cela, change le type de retour de la fonction `getAverage` pour la chaîne `"something"`.

# --hints--

Ta fonction `getAverage` doit renvoyer une chaîne.

```js
assert.isString(getAverage([2, 14, 26, 8]));
```

Ta fonction `getAverage` doit renvoyer la chaîne `"something"`.

```js
assert.strictEqual(getAverage([2, 14, 26, 8]), "something");
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  --fcc-editable-region--
  return sum / numbers.length;
  --fcc-editable-region--
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
