---
id: 699f1d72e906e4616155290d
title: Étape 1
challengeType: 1
dashedName: step-1
---

# --description--

Dans cet atelier, tu vas continuer à apprendre les annotations de type et leur fonctionnement avec les fonctions en construisant une boîte à outils mathématique type-safe.

Commence par créer une fonction appelée `square` qui accepte un paramètre et renvoie le carré d’un nombre.

# --hints--

Tu dois avoir une fonction appelée `square`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.square);
```

Ta fonction `square` doit avoir un paramètre.

```js
const explorer = await __helpers.Explorer(code);
assert.lengthOf(explorer.allFunctions.square.parameters, 1);
```

Ta fonction `square` doit renvoyer le carré d’un nombre.

```js
assert.strictEqual(square(2), 4);
assert.strictEqual(square(3), 9);
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
