---
id: 699f4ec939f10da62fb56c3e
title: Étape 15
challengeType: 1
dashedName: step-15
---

# --description--

Il est maintenant temps de tester ta fonction.

Crée une variable appelée `raisedResult` et assigne-lui le résultat de l’appel de `raiseTo` avec les arguments `2` et `3`.

Puis affiche `raisedResult` dans la console.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Tu dois avoir une variable appelée `raisedResult`.

```js
assert.isDefined(raisedResult);
```

Ta variable `raisedResult` doit être un nombre.

```js
assert.isNumber(raisedResult);
```

Tu dois afficher `raisedResult` dans la console.

```js
assert.equal(getLogs()[2], raisedResult);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent: number): number {
  return base ** exponent;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
