---
id: 699f3cb8d057f29954205d7d
title: Étape 4
challengeType: 1
dashedName: step-4
---

# --description--

Comme tu te souviens des leçons précédentes, tu peux ajouter des annotations de type aux paramètres de fonction comme ceci :

```ts
function add(num1: number, num2: number) {
  return num1 + num2;
}
```

Dans cet exemple, TypeScript affichera une erreur si tu essaies d’appeler la fonction `add` avec des arguments qui ne sont pas des nombres.

Mets à jour ta fonction `square` pour y inclure une annotation de type pour le paramètre `num`. Le type doit être `number`.

# --hints--

Tu dois ajouter une annotation de type au paramètre `num` dans la fonction `square`. Le type doit être un `number`.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.square;
assert.isTrue(parameters[0].matches("num: number"));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
function square(num) {
  return num * num;
}

--fcc-editable-region--
const result = square("something");

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square("something");

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
