---
id: 699f458e47fb88c756fdba81
title: Étape 9
challengeType: 1
dashedName: step-9
---

# --description--

Pour l’instant, le paramètre `numbers` a un type `any` implicite. Pour corriger cela, tu peux typer explicitement le paramètre comme un tableau de nombres.

Rappelle-toi des leçons précédentes : tu peux typer un tableau de nombres comme ceci :

```ts
const numbers: number[] = [1, 2, 3];
```

Mets à jour ta fonction `getAverage` pour typer explicitement le paramètre `numbers` comme un tableau de nombres.

# --hints--

Ta fonction `getAverage` doit avoir un paramètre nommé `numbers` qui est explicitement typé comme un tableau de nombres.

```js
const explorer = await __helpers.Explorer(code);
const parameters = explorer.allFunctions.getAverage.parameters;
assert.isTrue(parameters[0].matches("numbers: number[]"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}
--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
