---
id: 699f4639ae0016caaf81b7c6
title: Étape 10
challengeType: 1
dashedName: step-10
---

# --description--

Il est maintenant temps de tester ta fonction.

Commence par créer une variable appelée `avgResult` et assigne-lui l’appel de fonction `getAverage`. L’argument de cet appel de fonction doit être `[2, 14, 26, 8]`.

Puis affiche la valeur de `avgResult` dans la console.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Tu dois avoir une variable appelée `avgResult`.

```js
assert.isDefined(avgResult);
```

Tu dois assigner `getAverage([2, 14, 26, 8])` à ta variable `avgResult`.

```js
assert.strictEqual(avgResult, getAverage([2, 14, 26, 8]));
```

Ta variable `avgResult` doit être un nombre.

```js
assert.isNumber(avgResult);
```

Tu dois afficher `avgResult` dans la console.

```js
assert.equal(getLogs()[1], avgResult);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
