---
id: 619665c9abd72906f3ad30f9
title: Étape 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Tu vas construire un Flappy Penguin heureux, tout en continuant à pratiquer les transformations et animations CSS.

Commence par lier ta feuille de style à la page.

# --hints--

Ton code devrait avoir un élément `link`.

```js
assert.exists(document.querySelector('link'));
```

Ton élément `link` devrait être dans ton élément `head`.

```js
assert.match(code, /<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i);
```

Ton élément `link` devrait avoir un attribut `rel` avec la valeur `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

Ton élément `link` devrait avoir un attribut `href` avec la valeur `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
