---
id: 61969aa6acef5b12200f672e
title: Étape 13
challengeType: 0
dashedName: step-13
---

# --description--

Pour créer un peu de décor en arrière-plan, tu vas ajouter deux montagnes.

Au-dessus de l'élément `.penguin`, ajoute un `div` avec une `class` de `left-mountain`.

# --hints--

Tu devrais ajouter un nouveau `div` dans `body`. Il fallait voir --fcc-expected-- éléments `div`, mais --fcc-actual-- ont été trouvés.

```js
assert.lengthOf(document.querySelectorAll('body > div'), 3);
```

Tu devrais donner au `div` une `class` de `left-mountain`.

```js
assert.include(document.querySelector('body > div:not(.ground, .penguin)')?.className, 'left-mountain');
```

Tu devrais placer `.left-mountain` avant `.penguin`.

```js
assert.strictEqual(document.querySelector('.penguin')?.previousElementSibling, document.querySelector('.left-mountain'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

--fcc-editable-region--
  <body>

    <div class="penguin"></div>
    <div class="ground"></div>
  </body>
--fcc-editable-region--
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

.penguin {
  width: 300px;
  height: 300px;
  margin: auto;
  margin-top: 75px;
}

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
}
```
