---
id: 619691693bc14b0e528f5a20
title: Étape 7
challengeType: 0
dashedName: step-7
---

# --description--

Cible l'élément `.ground` et définis son `width` pour qu'il occupe toute la largeur de la fenêtre d'affichage. Ensuite, définis le `height` sur `400px`.

# --hints--

Tu devrais utiliser le sélecteur `.ground`.

```js
assert.match(code, /\.ground\s*\{/);
```

Tu devrais donner à l'élément `.ground` un `width` de `100vw`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.ground')?.width, '100vw');
```

Tu devrais donner à l'élément `.ground` un `height` de `400px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.ground')?.height, '400px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

--fcc-editable-region--

--fcc-editable-region--
```
