---
id: 61967e74a8e3690ab6292daa
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

Cible l'élément `body` pour définir le `background` sur un dégradé linéaire incliné de 45 degrés dans le sens horaire, qui commence à `rgb(118, 201, 255)` et se terminé à `rgb(247, 255, 222)`.

# --hints--

Tu devrais utiliser le sélecteur de l'élément `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

Tu devrais utiliser la propriété `background` dans le sélecteur `body`.

```js
assert.isTrue(new __helpers.CSSHelp(document).isPropertyUsed('background'));
```

Tu devrais définir `background` sur `linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222))`.

```js
assert.include(['linear-gradient(45deg,rgb(118,201,255),rgb(247,255,222))', 'rgba(0,0,0,0)linear-gradient(45deg,rgb(118,201,255),rgb(247,255,222))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
