---
id: 699b5b7cb68b6d35afca61b7
title: Étape 4
challengeType: 1
dashedName: step-4
---

# --description--

Même s’il est possible de continuer à ajouter des propriétés comme ça à l’objet `profile`, ce serait bien de restreindre les types de propriétés autorisées. Il serait préférable de définir la forme de l’objet `profile` à l’avance.

Dans les leçons précédentes, tu as appris à travailler avec des types d’objet comme ceci :

```ts
const composer: {
  name: string;
  instrument: string;
  isActive: boolean;
} = {
  name: "Ludwig van Beethoven",
  instrument: "Piano",
  isActive: false
};
```

Dans cet exemple, `composer` est un objet avec un type d’objet en ligne. Cet objet a trois propriétés requises : `name`, `instrument` et `isActive`. Les deux premières propriétés doivent être de type `string` et `isActive` doit être de type `boolean`.

Mets à jour ton objet `profile` existant pour y inclure un type d’objet en ligne. `username` doit être typé en `string`, `age` en `number` et `isLoggedIn` en `boolean`.

# --hints--

Ton objet `profile` doit avoir un type d’objet en ligne avec une propriété `username` définie sur `string`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "username", type: "string" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

Ton objet `profile` doit avoir un type d’objet en ligne avec une propriété `age` définie sur `number`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "age", type: "number" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

Ton objet `profile` doit avoir un type d’objet en ligne avec une propriété `isLoggedIn` définie sur `boolean`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "isLoggedIn", type: "boolean" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
const profile = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  mood: null
}
--fcc-editable-region--

console.log(profile);
```
