---
id: 699b68f94a0d154c2a7a5ce0
title: Étape 9
challengeType: 1
dashedName: step-9
---

# --description--

Maintenant que l’objet `profile` est complet, il est temps de construire un nouvel objet qui représente les rôles utilisateur.

Crée une nouvelle variable appelée `userRoles` et assigne-lui un objet. Ton objet doit avoir les propriétés et valeurs suivantes :

```md
admin: "full-access",
editor: "limited-access",
viewer: "read-only"
```

# --hints--

Tu dois avoir une variable appelée `userRoles`.

```js
assert.isDefined(userRoles);
```

Ta variable `userRoles` doit être un objet.

```js
assert.isObject(userRoles);
```

Ton objet `userRoles` doit avoir une propriété `admin`.

```js
assert.property(userRoles, "admin");
```

Ta propriété `admin` doit avoir la valeur `"full-access"`.

```js
assert.propertyVal(userRoles, "admin", "full-access");
```

Ton objet `userRoles` doit avoir une propriété `editor`.

```js
assert.property(userRoles, "editor");
```

Ta propriété `editor` doit avoir la valeur `"limited-access"`.

```js
assert.propertyVal(userRoles, "editor", "limited-access");
```

Ton objet `userRoles` doit avoir une propriété `viewer`.

```js
assert.property(userRoles, "viewer");
```

Ta propriété `viewer` doit avoir la valeur `"read-only"`.

```js
assert.propertyVal(userRoles, "viewer", "read-only");
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--

--fcc-editable-region--
```
