---
id: 699b64096301fd75e3efeb1b
title: Étape 7
challengeType: 1
dashedName: step-7
---

# --description--

Si tu ouvres à nouveau la console, tu verras un message d’erreur différent :

```md
Property 'bio' is missing in type '{ username: string; age: number; isLoggedIn: false; }' but required in type '{ username: string; age: number; isLoggedIn: boolean; bio: string; }'.
```

Comme tu te souviens des leçons précédentes, tu peux marquer des propriétés comme optionnelles en utilisant le symbole `?` à côté du nom de la propriété, comme ceci :

```ts
const musician: {
  name: string;
  instrument: string;
  isActive: boolean;
  numberOfAlbums?: number;
} = {
  name: "Ludwig van Beethoven",
  instrument: "Piano",
  isActive: false
};
```

Dans cet exemple, `numberOfAlbums` est marqué comme optionnel, donc TypeScript ne lèvera pas d’erreur dans ce cas.

Mets à jour ta propriété `bio` pour qu’elle soit marquée comme optionnelle.

# --hints--

Ta propriété `bio` doit être marquée comme optionnelle.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "bio", type: "string", isOptional: true };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
--fcc-editable-region--
  bio: string
--fcc-editable-region--
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};

console.log(profile);
```
