---
id: 699b65caf17a8e8d340af27d
title: Étape 8
challengeType: 1
dashedName: step-8
---

# --description--

Il est maintenant temps d’ajouter la dernière propriété à l’objet `profile`.

Dans les leçons précédentes, tu as appris à travailler avec les types de tableaux. Tu as vu différents types de tableaux comme `string[]` ou `number[]` :

```ts
const scores: number[] = [95, 87, 100];
```

Dans cet exemple, le tableau `scores` n’accepte que des nombres. Si tu essaies d’ajouter un autre type comme une `string` ou `null`, TypeScript lèvera une erreur.

Mets à jour ton objet `profile` pour y ajouter une propriété `programmingLanguages`, typée comme `string[]`. Dans le littéral d’objet, définis cette propriété sur `["JavaScript", "Python", "C++"]`.

# --hints--

Tu dois avoir une propriété `programmingLanguages` à l’intérieur de ton objet `profile`.

```js
assert.property(profile, "programmingLanguages");
```

Ta propriété `programmingLanguages` doit avoir la valeur `["JavaScript", "Python", "C++"]`.

```js
assert.deepEqual(profile.programmingLanguages, ["JavaScript", "Python", "C++"]);
```

Ta propriété `programmingLanguages` doit être typée comme `string[]`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "programmingLanguages", type: "string[]" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};
--fcc-editable-region--

console.log(profile);
```
