---
id: 699b4e284291adcfcb90df47
title: Étape 1
challengeType: 1
dashedName: step-1
---

# --description--

Dans cet atelier, tu vas t’entraîner à travailler avec les types primitifs ainsi qu’avec les types objet et tableau en construisant un profil utilisateur type-safe.

Commence par créer une variable appelée `profile` et assigne-lui un objet. Ton objet doit avoir ces trois propriétés et valeurs :

```md
username: "codeLearner",
age: 25,
isLoggedIn: false,
```

# --hints--

Tu dois avoir une variable appelée `profile`.

```js
assert.exists(profile);
```

Ta variable `profile` doit être un objet.

```js
assert.isObject(profile);
```

Tu dois avoir une propriété `username` dans ton objet `profile`.

```js
assert.property(profile, "username");
```

Ta propriété `username` doit avoir la valeur `"codeLearner"`.

```js
assert.propertyVal(profile, "username", "codeLearner");
```

Ton objet `profile` doit avoir une propriété `age`.

```js
assert.property(profile, "age");
```

Ta propriété `age` doit avoir la valeur `25`.

```js
assert.propertyVal(profile, "age", 25);
```

Ton objet `profile` doit avoir une propriété `isLoggedIn`.

```js
assert.property(profile, "isLoggedIn");
```

Ta propriété `isLoggedIn` doit avoir la valeur `false`.

```js
assert.propertyVal(profile, "isLoggedIn", false);
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--

--fcc-editable-region--
```
