---
id: 699b62c99ec8c63470313365
title: Étape 6
challengeType: 1
dashedName: step-6
---

# --description--

Il est maintenant temps d’ajouter une autre propriété à ton objet `profile`.

Commence par mettre à jour le type d’objet en ligne pour y inclure une propriété `bio` avec un type défini sur `string`.

# --hints--

Tu dois mettre à jour le type d’objet en ligne pour y inclure une propriété `bio` définie sur `string`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "bio", type: "string" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
--fcc-editable-region--
  username: string;
  age: number;
  isLoggedIn: boolean;

--fcc-editable-region--
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};

console.log(profile);
```
