---
id: 699b6dd1a2961bd38491ef2e
title: Étape 11
challengeType: 1
dashedName: step-11
---

# --description--

L’objet `profile` avait des clés et des types stricts. Si tu essayais d’ajouter une propriété non définie, ou d’assigner un type incompatible, TypeScript lèverait une erreur.

Ce serait bien si l’objet `userRoles` n’était pas aussi strict et avait des clés plus flexibles.

Comme tu te souviens des leçons précédentes, tu peux obtenir ce résultat en utilisant le type `Record` :

```ts
const studentGrades: Record<string, number> = {
  Math: 95,
  English: 88,
  Science: 92
};
```

Dans cet exemple `studentGrades`, TypeScript n’a pas besoin de connaître toutes les matières à l’avance. Cela te permet d’ajouter plus tard des matières comme `Art` ou `Music`. TypeScript se soucie seulement que les valeurs de chaque clé soient un nombre.

Mets à jour ton objet `userRoles` en ajoutant un type `Record` en ligne. Définis-le sur `<string, string>` pour que chaque clé puisse être n’importe quelle chaîne, et que chaque valeur doive être une chaîne.

# --hints--

Ton objet `userRoles` doit avoir un type `Record` en ligne défini sur `Record<string, string>`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.annotation.matches("Record<string, string>"));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--
const userRoles = {
--fcc-editable-region--
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only"
};

console.log(userRoles);
```
