---
id: 699b5370a43e299b4673bcaa
title: Étape 2
challengeType: 1
dashedName: step-2
---

# --description--

Il serait utile de voir les propriétés de l’objet `profile` dans la console pendant que tu le construis.

Ajoute un `console.log` qui affiche `profile` dans la console.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Tu dois afficher `profile` dans la console.

```js
assert.equal(getLogs()[0], profile);
```

# --seed--

## --seed-contents--

```ts
const profile = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
}

--fcc-editable-region--

--fcc-editable-region--
```
