---
id: 699b71d98c820ac7ab9eea2a
title: Étape 13
challengeType: 1
dashedName: step-13
---

# --description--

Si tu ouvres à nouveau la console, tu verras un nouveau message d’erreur :

```md
Type 'number' is not assignable to type 'string'.
```

Cela se produit parce que tu tentes d’ajouter une propriété `guest` avec une valeur numérique.

Pour supprimer l’erreur TypeScript, change la valeur de la propriété `guest` en `"read-only"`.

Avec ce dernier changement, ton atelier de profil utilisateur est terminé !

# --hints--

Ta propriété `guest` pour l’objet `userRoles` doit avoir la valeur `"read-only"`.

```js
assert.propertyVal(userRoles, "guest", "read-only");
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only",
  moderator: "medium-access",
--fcc-editable-region--
  guest: 3
--fcc-editable-region--
};

console.log(userRoles);
```

# --solutions--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only",
  moderator: "medium-access",
  guest: "read-only"
};

console.log(userRoles);
```
