---
id: 699b70b10ae344c802d106b0
title: Étape 12
challengeType: 1
dashedName: step-12
---

# --description--

Il est maintenant temps d’ajouter quelques propriétés de plus à l’objet `userRoles`.

Ajoute une propriété `moderator` avec la valeur `"medium-access"`.

Puis en dessous, ajoute une propriété `guest` avec la valeur `3`.

# --hints--

Ton objet `userRoles` doit avoir une propriété `moderator`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.exists(userRoles.objectProps.moderator);
```

Ta propriété `moderator` doit avoir la valeur `"medium-access"`.

```js
const explorer = await __helpers.Explorer(code);

const { userRoles } = explorer.variables;

assert.isTrue(userRoles.objectProps.moderator.value.matches("'medium-access'"));

```

Ton objet `userRoles` doit avoir une propriété `guest`.

```js
const explorer = await __helpers.Explorer(code);

const { userRoles } = explorer.variables;

assert.exists(userRoles.objectProps.guest);

```

Ta propriété `guest` doit avoir la valeur du nombre `3`.

```js
const explorer = await __helpers.Explorer(code);

const { userRoles } = explorer.variables;

assert.isTrue(userRoles.objectProps.guest.value.matches("3"));

```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--
const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only"

};
--fcc-editable-region--

console.log(userRoles);
```
