---
id: 699b61ca8c9cb9bc1968364e
title: Étape 5
challengeType: 1
dashedName: step-5
---

# --description--

Si tu ouvres la console, tu devrais voir le message d’erreur suivant :

```md
Object literal may only specify known properties, and 'mood' does not exist in type '{ username: string; age: number; isLoggedIn: boolean; }'.
```

En JavaScript classique, tu pouvais ajouter les propriétés et valeurs que tu voulais à l’objet `profile`. Mais en TypeScript, le compilateur vérifie que l’objet n’a que les propriétés que tu as définies (`username`, `age` et `isLoggedIn` dans ce cas). Comme `mood` n’était pas inclus dans le type, TypeScript le traite comme une propriété invalide et lève une erreur.

Retire `mood: null` de ton code.

# --hints--

Tu ne dois plus avoir `mood: null` dans ton code.

```js
assert.isUndefined(profile.mood);
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
--fcc-editable-region--
  mood: null
--fcc-editable-region--
};

console.log(profile);
```
