---
id: 66a37f37ef5823a313de8c26
title: Construire une orbite lunaire
challengeType: 25
dashedName: build-a-moon-orbit
demoType: onClick
---

# --description--

Dans ce lab, tu vas créer une animation simple de l'orbite de la Lune autour de la Terre avec HTML et CSS. La Terre sera au centre du système, et la Lune orbitera autour d'elle.

**Objectif :** Réponds aux user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu dois aligner tous les éléments au centre de la page en définissant la hauteur à `100%` de la fenêtre d'affichage et en définissant une mise en page flexbox sur l'élément `body`.

2. Tu dois avoir un `div` avec la classe `space`.

3. À l'intérieur de l'élément `.space`, il doit y avoir deux autres éléments `div` avec les classes `earth` et `orbit`, respectivement et dans cet ordre.

4. À l'intérieur de l'élément `.orbit`, il doit y avoir un autre `div` avec la classe `moon`.

5. L'élément `div` avec une `class` de `space` doit être centré sur la page et avoir une largeur et une hauteur de `200px`.

6. L'élément `div` avec une `class` de `space` doit utiliser le positionnement `relative`.

7. L'élément `.earth` doit utiliser le positionnement `absolute`. Place le centre de l'élément `.earth` au milieu de son parent sur les axes vertical (top) et horizontal (left). Ensuite, décale l'élément `.earth` vers l'arrière de la moitié de sa propre largeur et hauteur, afin de le centrer dans son parent, l'élément `.space`.

8. L'élément `.earth` doit avoir une largeur et une hauteur de `100px`.

9. L'élément `.orbit` doit avoir une largeur et une hauteur de `200px`.

10. L'élément `.orbit` doit être positionné avec le positionnement `absolute`. Son coin inférieur droit doit être au centre de l'élément `.space` à l'aide d'une propriété `transform` qui le décale de `-50%` sur les axes vertical et horizontal.

11. Le chemin de l'orbite de la Lune autour de la Terre doit être un cercle.

12. L'élément `.moon` doit être positionné avec le positionnement `absolute` et avoir une largeur et une hauteur de `30px`. L'élément `.moon` doit se positionner en haut de l'élément `.orbit` et être centré horizontalement.

13. Tu dois ajuster davantage le positionnement horizontal de l'élément `.moon` en déplaçant l'élément vers la gauche de la moitié de sa largeur.

14. Le sélecteur `.earth` doit avoir une couleur de fond et un `border-radius` de `50%`.

15. Le sélecteur `.moon` doit avoir une couleur de fond et un `border-radius` de `50%`.

16. Tu dois définir une animation `@keyframes orbit` qui fait tourner l'élément `.orbit` de 360 degrés autour de son centre. Tu dois appliquer cette animation à l'élément `.orbit` avec une durée de `5` secondes, une fonction de temporisation linéaire et des répétitions infinies.

**Remarque :** Assure-toi de lier ta feuille de style dans ton HTML et d'appliquer ton CSS.

# --hints--

Le contenu de l'élément `body` doit être centré sur la page en définissant la hauteur à `100%` de la hauteur de la fenêtre d'affichage, `display` à `flex`, ainsi que `justify-content` et `align-items` à `center`.

```js

assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.display, 'flex');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.justifyContent, 'center');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.alignItems, 'center');
```

Tu dois avoir un `div` principal avec la classe `space`.

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('class'), 'space');
```

L'élément `.space` doit avoir une largeur et une hauteur de `200px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.width, '200px')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.height, '200px');
```

L'élément `.space` doit avoir une propriété `position` avec la valeur `relative`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.position, 'relative')
```

Tu dois avoir un `div` avec une classe `earth` à l'intérieur du `div` avec une `class` de `space`.

```js
assert.exists(document.querySelector('div.space div.earth'));
```

L'élément `.earth` doit être le premier enfant de l'élément `.space`.

```js
assert.strictEqual(document.querySelector('.space > div')?.getAttribute('class'), 'earth');
```

L'élément `.earth` doit avoir une largeur et une hauteur de `100px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.width, '100px');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.height, '100px');
```

L'élément `.earth` doit utiliser le positionnement `absolute`.


```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.position, 'absolute');
```

L'élément `.earth` doit avoir une position `top` et `left` de `50%`.


```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.top, '50%') 
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.left, '50%');
```

L'élément `.earth` doit être déplacé vers la gauche et vers le haut de `50%` en définissant la propriété `transform` sur `translate(-50%, -50%)`.


```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.transform, 'translate(-50%, -50%)');
```

Tu dois avoir un `div` avec la classe `orbit` à l'intérieur de l'élément `.space`.

```js
assert.exists(document.querySelector('div.space div.orbit'));
```

L'élément `.orbit` doit venir après l'élément `.earth`.

```js
assert.isTrue(document.querySelector('.earth')?.nextElementSibling?.classList?.contains('orbit'));
```

L'élément `.orbit` doit avoir une largeur et une hauteur de `200px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.width, '200px')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.height, '200px');
```

L'élément `.orbit` doit utiliser le positionnement `absolute`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.position, 'absolute');
```

L'élément `.orbit` doit être positionné au centre de l'élément `.space` à l'aide d'une propriété `transform` qui le décale de `-50%` sur les axes vertical et horizontal.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.transform, 'translate(-50%, -50%)');
```

Tu dois avoir un `div` avec la classe `moon` à l'intérieur de l'élément `.orbit`.

```js
assert.exists(document.querySelector('div.space div.orbit div.moon'));
```

L'élément `.moon` doit avoir une largeur et une hauteur de `30px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.width, '30px') 
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.height, '30px');
```

L'élément `.moon` doit être positionné avec le positionnement `absolute`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.position, 'absolute');
```

Le bord supérieur de l'élément `.moon` doit être aligné avec le bord supérieur de son bloc conteneur.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.top, '0px');
```

Tu dois déplacer le bord gauche de l'élément `.moon` au centre du bloc conteneur (`50%` depuis le côté gauche).

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.left, '50%');
```

L'élément `.moon` doit être ajusté horizontalement en le translatant de `-50%` sur l'axe x.

```js
const transformValue = new __helpers.CSSHelp(document).getStyle('.moon')?.transform;
const validTransforms = ['translateX(-50%)', 'translate(-50%)'];
assert.oneOf(transformValue, validTransforms);
```

Le sélecteur `.earth` doit avoir une couleur de fond.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.earth')?.backgroundColor);
```

Le sélecteur `.earth` doit avoir un `border-radius` de `50%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.borderRadius, '50%');
```

Le sélecteur `.moon` doit avoir une couleur de fond.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.moon')?.backgroundColor);
```

Le sélecteur `.moon` doit avoir un `border-radius` de `50%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.borderRadius, '50%');
```

Tu dois avoir une règle `@keyframes`.

```js
assert.lengthOf(new __helpers.CSSHelp(document).getCSSRules('keyframes'), 1);
```

Ta nouvelle règle `@keyframes` doit être nommée `orbit`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.name, 'orbit');
```

Ta règle `@keyframes orbit` doit avoir un sélecteur `0%`.

```js
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules; 
assert.strictEqual(rules?.[0]?.keyText, '0%');
```

Ta règle `@keyframes orbit` doit avoir un sélecteur `100%`.

```js 
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules
assert.isTrue(rules?.[0]?.keyText === '100%' || rules?.[1]?.keyText === '100%');
```

Ton sélecteur `0%` doit avoir une propriété `transform` définie sur `rotate(0deg) translate(-50%, -50%)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[0]?.style?.transform, 'rotate(0deg) translate(-50%, -50%)');
```

Ton sélecteur `100%` doit venir après ton sélecteur `0%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.keyText, '100%')
```

Ton sélecteur `100%` doit avoir une propriété `transform` définie sur `rotate(360deg) translate(-50%, -50%)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.style?.transform, 'rotate(360deg) translate(-50%, -50%)')
```

Ton sélecteur `.orbit` doit avoir une propriété `animation` définie sur `orbit 5s linear infinite`.

```js

const orbitStyles = new __helpers.CSSHelp(document).getStyle('.orbit');
assert.strictEqual(orbitStyles?.animationName, "orbit");
assert.strictEqual(orbitStyles?.animationDuration, "5s");
assert.strictEqual(orbitStyles?.animationTimingFunction, "linear");
assert.strictEqual(orbitStyles?.animationIterationCount, "infinite");

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Moon Orbit</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Moon Orbiting Earth</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="space">
        <div class="earth"></div>
        <div class="orbit">
            <div class="moon"></div>
        </div>
    </div>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background-color: #000;
}

.space {
    position: relative;
    width: 200px;
    height: 200px;
}

.earth {
    width: 100px;
    height: 100px;
    background-color: blue;
    border-radius: 50%;
    position: absolute;
    top: 50%;
    left: 50%;
    transform: translate(-50%, -50%);
    box-shadow: 0 0 20px rgba(0, 0, 255, 0.5);
}

.orbit {
    width: 200px;
    height: 200px;
    position: absolute;
    transform: translate(-50%, -50%);
    animation: orbit 5s linear infinite;
}

.moon {
    width: 30px;
    height: 30px;
    background-color: gray;
    border-radius: 50%;
    position: absolute;
    top: 0;
    left: 50%;
    transform: translateX(-50%);
    box-shadow: 0 0 10px rgb(255, 255, 255);
}


@keyframes orbit {
    0% {
        transform: rotate(0deg) translate(-50%, -50%);
    }

    100% {
        transform: rotate(360deg) translate(-50%, -50%);
    }
}
```
