---
id: 587d78af367417b2b2512b03
title: Construire un formulaire de sondage
challengeType: 25
dashedName: build-a-survey-form
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**Objectif :** Respecte les rﾃｩcits utilisateur ci-dessous et fais passer tous les tests pour terminer le lab.

**Rﾃｩcits utilisateur :**

1. Tu dois avoir un titre de page dans un ﾃｩlﾃｩment `h1` avec un `id` de `title`.
1. Tu dois avoir une courte explication dans un ﾃｩlﾃｩment `p` avec un `id` de `description`.
1. Tu dois avoir un ﾃｩlﾃｩment `form` avec un `id` de `survey-form`.
1. Dans l窶凖ｩlﾃｩment de formulaire, tu dois avoir un champ `input` obligatoire pour saisir ton nom, avec un `id` de `name` et un `type` de `text`.
1. Dans l窶凖ｩlﾃｩment de formulaire, tu dois avoir un champ `input` obligatoire pour saisir ton e-mail, avec un `id` de `email`.
1. Si tu saisis un e-mail mal formatﾃｩ, tu dois voir une erreur de validation HTML5.
1. Dans l窶凖ｩlﾃｩment de formulaire, tu dois avoir un champ `input` pour saisir un nombre, avec un `id` de `number`.
1. L窶冓nput de nombre ne doit pas accepter autre chose que des nombres, soit en empﾃｪchant leur saisie, soit en affichant une erreur de validation HTML5 selon ton navigateur.
1. Si tu saisis des nombres hors de la plage de l窶冓nput de nombre, dﾃｩfinie par les attributs `min` et `max`, tu dois voir une erreur de validation HTML5.
1. Pour les champs de nom, d窶册-mail et de nombre, tu dois avoir dans le formulaire des ﾃｩlﾃｩments `label` correspondants qui dﾃｩcrivent le rﾃｴle de chaque champ avec les ids suivants : `name-label`, `email-label` et `number-label`.
1. Pour les champs de nom, d窶册-mail et de nombre, tu dois avoir un texte de placeholder qui donne une description ou des instructions pour chaque champ.
1. Dans l窶凖ｩlﾃｩment de formulaire, tu dois avoir un ﾃｩlﾃｩment dﾃｩroulant `select` avec un `id` de `dropdown` et au moins deux options parmi lesquelles choisir.
1. Dans l窶凖ｩlﾃｩment de formulaire, tu peux sﾃｩlectionner une option dans un groupe d窶兮u moins deux boutons radio regroupﾃｩs avec l窶兮ttribut `name`.
1. Dans l窶凖ｩlﾃｩment de formulaire, tu peux sﾃｩlectionner plusieurs champs dans une sﾃｩrie de cases ﾃ cocher, chacune devant avoir un attribut `value`.
1. Dans l窶凖ｩlﾃｩment de formulaire, tu dois avoir un `textarea` pour des commentaires supplﾃｩmentaires.
1. Dans l窶凖ｩlﾃｩment de formulaire, tu dois avoir un bouton avec un `id` de `submit` pour soumettre tous les inputs.

# --hints--

Tu dois avoir un ﾃｩlﾃｩment `h1` avec un `id` de `title`.

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'H1');
```

Ton `#title` ne doit pas ﾃｪtre vide.

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Tu dois avoir un ﾃｩlﾃｩment `p` avec un `id` de `description`.

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'P');
```

Ton `#description` ne doit pas ﾃｪtre vide.

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Tu dois avoir un ﾃｩlﾃｩment `form` avec un `id` de `survey-form`.

```js
const el = document.getElementById('survey-form');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'FORM');
```

Tu dois avoir un ﾃｩlﾃｩment `input` avec un `id` de `name`.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

Ton `#name` doit avoir un `type` de `text`.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.type, 'text');
```

Ton `#name` doit ﾃｪtre obligatoire.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isTrue(el.required);
```

Ton `#name` doit ﾃｪtre un descendant de `#survey-form`.

```js
const el = document.querySelector('#survey-form #name');
assert.isNotNull(el);
```

Tu dois avoir un ﾃｩlﾃｩment `input` avec un `id` de `email`.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

Ton `#email` doit avoir un `type` de `email`.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.type, 'email');
```

Ton `#email` doit ﾃｪtre obligatoire.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isTrue(el.required);
```

Ton `#email` doit ﾃｪtre un descendant de `#survey-form`.

```js
const el = document.querySelector('#survey-form #email');
assert.isNotNull(el);
```

Tu dois avoir un ﾃｩlﾃｩment `input` avec un `id` de `number`.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

Ton `#number` doit ﾃｪtre un descendant de `#survey-form`.

```js
const el = document.querySelector('#survey-form #number');
assert.isNotNull(el);
```

Ton `#number` doit avoir un `type` de `number`.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.type, 'number');
```

Ton `#number` doit avoir un attribut `min` avec une valeur numﾃｩrique.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.min);
assert.isTrue(isFinite(el.min));
```

Ton `#number` doit avoir un attribut `max` avec une valeur numﾃｩrique.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.max);
assert.isTrue(isFinite(el.max));
```

Tu dois avoir un ﾃｩlﾃｩment `label` avec un `id` de `name-label`.

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

Tu dois avoir un ﾃｩlﾃｩment `label` avec un `id` de `email-label`.

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

Tu dois avoir un ﾃｩlﾃｩment `label` avec un `id` de `number-label`.

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

Ton `#name-label` doit contenir du texte qui dﾃｩcrit l窶冓nput.

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Ton `#email-label` doit contenir du texte qui dﾃｩcrit l窶冓nput.

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Ton `#number-label` doit contenir du texte qui dﾃｩcrit l窶冓nput.

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Ton `#name-label` doit ﾃｪtre un descendant de `#survey-form`.

```js
const el = document.querySelector('#survey-form #name-label');
assert.isNotNull(el);
```

Ton `#email-label` doit ﾃｪtre un descendant de `#survey-form`.

```js
const el = document.querySelector('#survey-form #email-label');
assert.isNotNull(el);
```

Ton `#number-label` doit ﾃｪtre un descendant de `#survey-form`.

```js
const el = document.querySelector('#survey-form #number-label');
assert.isNotNull(el);
```

Ton `#name` doit avoir un attribut `placeholder` et une valeur.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

Ton `#email` doit avoir un attribut `placeholder` et une valeur.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

Ton `#number` doit avoir un attribut `placeholder` et une valeur.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

Tu dois avoir un champ `select` avec un `id` de `dropdown`.

```js
const el = document.getElementById('dropdown');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'SELECT');
```

Ton `#dropdown` doit avoir au moins deux ﾃｩlﾃｩments `option` sﾃｩlectionnables, donc non dﾃｩsactivﾃｩs.

```js
const els = document.querySelectorAll('#dropdown option:not([disabled])');
assert.isAtLeast(els.length, 2);
```

Ton `#dropdown` doit ﾃｪtre un descendant de `#survey-form`.

```js
const el = document.querySelector('#survey-form #dropdown');
assert.isNotNull(el);
```

Tu dois avoir au moins deux ﾃｩlﾃｩments `input` avec un `type` de `radio`, c窶册st-ﾃ-dire des boutons radio.

```js
const els = document.querySelectorAll('input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

Tu dois avoir au moins deux boutons radio descendants de `#survey-form`.

```js
const els = document.querySelectorAll('#survey-form input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

Tous tes boutons radio doivent avoir un attribut `value` et une valeur.

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][value=""], input[type="radio"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

Tous tes boutons radio doivent avoir un attribut `name` et une valeur.

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][name=""], input[type="radio"]:not([name])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

Chaque groupe de boutons radio doit avoir au moins 2 boutons radio.

```js
const radioButtons = document.querySelectorAll('input[type="radio"]');
const groups = {};

if (radioButtons) {
  radioButtons.forEach(el => {
    if (!groups[el.name]) groups[el.name] = [];
    groups[el.name].push(el);
  });
}

const groupKeys = Object.keys(groups);

groupKeys.forEach(key => {
  if (groups[key].length < 2) assert.fail();
});

assert.isAbove(groupKeys.length, 0);
```

Tu dois avoir au moins deux ﾃｩlﾃｩments `input` avec un `type` de `checkbox`, c窶册st-ﾃ-dire des cases ﾃ cocher, descendants de `#survey-form`.

```js
const els = document.querySelectorAll('#survey-form input[type="checkbox"]');
assert.isAtLeast(els.length, 2);
```

Toutes tes cases ﾃ cocher dans `#survey-form` doivent avoir un attribut `value` et une valeur.

```js
const els1 = document.querySelectorAll('#survey-form input[type="checkbox"]');
const els2 = document.querySelectorAll(
  '#survey-form input[type="checkbox"][value=""], #survey-form input[type="checkbox"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

Tu dois avoir au moins un ﾃｩlﾃｩment `textarea` descendant de `#survey-form`.

```js
const el = document.querySelector('#survey-form textarea');
assert.isNotNull(el);
```

Tu dois avoir un ﾃｩlﾃｩment `input` ou `button` avec un `id` de `submit`.

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.isTrue(el.tagName === 'INPUT' || el.tagName === 'BUTTON');
```

Ton `#submit` doit avoir un `type` de `submit`.

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.strictEqual(el.type, 'submit');
```

Ton `#submit` doit ﾃｪtre un descendant de `#survey-form`.

```js
const el = document.querySelector('#survey-form #submit');
assert.isNotNull(el);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Survey Form</title>
  </head>

  <body></body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html>
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="container">
      <header class="header">
        <h1 id="title" class="text-center">freeCodeCamp Survey Form</h1>
        <p id="description" class="description text-center">
          Thank you for taking the time to help us improve the platform
        </p>
      </header>
      <form id="survey-form">
        <div class="form-group">
          <label id="name-label" for="name">Name</label>
          <input
            type="text"
            name="name"
            id="name"
            class="form-control"
            placeholder="Enter your name"
            required
          />
        </div>
        <div class="form-group">
          <label id="email-label" for="email">Email</label>
          <input
            type="email"
            name="email"
            id="email"
            class="form-control"
            placeholder="Enter your Email"
            required
          />
        </div>
        <div class="form-group">
          <label id="number-label" for="number"
            >Age<span class="clue">(optional)</span></label
          >
          <input
            type="number"
            name="age"
            id="number"
            min="10"
            max="99"
            class="form-control"
            placeholder="Age"
          />
        </div>
        <div class="form-group">
          <p>Which option best describes your current role?</p>
          <select id="dropdown" name="role" class="form-control" required>
            <option disabled selected value>Select current role</option>
            <option value="student">Student</option>
            <option value="job">Full Time Job</option>
            <option value="learner">Full Time Learner</option>
            <option value="preferNo">Prefer not to say</option>
            <option value="other">Other</option>
          </select>
        </div>

        <div class="form-group">
          <p>Would you recommend freeCodeCamp to a friend?</p>
          <label>
            <input
              name="user-recommend"
              value="definitely"
              type="radio"
              class="input-radio"
              checked
            />Definitely</label
          >
          <label>
            <input
              name="user-recommend"
              value="maybe"
              type="radio"
              class="input-radio"
            />Maybe</label
          >

          <label
            ><input
              name="user-recommend"
              value="not-sure"
              type="radio"
              class="input-radio"
            />Not sure</label
          >
        </div>

        <div class="form-group">
          <p>What is your favorite feature of freeCodeCamp?</p>
          <select id="most-like" name="mostLike" class="form-control" required>
            <option disabled selected value>Select an option</option>
            <option value="challenges">Challenges</option>
            <option value="projects">Projects</option>
            <option value="community">Community</option>
            <option value="openSource">Open Source</option>
          </select>
        </div>

        <div class="form-group">
          <p>
            What would you like to see improved?
            <span class="clue">(Check all that apply)</span>
          </p>

          <label
            ><input
              name="prefer"
              value="front-end-projects"
              type="checkbox"
              class="input-checkbox"
            />Front-end Projects</label
          >
          <label>
            <input
              name="prefer"
              value="back-end-projects"
              type="checkbox"
              class="input-checkbox"
            />Back-end Projects</label
          >
          <label
            ><input
              name="prefer"
              value="data-visualization"
              type="checkbox"
              class="input-checkbox"
            />Data Visualization</label
          >
          <label
            ><input
              name="prefer"
              value="challenges"
              type="checkbox"
              class="input-checkbox"
            />Challenges</label
          >
          <label
            ><input
              name="prefer"
              value="open-source-community"
              type="checkbox"
              class="input-checkbox"
            />Open Source Community</label
          >
          <label
            ><input
              name="prefer"
              value="gitter-help-rooms"
              type="checkbox"
              class="input-checkbox"
            />Gitter help rooms</label
          >
          <label
            ><input
              name="prefer"
              value="videos"
              type="checkbox"
              class="input-checkbox"
            />Videos</label
          >
          <label
            ><input
              name="prefer"
              value="city-meetups"
              type="checkbox"
              class="input-checkbox"
            />City Meetups</label
          >
          <label
            ><input
              name="prefer"
              value="wiki"
              type="checkbox"
              class="input-checkbox"
            />Wiki</label
          >
          <label
            ><input
              name="prefer"
              value="forum"
              type="checkbox"
              class="input-checkbox"
            />Forum</label
          >
          <label
            ><input
              name="prefer"
              value="additional-courses"
              type="checkbox"
              class="input-checkbox"
            />Additional Courses</label
          >
        </div>

        <div class="form-group">
          <p>Any comments or suggestions?</p>
          <textarea
            id="comments"
            class="input-textarea"
            name="comment"
            placeholder="Enter your comment here..."
          ></textarea>
        </div>

        <div class="form-group">
          <button type="submit" id="submit" class="submit-button">
            Submit
          </button>
        </div>
      </form>
    </div>
  </body>
</html>
```
