---
id: 66bb6a9c2dd58b73cd759034
title: Crée une application de carnet de notes
challengeType: 26
dashedName: build-a-gradebook-app
---

# --description--

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois avoir une fonction nommée `getAverage` qui prend un tableau de notes de test en paramètre et renvoie la note moyenne.
1. Tu dois avoir une fonction nommée `getGrade` qui prend la note d'un étudiant en paramètre et renvoie une chaîne représentant une note alphabétique en fonction du score. Voici les scores et leurs notes alphabétiques correspondantes :

    | Plage de score | Note    |
    | -------------- | ------- |
    | `100`          | `"A+"`  |
    | `90 - 99`      | `"A"`   |
    | `80 - 89`      | `"B"`   |
    | `70 - 79`      | `"C"`   |
    | `60 - 69`      | `"D"`   |
    | `0 - 59`       | `"F"`   |

1. Tu dois avoir une fonction nommée `hasPassingGrade` qui prend un score en paramètre et renvoie soit `true` soit `false` selon que le score correspond à une note de passage ou non.
1. La fonction `hasPassingGrade` doit utiliser la fonction `getGrade` pour obtenir la note alphabétique, et s'en servir pour déterminer si la note est suffisante. Une note de passage est toute note différente de `"F"`.
1. Tu dois avoir une fonction nommée `studentMsg` qui prend un tableau de scores et la note d'un étudiant en paramètres. La fonction doit renvoyer une chaîne au format :

    - `"Class average: average-goes-here. Your grade: grade-goes-here. You passed the course."`, si l'étudiant a réussi le cours.
    - `"Class average: average-goes-here. Your grade: grade-goes-here. You failed the course."`, si l'étudiant a échoué au cours.

    Remplace `average-goes-here` par la moyenne de tous les scores et `grade-goes-here` par la note de l'étudiant.
    Utilise `getAverage` pour obtenir la note moyenne et `getGrade` pour obtenir la note de l'étudiant.

# --hints--

Tu dois avoir une fonction nommée `getAverage`.

```js
assert.isFunction(getAverage);
```

Ta fonction `getAverage` doit renvoyer un nombre.

```js
assert.isNumber(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]));
```

`getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89])` doit renvoyer `71.7`.

```js
assert.strictEqual(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 71.7);
```

`getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95])` doit renvoyer `85.4`.

```js
assert.strictEqual(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]), 85.4);
```

`getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100])` doit renvoyer `92.4`.

```js
assert.strictEqual(getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100]), 92.4);
```

`getAverage([10, 20, 30, 40, 55, 65, 75, 83])` doit renvoyer `47.25`.

```js
assert.strictEqual(getAverage([10, 20, 30, 40, 55, 65, 75, 83]), 47.25);
```

`getAverage([10, 20, 30, 40, 50, 60, 70, 97])` doit renvoyer `47.125`.

```js
assert.strictEqual(getAverage([10, 20, 30, 40, 50, 60, 70, 97]), 47.125);
```

Ta fonction `getAverage` doit renvoyer la note moyenne des scores de test.

```js
assert.strictEqual(getAverage([52, 56, 60, 65, 70, 75, 80, 85, 90, 95]), 72.8);
assert.strictEqual(getAverage([45, 50, 55, 60, 65, 70, 75, 80, 85, 90]), 67.5);
assert.strictEqual(getAverage([38, 42, 46, 50, 54, 58, 62, 66, 70, 74]), 56);
```

Tu dois avoir une fonction nommée `getGrade`.

```js
assert.isFunction(getGrade);
```

Ta fonction `getGrade` doit renvoyer `"A+"` si le score est `100`.

```js
assert.strictEqual(getGrade(100), "A+");
```

Ta fonction `getGrade` doit renvoyer `"A"` si le score est entre `90` et `99`.

```js
assert.strictEqual(getGrade(90), "A");
assert.strictEqual(getGrade(94), "A");
assert.strictEqual(getGrade(99), "A");
```

Ta fonction `getGrade` doit renvoyer `"B"` si le score est entre `80` et `89`.

```js
assert.strictEqual(getGrade(80), "B");
assert.strictEqual(getGrade(83), "B");
assert.strictEqual(getGrade(88), "B");
```

Ta fonction `getGrade` doit renvoyer `"C"` si le score est entre `70` et `79`.

```js
assert.strictEqual(getGrade(70), "C");
assert.strictEqual(getGrade(75), "C");
assert.strictEqual(getGrade(79), "C");
```

Ta fonction `getGrade` doit renvoyer `"D"` si le score est entre `60` et `69`.

```js
assert.strictEqual(getGrade(60), "D");
assert.strictEqual(getGrade(63), "D");
assert.strictEqual(getGrade(68), "D");
```

Ta fonction `getGrade` doit renvoyer `"F"` si le score est entre `0` et `59`.

```js
assert.strictEqual(getGrade(0), "F");
assert.strictEqual(getGrade(30), "F");
assert.strictEqual(getGrade(43), "F");
assert.strictEqual(getGrade(59), "F");
```


Tu dois avoir une fonction nommée `hasPassingGrade`.

```js
assert.isFunction(hasPassingGrade);
```

Ta fonction `hasPassingGrade` doit renvoyer une valeur booléenne.

```js
assert.isBoolean(hasPassingGrade(100));
```

Ta fonction `hasPassingGrade` doit renvoyer `true` si la note est un `"A"`.

```js
assert.isTrue(hasPassingGrade(100));
```

Ta fonction `hasPassingGrade` doit renvoyer `false` si la note est un `"F"`.

```js
assert.isFalse(hasPassingGrade(53));
```

Ta fonction `hasPassingGrade` doit renvoyer `true` pour toutes les notes de passage.

```js
assert.isTrue(hasPassingGrade(87));
assert.isTrue(hasPassingGrade(60));
assert.isTrue(hasPassingGrade(73));
assert.isTrue(hasPassingGrade(96));
```

Tu dois avoir une fonction nommée `studentMsg`.

```js
assert.isFunction(studentMsg);
```

`studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37)` doit renvoyer le message suivant : `"Class average: 71.7. Your grade: F. You failed the course."`.

```js
assert.strictEqual(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37), "Class average: 71.7. Your grade: F. You failed the course.");
```

`studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100)` doit renvoyer le message suivant : `"Class average: 50.8. Your grade: A+. You passed the course."`.

```js
assert.strictEqual(studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100), "Class average: 50.8. Your grade: A+. You passed the course.");
```

`studentMsg([12, 22, 32, 42, 52, 62, 72, 92], 85)` doit renvoyer le message suivant : `"Class average: 48.25. Your grade: B. You passed the course."`.

```js
assert.strictEqual(studentMsg([12, 22, 32, 42, 52, 62, 72, 92], 85), "Class average: 48.25. Your grade: B. You passed the course.");
```

`studentMsg([15, 25, 35, 45, 55, 60, 70, 60], 75)` doit renvoyer le message suivant : `"Class average: 45.625. Your grade: C. You passed the course."`.

```js
assert.strictEqual(studentMsg([15, 25, 35, 45, 55, 60, 70, 60], 75), "Class average: 45.625. Your grade: C. You passed the course.");
```

Ta fonction `studentMsg` doit renvoyer le message correct en fonction de la note de l'étudiant et de la moyenne de la classe.

```js
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 92), "Class average: 70.25. Your grade: A. You passed the course.");
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 57), "Class average: 70.25. Your grade: F. You failed the course.");
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function getAverage(scores) {
    let sum = 0;

    for (const score of scores) {
      sum += score;
    }

    return sum / scores.length;
  }

  function getGrade(score) {
    if (score === 100) {
      return "A+";
    } else if (score >= 90) {
      return "A";
    } else if (score >= 80) {
      return "B";
    } else if (score >= 70) {
      return "C";
    } else if (score >= 60) {
      return "D";
    } else {
      return "F";
    }
  }

  function hasPassingGrade(score) {
    return getGrade(score) !== "F";
  }

  function studentMsg(totalScores, studentScore) {
    let average = getAverage(totalScores);
    let grade = getGrade(studentScore);

    return `Class average: ${average}. Your grade: ${grade}. You ${
      hasPassingGrade(studentScore) ? "passed" : "failed"
    } the course.`;
  }
```
