---
id: 66faac4139dbbd5dd632c860
title: Construire un formulaire de candidature
challengeType: 25
dashedName: lab-job-application-form
demoType: onClick
---

# --description--

**Objectif :** Remplir les user stories ci-dessous et faire passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu devrais avoir un élément `div` avec la classe `container`.
1. À l'intérieur de l'élément `div`, tu devrais avoir un élément `form`.
1. Le formulaire devrait contenir un élément `input` avec le type `text` et l'id `name` pour entrer le nom complet de l'utilisateur.
1. Tu devrais avoir un autre élément `input` avec le type `email` et l'id `email` pour entrer l'adresse email de l'utilisateur.
1. Le formulaire devrait inclure un élément `select` avec l'id `position` qui permet aux utilisateurs de sélectionner un poste.
1. Tu devrais avoir un élément `fieldset` avec la classe `radio-group`.
1. À l'intérieur de `.radio-group` tu devrais avoir un ensemble d'éléments `input` avec le type `radio` et des labels pertinents pour sélectionner les options de disponibilité (par ex., Full-Time, Part-Time). Le `name` du groupe devrait être `availability`.
1. Tu devrais avoir un élément `textarea` avec l'id `message` pour entrer un message.
1. Tu devrais associer chaque élément `input` à un élément `label`.
1. Tu devrais avoir un élément `button` avec le type `submit` pour soumettre le formulaire.
1. Ajoute une pseudo-classe `:focus` aux éléments `input` et `textarea` pour changer leur couleur de bordure quand ils sont focalisés.
1. Les éléments `input`, `select` et `textarea` devraient avoir une pseudo-classe `:invalid` qui change la couleur de la bordure en rouge quand une entrée invalide est détectée.
1. Les éléments `input`, `select` et `textarea` devraient avoir une pseudo-classe `:valid` qui change la couleur de la bordure en vert quand une entrée valide est saisie.
1. L'élément `button` devrait avoir une pseudo-classe `:hover` qui change la couleur d'arrière-plan quand il est survolé.
1. Utilise la pseudo-classe `:checked` sur `.radio-group input[type="radio"]` pour ajouter une couleur de bordure, une couleur d'arrière-plan et une ombre de boîte quand le bouton radio est sélectionné.
1. Utilise la pseudo-classe `:checked` sur les boutons radio pour changer la couleur du texte du `label` associé quand l'option est sélectionnée.
1. Ajoute une pseudo-classe `:first-of-type` à l'élément `input` pour styliser le premier champ d'entrée différemment. (par ex., coins arrondis).

**Note :** Assure-toi de lier ta feuille de style dans ton HTML et d'appliquer ton CSS.

# --hints--

Tu devrais avoir un élément `div` avec la classe `container`.

```js
assert.exists(document.querySelector("div.container"));
```

À l'intérieur de l'élément `div`, tu devrais avoir un élément `form`.

```js
assert.exists(document.querySelector("div.container > form"));
```

Le formulaire devrait contenir un élément `input` avec le type `text` et l'id `name` pour entrer le nom complet de l'utilisateur.

```js
assert.exists(document.querySelector("div.container > form input#name[type='text']"));
```

Tu devrais avoir un autre élément `input` avec le type `email` et l'id `email` pour entrer l'adresse email de l'utilisateur.

```js
assert.exists(document.querySelector("div.container > form input#email[type='email']"));
```

Le formulaire devrait inclure un élément `select` avec l'id `position` avec quelques éléments `option`.

```js
assert.exists(document.querySelector("div.container > form select#position"));
assert.isNotEmpty(document.querySelectorAll("div.container > form select#position > option"))
```

Tu devrais avoir un élément `fieldset` ou `section` avec la classe `.radio-group`.

```js
assert.lengthOf(document.querySelectorAll('form fieldset.radio-group, form section.radio-group'), 1)
```

À l'intérieur de `.radio-group` tu devrais avoir un groupe d'éléments `input` avec le type `radio` pour sélectionner les options de disponibilité. Le `name` du groupe devrait être `availability`.

```js
assert.isAtLeast(document.querySelectorAll("div.container > form .radio-group input[type='radio'][name='availability']").length, 2);
assert.isEmpty(document.querySelectorAll("div.container > form .radio-group input[type='radio']:not([name='availability'])"))
```

Tu devrais avoir un élément `textarea` avec l'id `message` pour entrer un message.

```js
assert.isNotEmpty(document.querySelectorAll("div.container > form textarea#message"));
```

Tu devrais associer chaque élément `input` à un élément `label`.

```js
const els = document.querySelectorAll('input');
assert.isNotEmpty(els);
els.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label); 
  assert.equal(label.tagName, "LABEL");
})
```

Tu devrais avoir un élément `button` avec le type `submit` pour soumettre le formulaire.

```js
assert.isNotEmpty(document.querySelectorAll("div.container > form button[type='submit']"));
```

Tu devrais ajouter une pseudo-classe `:focus` aux éléments `input` et `textarea` pour changer leur couleur de bordure quand ils sont focalisés. Utilise un sélecteur de liste dans l'ordre donné.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input:focus, textarea:focus').getPropertyValue('border-color'))
```

Les éléments `input`, `select` et `textarea` devraient avoir une pseudo-classe `:invalid` qui change la couleur de la bordure en rouge quand une entrée invalide est détectée. Utilise un sélecteur de liste dans l'ordre donné.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input:invalid, select:invalid, textarea:invalid').getPropertyValue('border-color'), 'red')
```

Les éléments `input`, `select` et `textarea` devraient avoir une pseudo-classe `:valid` qui change la couleur de la bordure en vert quand une entrée valide est saisie. Utilise un sélecteur de liste dans l'ordre donné.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input:valid, select:valid, textarea:valid').getPropertyValue('border-color'), 'green')
```

L'élément `button` devrait avoir une pseudo-classe `:hover` qui change la couleur d'arrière-plan quand il est survolé.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('button:hover').getPropertyValue('background-color'))
```

Tu devrais utiliser la pseudo-classe `:checked` sur `.radio-group input[type="radio"]` pour ajouter une couleur de bordure quand le bouton radio est sélectionné.

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('border-color'));
```

Tu devrais utiliser la pseudo-classe `:checked` sur `.radio-group input[type="radio"]` pour ajouter une couleur d'arrière-plan quand le bouton radio est sélectionné.

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('background-color'));
```

Tu devrais utiliser la pseudo-classe `:checked` sur `.radio-group input[type="radio"]` pour ajouter une ombre de boîte quand le bouton radio est sélectionné.

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('box-shadow'));
```

Tu devrais utiliser la pseudo-classe `:checked` sur les boutons radio pour changer la couleur du texte du `label` associé quand l'option est sélectionnée.

```js
const els = document.querySelectorAll('input[type="radio"]');
assert.isNotEmpty(els);
els.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.equal(label.tagName, "LABEL");
  input.checked = false;
  const colorBefore = getComputedStyle(label).color;
  input.checked = true;
  const colorAfter = getComputedStyle(label).color;
  assert.notEqual(colorAfter, colorBefore);
})
```

Ajoute une pseudo-classe `:first-of-type` aux éléments `input` pour styliser le premier champ d'entrée différemment.

```js
const pseudo = new __helpers.CSSHelp(document).getStyle('input:first-of-type');
assert.exists(pseudo);
assert.isNotEmpty([...pseudo]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Job Application Form</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Job Application</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="container">
        <h1>Job Application Form</h1>
        <form id="job-application">
            <label for="name">Full Name:</label>
            <input type="text" id="name" name="name" placeholder="Enter your name" required>

            <label for="email">Email:</label>
            <input type="email" id="email" name="email" placeholder="Enter your email" required>

            <label for="position">Position:</label>
            <select id="position" name="position" required>
                <option value="">Select a position</option>
                <option value="developer">Developer</option>
                <option value="designer">Designer</option>
                <option value="manager">Manager</option>
            </select>

            <fieldset class="radio-group">
              <legend>Availability:</legend>
                <input type="radio" id="fullTime" name="availability" value="fullTime" checked>
                <label for="fullTime">Full-Time</label>
                <input type="radio" id="partTime" name="availability" value="partTime">
                <label for="partTime">Part-Time</label>
            </fieldset>

            <label for="message">Why do you want this job?</label>
            <textarea id="message" name="message" rows="5" placeholder="Write your motivation" required></textarea>

            <button type="submit">Submit Application</button>
        </form>
    </div>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  background-color: white;
  padding: 20px;
}

.container {
  max-width: 600px;
  margin: 0 auto;
  background-color: white;
  padding: 20px;
  border-radius: 8px;
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
}

h1 {
  text-align: center;
}

label {
  display: block;
  margin-bottom: 5px;
  font-weight: bold;
}

input, select, textarea {
  width: 100%;
  padding: 10px;
  margin-bottom: 15px;
  border: 1px solid #ccc;
  border-radius: 4px;
  box-sizing: border-box;
}

.radio-group {
  display: flex;
  align-items: center;
  margin-top: 15px;
}
.radio-group input[type="radio"] {
  margin-right: 10px;
  -webkit-appearance: none;
  -moz-appearance: none;
  appearance: none;
  width: 20px;
  height: 20px;
  border: 2px solid #ccc;
  border-radius: 50%;
  outline: none;
  transition: border-color 0.2s;
}

.radio-group label {
  margin: 0;
  font-weight: normal;
  cursor: pointer;
}

button {
  width: 100%;
  padding: 10px;
  background-color: green;
  border: none;
  color: white;
  border-radius: 4px;
  font-size: 16px;
  cursor: pointer;
}

input:focus, textarea:focus {
  border-color: blue;
  outline: none;
}

input:invalid, select:invalid, textarea:invalid {
  border-color: red;
}

input:valid, select:valid, textarea:valid {
  border-color: green;
}

.radio-group input[type="radio"]:checked {
  border-color: green;
  background-color: green;
  box-shadow: inset 0 0 0 4px white;
}

button:hover {
  background-color: black;
}

input[type="radio"]:checked + label {
  color: green;
}

button:disabled {
  background-color: grey;
}

input:first-of-type {
  border-top-left-radius: 4px;
  border-top-right-radius: 4px;
}

textarea:nth-child(5) {
  border-bottom-left-radius: 4px;
  border-bottom-right-radius: 4px;
}
```
