---
id: 6716249b5405164036fd0b0d
title: Crée un visualiseur de tri
challengeType: 25
dashedName: build-a-sorting-visualizer
demoType: onClick
---

# --description--

L'algorithme du tri à bulles (Bubble Sort) trie une séquence d'entiers en comparant des couples d'éléments adjacents à partir du début de la séquence. Si le premier élément est plus grand que le second, il les échange. Puis il poursuit avec le couple suivant. Lorsque le dernier élément de la séquence est atteint, il recommence un nouveau cycle depuis le début de la séquence, et répète le processus jusqu'à ce que les éléments soient triés. L'algorithme s'arrête après qu'un cycle se termine sans aucun échange.

Pour ce lab, tout le HTML et le CSS t'ont été fournis. Tu vas utiliser JavaScript pour terminer le visualiseur de tri à bulles afin qu'il visualise chaque étape nécessaire à l'algorithme Bubble Sort pour trier un tableau de cinq entiers.

**Objectif :** Remplis les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu dois avoir une fonction nommée `generateElement` qui renvoie un entier aléatoire entre `1` et `100`, inclus.
1. Tu dois avoir une fonction nommée `generateArray` qui utilise la fonction `generateElement` pour renvoyer un tableau contenant cinq entiers aléatoires.
1. Tu dois avoir une fonction nommée `generateContainer` qui crée et renvoie un élément `div` vide.
1. Tu dois avoir une fonction nommée `fillArrContainer` qui prend un élément HTML comme premier argument et un tableau comme second argument.
1. `fillArrContainer` doit prendre un élément comme premier paramètre et un tableau d'entiers comme second paramètre, puis remplir l'élément avec cinq éléments `span`, chaque `span` affichant l'un des entiers du tableau.
1. Tu dois avoir une fonction nommée `isOrdered` qui prend deux entiers et renvoie un booléen indiquant si le premier entier est inférieur ou égal au second.
1. Tu dois avoir une fonction nommée `swapElements` qui prend un tableau d'entiers et un indice numérique.
1. La fonction `swapElements` doit modifier le tableau en place en échangeant l'élément à l'indice passé et l'élément suivant si `isOrdered` renvoie `false`.
1. Tu dois avoir une fonction nommée `highlightCurrentEls` qui prend un élément HTML et un indice numérique.
1. La fonction `highlightCurrentEls` doit définir la `border` de l'enfant de l'élément donné à l'indice donné, et de l'enfant immédiatement après l'indice, avec un style `dashed`, une couleur `red`, et une largeur de ton choix.
1. Lorsque tu cliques sur `#generate-btn`, tu dois utiliser la fonction `fillArrContainer` pour remplir `#starting-array` avec cinq éléments `span`, chacun avec un nombre aléatoire comme texte. S'ils sont présents, les autres éléments doivent être retirés de `#array-container`.
1. Tu dois implémenter l'algorithme Bubble Sort de sorte qu'après avoir cliqué sur `#sort-btn`, `#array-container` contienne un élément `div` pour chacune des étapes requises par l'algorithme Bubble Sort pour trier le tableau de départ, y compris le `div` représentant le tableau de départ et un `div` représentant le tableau trié. Les fonctions que tu as créées jusqu'ici peuvent être utiles ici.
1. Chaque `div` doit contenir cinq éléments `span`, représentant le tableau dans son état actuel de tri.
1. Après avoir cliqué sur `#sort-btn`, `#starting-array` doit représenter l'étape de départ avec le tableau initial et les deux premiers entiers surlignés.
1. Pour chaque étape de tri, tu dois utiliser `highlightCurrentEls` pour surligner les deux nombres en cours de comparaison, et les échanger à l'étape suivante en utilisant `swapElements`.

# --hints--

Tu dois avoir une fonction nommée `generateElement`.

```js
assert.isFunction(generateElement);
```

Ta fonction `generateElement` doit renvoyer un entier aléatoire entre `1` et `100` inclus.

```js
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {
    assert.strictEqual(generateElement(), 26);
    assert.strictEqual(generateElement(), 9);
} finally {
    randomMocker.restore();
}
```

Tu dois avoir une fonction nommée `generateArray`.

```js
assert.isFunction(generateArray)
```

Ta fonction `generateArray` doit utiliser la fonction `generateElement`.

```js
let flag = false;
const temp = generateElement;
generateElement = () => flag = true;
try {
    generateArray();
    assert.isTrue(flag);
} finally {
    generateElement = temp;
}
```

Ta fonction `generateArray` doit renvoyer un tableau contenant cinq entiers aléatoires entre `1` et `100`.

```js
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {
    assert.deepEqual(generateArray(), [26, 9, 58, 23, 38])
} finally {
    randomMocker.restore();
}
```

Tu dois avoir une fonction nommée `generateContainer`.

```js
assert.isFunction(generateContainer);
```

Ta fonction `generateContainer` doit renvoyer un élément `div` vide.

```js
const div = generateContainer();
assert.equal(div.tagName, "DIV");
assert.isEmpty(div.children);
```

Tu dois avoir une fonction nommée `fillArrContainer`.

```js
assert.isFunction(fillArrContainer);
```

Ta fonction `fillArrContainer()` doit prendre un élément comme premier paramètre et un tableau d'entiers comme second paramètre, puis remplir l'élément avec cinq éléments `span`, chaque `span` affichant l'un des entiers du tableau.

```js
const testDiv = document.createElement("div");
const testArr = [15, 98, 17, 5, 63]
fillArrContainer(testDiv, testArr);
const children = testDiv.children;
assert.lengthOf(children, 5);
Array.from(children).forEach((el, i) => {
    assert.equal(el.tagName, "SPAN");
    assert.equal(el.innerText.trim(), testArr[i])
})
```

Tu dois avoir une fonction nommée `isOrdered`.

```js
assert.isFunction(isOrdered);
```

Ta fonction `isOrdered` doit prendre deux entiers et renvoyer un booléen indiquant si le premier entier est inférieur ou égal au second.

```js
assert.lengthOf(isOrdered, 2);
assert.isTrue(isOrdered(2, 60));
assert.isFalse(isOrdered(10, 3));
assert.isTrue(isOrdered(5, 5));
```

Tu dois avoir une fonction nommée `swapElements`.

```js
assert.isFunction(swapElements);
```

Ta fonction `swapElements` prend un tableau d'entiers et un indice numérique comme arguments. Elle doit modifier le tableau passé en place en échangeant l'élément à l'indice donné et l'élément suivant si le premier élément est plus grand que le second.

```js
const testArr = [22, 4, 87, 47, 33];
swapElements(testArr, 0);
assert.deepEqual(testArr, [4, 22, 87, 47, 33]);
swapElements(testArr, 1);
assert.deepEqual(testArr, [4, 22, 87, 47, 33]);
swapElements(testArr, 2);
assert.deepEqual(testArr, [4, 22, 47, 87, 33]);
swapElements(testArr, 3);
assert.deepEqual(testArr, [4, 22, 47, 33, 87]);
```

Tu dois avoir une fonction nommée `highlightCurrentEls`.

```js
assert.isFunction(highlightCurrentEls);
```

Ta fonction `highlightCurrentEls` doit donner aux descendants de l'élément spécifié, situés à l'indice donné et à l'indice suivant, une bordure `dashed`, `red`, et d'une largeur de ton choix.

```js
const testDiv = document.createElement("div");
document.querySelector("body").appendChild(testDiv)
for (let i = 0; i < 5; i++) {
    testDiv.appendChild(document.createElement("span"));
}
const redBorderRegex = /dashed (rgb\(255,\s*0,\s*0\)|#FF0000|#F00|hsl\(0,\s*100%,\s*50%\))/;
const revertBorder = () => {
    for (const el of children) {
        el.style.border = "revert";
    }
}
const children = testDiv.children;

for (let i = 0; i < 3; i++) {
    highlightCurrentEls(testDiv, i);
    for (let j = 0; j < 5; j++) {
        let b = getComputedStyle(children[j]).border;
        if (j == i || j == i + 1) {
            assert.match(b, redBorderRegex);
            assert.isAbove(parseFloat(b), 0);
        } else {
            assert.notMatch(b, redBorderRegex);
        }
    }
    revertBorder();
}
testDiv.remove();
```

Lorsque tu cliques sur `#generate-btn`, tu dois remplir `#starting-array` avec cinq éléments `span`, chacun avec un nombre aléatoire entre `1` et `100` comme texte.

```js
const genBtn = document.querySelector("#generate-btn");
genBtn.click();
const children = document.querySelector("#starting-array").querySelectorAll("span");
assert.lengthOf(children, 5);
Array.from(children).forEach(el => {
    assert.equal(el.tagName, "SPAN");
    const num = Number(el.innerText.trim());
    assert.isAtMost(num, 100);
    assert.isAtLeast(num, 1);
})
```

Lorsque `#starting-array` contient déjà un tableau généré, ou que `#array-container` contient le tableau trié, cliquer sur `#generate-btn` doit retirer les autres éléments de `#array-container`, en ne laissant que `#starting-array` avec des nombres nouvellement générés.

```js
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
genBtn.dispatchEvent(new Event("click"));

const prevNumbers = Array.from(document.querySelector("#starting-array").querySelectorAll("span")).map(el => Number(el.innerText.trim()));

sortBtn.dispatchEvent(new Event("click"));
genBtn.dispatchEvent(new Event("click"));

const container = document.querySelector("#array-container");
assert.lengthOf(container.children, 1);
const numbers = Array.from(document.querySelector("#starting-array").querySelectorAll("span")).map(el => Number(el.innerText.trim()));
assert.lengthOf(numbers, 5);
assert.isTrue(prevNumbers.some((num, index) => num !== numbers[index]))
```

Après avoir cliqué sur `#sort-btn`, `#array-container` doit contenir autant d'éléments `div` que d'étapes requises par l'algorithme Bubble Sort pour trier le tableau de départ, y compris le `div` représentant le tableau de départ et un `div` représentant le tableau trié.

```js
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
// using randomMocker to be sure that the starting array requires 13 steps to be sorted
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {
    genBtn.dispatchEvent(new Event("click"));
    sortBtn.dispatchEvent(new Event("click"));
    const container = document.querySelector("#array-container");
    assert.lengthOf(container.children, 13)
    Array.from(container.children).forEach(el => {assert.equal(el.tagName, "DIV")})
} finally {
    randomMocker.restore();
}
```

Après avoir cliqué sur `#sort-btn`, chaque `div` dans `#array-container` doit contenir cinq `span`, chacun avec un nombre comme texte, et arrangés pour représenter les étapes requises par l'algorithme Bubble Sort pour trier le tableau de départ.

```js
const finalArr = [9, 23, 26, 38, 58];
const arrays = [
    [26, 9, 58, 23, 38],
    [9, 26, 58, 23, 38],
    [9, 26, 58, 23, 38],
    [9, 26, 23, 58, 38],
    [9, 26, 23, 38, 58],
    [9, 26, 23, 38, 58],
    finalArr,
    finalArr,
    finalArr,
    finalArr,
    finalArr,
    finalArr,
    finalArr
]
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {
    genBtn.dispatchEvent(new Event("click"));
    sortBtn.dispatchEvent(new Event("click"));
    const container = document.querySelector("#array-container");
    assert.isNotEmpty(container.children);
    Array.from(container.children).forEach((el, i) => {
        Array.from(el.children).forEach((j, k) => {
            assert.strictEqual(Number(j.innerText.trim()), arrays[i][k])
        })
    })
} finally {
    randomMocker.restore();
}
```

Lorsque tu cliques sur `#sort-btn`, tu dois utiliser la fonction `highlightCurrentEls` pour surligner les éléments comparés à chaque étape.

```js
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
let flag = false;
const temp = highlightCurrentEls;
highlightCurrentEls = () => flag = true;
try {
    genBtn.dispatchEvent(new Event("click"));
    sortBtn.dispatchEvent(new Event("click"));
    assert.isTrue(flag);
} finally {
    highlightCurrentEls = temp;
}
```

Après avoir cliqué sur `#sort-btn`, `#starting-array` doit représenter l'étape de départ avec le tableau initial et les deux premiers entiers surlignés à l'aide de `highlightCurrentEls`.

```js
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
genBtn.dispatchEvent(new Event("click"));
sortBtn.dispatchEvent(new Event("click"));
const firstContainer = document.querySelector("#starting-array");
const children = firstContainer.children
const redBorderRegex = /dashed (rgb\(255,\s*0,\s*0\)|#FF0000|hsl\(0,\s*100%,\s*50%\))/;
assert.match(getComputedStyle(children[0]).border, redBorderRegex);
assert.match(getComputedStyle(children[1]).border, redBorderRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <meta http-equiv="X-UA-Compatible" content="ie=edge">
    <title>Sorting Visualizer</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main>
        <div id="array-container">
            <div id="starting-array"></div>
        </div>
        <div id="btn-container">
            <button id="generate-btn" type="button">Generate Array</button>
            <button id="sort-btn" type="button">Sort Array</button>
        </div>
    </main>
    <script src="script.js"></script>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
}

main {
    height: 100vh;
    display: flex;
    justify-content: center;
    flex-direction: column;
    align-items: center;
}

#array-container {
    max-height: 95vh;
    display: flex;
    flex-direction: column;
    flex-wrap: wrap;
    gap: 2px;

}

#array-container>div {
    min-width: 8rem;
    height: 2rem;
    box-shadow: rgba(50, 50, 93, 0.25) 0px 2px 5px -1px, rgba(0, 0, 0, 0.3) 0px 1px 3px -1px;
    border-radius: 10px;
    margin-bottom: 0.2rem;
    border: 2px solid darkgray;
    display: flex;
    justify-content: space-evenly;
    align-items: center;
}

#starting-array {
    border: 4px solid darkblue !important;
}

#btn-container {
    display: flex;
    justify-content: space-around;
}

button {
    padding: 2px;
    margin: 5px;
}

span {
    border-radius: 2px;
    padding: 0.5px;
    margin: 0
}

@media (min-width: 430px) {
  #array-container>div {
    min-width: 12rem;
  }
  span {
    padding: 1px;
    margin: 1px;
  }
}
```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <meta http-equiv="X-UA-Compatible" content="ie=edge">
    <title>Sorting Visualizer</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main>
        <div id="array-container">
            <div id="starting-array"></div>
        </div>
        <div id="btn-container">
            <button id="generate-btn" type="button">Generate Array</button>
            <button id="sort-btn" type="button">Sort Array</button>
        </div>
    </main>
    <script src="script.js"></script>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
}

main {
    height: 100vh;
    display: flex;
    justify-content: center;
    flex-direction: column;
    align-items: center;
}

#array-container {
    max-height: 95vh;
    display: flex;
    flex-direction: column;
    flex-wrap: wrap;
    gap: 2px;

}

#array-container>div {
    min-width: 8rem;
    height: 2rem;
    box-shadow: rgba(50, 50, 93, 0.25) 0px 2px 5px -1px, rgba(0, 0, 0, 0.3) 0px 1px 3px -1px;
    border-radius: 10px;
    margin-bottom: 0.2rem;
    border: 2px solid darkgray;
    display: flex;
    justify-content: space-evenly;
    align-items: center;
}

#starting-array {
    border: 4px solid darkblue !important;
}

#btn-container {
    display: flex;
    justify-content: space-around;
}

button {
    padding: 2px;
    margin: 5px;
}

#sort-btn {
   display: none
}

span {
    border-radius: 2px;
    padding: 0.5px;
    margin: 0
}

@media (min-width: 430px) {
  #array-container>div {
    min-width: 12rem;
  }
  span {
    padding: 1px;
    margin: 1px;
  }
}
```

```js
const arrayContainer = document.getElementById("array-container");
const generateArrayBtn = document.getElementById("generate-btn");
const sortArrayBtn = document.getElementById("sort-btn");
const arrayLength = 5;
const minVal = 1;
const maxVal = 100;
const isStart = () => arrayContainer.children.length === 1;
const clearArrayContainer = () => { arrayContainer.innerHTML = '<div id="starting-array"></div>' };
const showSortBtn = () => {sortArrayBtn.style.display = "inline-block"};
const hideSortBtn = () => {sortArrayBtn.style.display = "none"};
const generateElement = () => Math.floor(Math.random() * maxVal + minVal);
const generateArray = () => Array.from({ length: arrayLength }, generateElement);
const fillArrContainer = (container, arr) => {
    container.innerHTML = "";
    arr.forEach(i => {
        const el = document.createElement("span");
        el.innerText = i;
        el.id = `number-${i}`
        container.appendChild(el);
    })
}
const isOrdered = (el1, el2) => el1 <= el2;

const swapElements = (arr, n = 0) => {
    if (n < arr.length - 1 && !isOrdered(arr[n], arr[n + 1])) {
        const temp = arr[n];
        arr[n] = arr[n + 1];
        arr[n + 1] = temp;
    }
}
const getLastChildren = () => arrayContainer.lastElementChild;
const getLastArr = () => {
    const els = Array.from(getLastChildren().children);
    const arr = els.map(el => Number(el.id.replace("number-", "")));
    return arr;
}
const generateContainer = () => {
    const container = document.createElement("div");
    arrayContainer.appendChild(container);
    return container;
}
const highlightCurrentEls = (container, n = 0) => {
    const children = container.children;
    children[n].style.border = "2px dashed red";
    children[n + 1].style.border = "2px dashed red";
}

const highlightSorted = () => {
    getLastChildren().style.border = "4px solid green";
}

const bubbleSort = () => {
    let swapped = true;
    while (swapped) {
        const startingArr = getLastArr();
        startingArr.forEach((_, i) => {
            if (i + 1 < arrayLength) {
                highlightCurrentEls(getLastChildren(), i);
                const arr = getLastArr();
                swapElements(arr, i);
                fillArrContainer(generateContainer(), arr);
            }
        })
        const lastArr = getLastArr()
        if (startingArr.every((el, i) => el === lastArr[i])) swapped = false;
    }
}


generateArrayBtn.addEventListener("click", () => {
    if (!isStart()) {
        clearArrayContainer();
    }
    fillArrContainer(document.getElementById("starting-array"), generateArray());
    showSortBtn();
});

sortArrayBtn.addEventListener("click", () => {
    bubbleSort();
    highlightSorted();
    hideSortBtn();
})
```
