---
id: 66db529d37ad966480ebb633
title: Crée un programme de sélection de déjeuner
challengeType: 26
dashedName: build-a-lunch-picker-program
---

# --description--

Dans ce lab, tu vas créer un programme qui aide à gérer des options de déjeuner. Tu travailleras avec un tableau de déjeuners, tu ajouteras et supprimeras des éléments du tableau, et tu sélectionneras une option de déjeuner au hasard.

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois créer une variable `lunches` et lui affecter un tableau vide qui servira à stocker les éléments de déjeuner.

2. Tu dois créer une fonction `addLunchToEnd` qui prend un tableau comme premier argument et une chaîne comme second argument. La fonction doit :
    - Ajouter la chaîne à la fin du tableau.
    - Afficher la chaîne `[Lunch Item] added to the end of the lunch menu.` dans la console, où `[Lunch Item]` est la chaîne passée à la fonction.
    - Retourner le tableau mis à jour.

3. Tu dois créer une fonction `addLunchToStart` qui prend un tableau comme premier argument et une chaîne comme second argument. La fonction doit :
    - Ajouter la chaîne au début du tableau.
    - Afficher la chaîne `[Lunch Item] added to the start of the lunch menu.` dans la console, où `[Lunch Item]` est la chaîne passée à la fonction.
    - Retourner le tableau mis à jour.

4. Tu dois créer une fonction `removeLastLunch` qui prend un tableau comme argument. La fonction doit :
   - Supprimer le dernier élément du tableau.
   - Si la suppression réussit, afficher la chaîne `[Lunch Item] removed from the end of the lunch menu.` dans la console, où `[Lunch Item]` est l'élément supprimé du tableau.
   - Si le tableau est vide, afficher la chaîne `"No lunches to remove."` dans la console.
   - Retourner le tableau mis à jour.

5. Tu dois créer une fonction `removeFirstLunch` qui prend un tableau comme argument. La fonction doit :
   - Supprimer le premier élément du tableau.
   - Si la suppression réussit, afficher la chaîne `[Lunch Item] removed from the start of the lunch menu.` dans la console, où `[Lunch Item]` est l'élément supprimé du tableau.
   - Si le tableau est vide, afficher la chaîne `"No lunches to remove."` dans la console.
   - Retourner le tableau mis à jour.

6. Tu dois créer une fonction `getRandomLunch` qui prend un tableau comme argument. La fonction doit :
   - Sélectionner un élément au hasard dans le tableau.
   - Si elle réussit, afficher la chaîne `Randomly selected lunch: [Lunch Item]` dans la console, où `[Lunch Item]` est un élément aléatoire du tableau.
   - Si le tableau est vide, afficher la chaîne `"No lunches available."` dans la console.

7. Tu dois créer une fonction `showLunchMenu` qui prend un tableau comme argument et :
   - S'il y a des éléments dans le tableau, affiche la chaîne `Menu items: [Lunch Item], [Lunch Item]...` dans la console, où chaque `[Lunch item]` est l'un des éléments du tableau, dans l'ordre.
   - Si le tableau est vide, affiche la chaîne `"The menu is empty."` dans la console.

# --hints--

Tu dois déclarer une variable `lunches` et lui affecter un tableau vide pour stocker les éléments de déjeuner.

```js
assert.isArray(lunches);
```

Tu dois définir une fonction `addLunchToEnd`.

```js
assert.isFunction(addLunchToEnd);
```

La fonction `addLunchToEnd` doit avoir deux paramètres.

```js
assert.lengthOf(addLunchToEnd, 2);
```

`addLunchToEnd(lunches, "Tacos")` doit afficher la chaîne `"Tacos added to the end of the lunch menu."` dans la console.

```js
const tempArr = [];
const temp = console.log;
const testLunches = []
try {
  console.log = obj => tempArr.push(obj);
  addLunchToEnd(testLunches, "Tacos");
  assert.strictEqual(tempArr[0], "Tacos added to the end of the lunch menu.");
  addLunchToEnd(testLunches, "Pizza");
  assert.strictEqual(tempArr[1], "Pizza added to the end of the lunch menu.");
} finally {
  console.log = temp;
}
```

`addLunchToEnd(["Pizza", "Tacos"], "Burger")` doit retourner `["Pizza", "Tacos", "Burger"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(addLunchToEnd(["Pizza", "Tacos"], "Burger"), ["Pizza", "Tacos", "Burger"]);
  assert.deepStrictEqual(addLunchToEnd(["Fries"], "Tacos"), ["Fries", "Tacos"]);
} finally {
  console.log = temp;
}
```

Tu dois définir une fonction `addLunchToStart`.

```js
assert.isFunction(addLunchToStart);
```

La fonction `addLunchToStart` doit avoir deux paramètres.

```js
assert.lengthOf(addLunchToStart, 2);
```

`addLunchToStart(lunches, "Sushi")` doit afficher la chaîne `"Sushi added to the start of the lunch menu."` dans la console.

```js
const tempArr = [];
const testLunches = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  addLunchToStart(testLunches, "Sushi");
  assert.strictEqual(tempArr[0], "Sushi added to the start of the lunch menu.");
  addLunchToStart(testLunches, "Burger");
  assert.strictEqual(tempArr[1], "Burger added to the start of the lunch menu.");
} finally {
  console.log = temp;
}
```

`addLunchToStart(["Burger", "Sushi"], "Pizza")` doit retourner `["Pizza", "Burger", "Sushi"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(addLunchToStart(["Burger", "Sushi"], "Pizza"), ["Pizza", "Burger", "Sushi"]);
  assert.deepStrictEqual(addLunchToStart(["Noodles"], "Sushi"), ["Sushi", "Noodles"]);
} finally {
  console.log = temp;
}
```

Tu dois définir une fonction `removeLastLunch`.

```js
assert.isFunction(removeLastLunch)
```

La fonction `removeLastLunch` doit avoir un paramètre.

```js
assert.lengthOf(removeLastLunch, 1);
```

Lorsque le tableau d'entrée est vide, la fonction `removeLastLunch` doit afficher la chaîne `"No lunches to remove."` dans la console.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches to remove.");
} finally {
  console.log = temp;
}
```

`removeLastLunch(["Stew", "Soup", "Toast"])` doit afficher la chaîne `"Toast removed from the end of the lunch menu."` dans la console.

```js
let testLunches = ["Stew", "Soup", "Toast"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[0], "Toast removed from the end of the lunch menu.");
  testLunches = ["Rice", "Pizza"];
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[1], "Pizza removed from the end of the lunch menu.");
} finally {
  console.log = temp;
}
```

`removeLastLunch(["Sushi", "Pizza", "Noodles"])` doit retourner `["Sushi", "Pizza"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(removeLastLunch(["Sushi", "Pizza", "Noodles"]), ["Sushi", "Pizza"]);
  assert.deepStrictEqual(removeLastLunch(["Pizza", "Burger"]), ["Pizza"]);
} finally {
  console.log = temp;
}
```

Tu dois définir une fonction `removeFirstLunch`.

```js
assert.isFunction(removeFirstLunch);
```

La fonction `removeFirstLunch` doit avoir un seul paramètre.

```js
assert.lengthOf(removeFirstLunch, 1);
```

Lorsque le tableau d'entrée est vide, la fonction `removeFirstLunch` doit afficher la chaîne `"No lunches to remove."` dans la console.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches to remove.");
} finally {
  console.log = temp;
}
```

`removeFirstLunch(["Salad", "Eggs", "Cheese"])` doit afficher la chaîne `"Salad removed from the start of the lunch menu."` dans la console.

```js
let testLunches = ["Salad", "Eggs", "Cheese"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[0], "Salad removed from the start of the lunch menu.");
  testLunches = ["Stew", "Soup", "Toast"]
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[1], "Stew removed from the start of the lunch menu.");
} finally {
  console.log = temp;
}
```

`removeFirstLunch(["Sushi", "Pizza", "Burger"])` doit retourner `["Pizza", "Burger"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(removeFirstLunch(["Sushi", "Pizza", "Burger"]), ["Pizza", "Burger"]);
  assert.deepStrictEqual(removeFirstLunch(["Pizza", "Burger"]), ["Burger"]);
} finally {
  console.log = temp;
}
```

`addLunchToEnd`, `addLunchToStart`, `removeLastLunch` et `removeFirstLunch` doivent retourner le même tableau passé en argument après l'avoir mis à jour.

```js
const temp = console.log;
console.log = () => {}
const testLunches = ["Spam"];
try {
  assert.equal(addLunchToEnd(testLunches, "Pizza"), testLunches);
  assert.equal(addLunchToStart(testLunches, "Caviar"), testLunches);
  assert.equal(removeLastLunch(testLunches), testLunches);
  assert.equal(removeFirstLunch(testLunches), testLunches);
} finally {
  console.log = temp;
}
```

Tu dois définir une fonction `getRandomLunch`.

```js
assert.isFunction(getRandomLunch);
```

La fonction `getRandomLunch` doit avoir un seul paramètre.

```js
assert.lengthOf(getRandomLunch, 1);
```

Lorsque le tableau d'entrée est vide, la fonction `getRandomLunch` doit afficher la chaîne `"No lunches available."` dans la console.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  getRandomLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches available.");
} finally {
  console.log = temp;
}
```

Lorsque le tableau d'entrée n'est pas vide, la fonction `getRandomLunch` doit afficher dans la console une chaîne au format `Randomly selected lunch: [Lunch Item]`.

```js
const testLunches = ["Fish", "Fries", "Roast"];
const tempRandom = Math.random;
const tempArr = [];
const temp = console.log;

try {
  console.log = obj => tempArr.push(obj);
  // check that it correctly outputs the first item
  Math.random = () => 0;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[0], `Randomly selected lunch: ${testLunches[0]}`);

  // Second item
  Math.random = () => 0.4;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[1], `Randomly selected lunch: ${testLunches[1]}`);

  // Third item
  Math.random = () => 0.8;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[2], `Randomly selected lunch: ${testLunches[2]}`);

} finally {
  Math.random = tempRandom;
  console.log = temp;
}
```

La fonction `getRandomLunch` ne doit pas modifier le tableau qui lui est passé en argument.

```js
const testLunches = ["Fish", "Fries", "Roast"];
let originalLunches = [...testLunches];
const tempRandom = Math.random;

try {
  Math.random = () => 0;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  Math.random = () => 0.4;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  Math.random = () => 0.8;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  const emptyLunches = [];
  let originalEmpty = [...emptyLunches];
  getRandomLunch(emptyLunches);
  assert.sameOrderedMembers(emptyLunches, originalEmpty);

} finally {
  Math.random = tempRandom;
}
```

Tu dois définir une fonction `showLunchMenu`.

```js
assert.isFunction(showLunchMenu);
```

La fonction `showLunchMenu` doit avoir un seul paramètre.

```js
assert.lengthOf(showLunchMenu, 1);
```

Lorsque le tableau d'entrée est vide, la fonction `showLunchMenu` doit afficher la chaîne `"The menu is empty."` dans la console.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], "The menu is empty.");
} finally {
  console.log = temp;
}
```

`showLunchMenu(["Greens", "Corns", "Beans"])` doit afficher `"Menu items: Greens, Corns, Beans"` dans la console.

```js
const testLunches = ["Greens", "Corns", "Beans"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], `Menu items: ${testLunches.join(", ")}`);
} finally {
  console.log = temp;
}
```

`showLunchMenu(["Pizza", "Burger", "Fries", "Salad"])` doit afficher `"Menu items: Pizza, Burger, Fries, Salad"` dans la console.

```js
const testLunches = ["Pizza", "Burger", "Fries", "Salad"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], `Menu items: ${testLunches.join(", ")}`);
} finally {
  console.log = temp;
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const lunches = [];

// Add a new lunch to the end of the list
function addLunchToEnd(lunches, newLunch) {
  lunches.push(newLunch);
  console.log(`${newLunch} added to the end of the lunch menu.`);
  return lunches;
}

// Add a new lunch to the start of the list
function addLunchToStart(lunches, newLunch) {
  lunches.unshift(newLunch);
  console.log(`${newLunch} added to the start of the lunch menu.`);
  return lunches;
}

// Remove the last lunch from the list
function removeLastLunch(lunches) {
  const removedLunch = lunches.pop();
  if (removedLunch) {
    console.log(`${removedLunch} removed from the end of the lunch menu.`);
  } else {
    console.log("No lunches to remove.");
  }
  return lunches;
}

// Remove the first lunch from the list
function removeFirstLunch(lunches) {
  const removedLunch = lunches.shift();
  if (removedLunch) {
    console.log(`${removedLunch} removed from the start of the lunch menu.`);
  } else {
    console.log("No lunches to remove.");
  }
  return lunches;
}

// Function to get a random lunch from the list
function getRandomLunch(lunches) {
  if (lunches.length === 0) {
    console.log("No lunches available.");
  }
  else {
    const randomIndex = Math.floor(Math.random() * lunches.length);
    const randomLunch = lunches[randomIndex];
    console.log(`Randomly selected lunch: ${randomLunch}`);
  }
}

// Function to display all the lunches in the list
function showLunchMenu(lunches) {
  if (lunches.length === 0) {
    console.log("The menu is empty.");
  }
  else {
    console.log(`Menu items: ${lunches.join(', ')}`);
  }
}

addLunchToEnd(lunches, "Tacos");
addLunchToStart(lunches, "Sushi");
removeLastLunch(lunches);
removeFirstLunch(lunches);
getRandomLunch(lunches);
showLunchMenu(lunches);

```
