---
id: 66f686b8ebdb982fa8e14330
title: Construire un programme de conversion de dates
challengeType: 26
dashedName: lab-date-conversion
---

# --description--

Dans ce lab, tu vas coder un programme de conversion de dates qui convertit une date donnée en différents formats. Par exemple, la date actuelle `Fri Sep 27 2024 16:04:43 GMT+0500 (Pakistan Standard Time)` serait formatée de ces 2 façons :

- `(MM/DD/YYYY): 9/27/2024`.
- `(Month Day, Year): September 27, 2024`.

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu dois créer une variable nommée `currentDate` et lui assigner la date et l'heure actuelles en utilisant l'objet `Date`.
2. Tu dois créer une variable nommée `currentDateFormat` et lui assigner la chaîne `Current Date and Time: [current date]`. Remplace `[current date]` par le résultat de la variable `currentDate`.
3. Tu dois afficher la valeur de `currentDateFormat` dans la console.
4. Tu dois créer une fonction nommée `formatDateMMDDYYYY` qui prend l'objet date en paramètre. Tu peux nommer ce paramètre comme tu le souhaites.
5. Ta fonction `formatDateMMDDYYYY` doit convertir la date actuelle au format `month/day/year` et renvoyer la chaîne `Formatted Date (MM/DD/YYYY): [month]/[day]/[year]`.
6. Tu dois créer une fonction nommée `formatDateLong` qui prend l'objet date en paramètre.
7. Ta fonction `formatDateLong` doit convertir la date actuelle au format `Month Day, Year` et renvoyer la chaîne `Formatted Date (Month Day, Year): [formatted date]`.

**Remarque :** pour que les tests passent, assure-toi d'utiliser `en-US` comme locale lors du formatage des dates.

# --hints--

Tu dois avoir une variable nommée `currentDate` qui contient la date et l'heure actuelles en utilisant l'objet `Date`.

```js
// excluded timestamps
const expectedDate = new Date();
const expectedDateOnly = new Date(expectedDate.getFullYear(), expectedDate.getMonth(), expectedDate.getDate());

const currentDateOnly = new Date(currentDate.getFullYear(), currentDate.getMonth(), currentDate.getDate());

assert.equal(currentDateOnly.toString(), expectedDateOnly.toString());
```

Tu dois avoir une variable nommée `currentDateFormat` qui contient la date actuelle au format `Current Date and Time: <ddd> <MMM> <dd> <yyyy> <HH>:<mm>:<ss> <TIMEZONE>`.

```js
const expectedDateString = `Current Date and Time: `;
assert.include(currentDateFormat, expectedDateString);

const currentTimestamp = new Date(currentDateFormat.replace(expectedDateString, '')).getTime();
const expectedTimestamp = new Date().getTime();
assert.approximately(currentTimestamp, expectedTimestamp, 2000);
```

Tu dois afficher la valeur de `currentDateFormat` dans la console.

```js
assert.match(code, /console\.log\(\s*currentDateFormat\s*\)/);
```

Tu dois avoir une fonction `formatDateMMDDYYYY`

```js
assert.isFunction(formatDateMMDDYYYY);
```

La fonction `formatDateMMDDYYYY` doit prendre un seul paramètre.

```js
assert.lengthOf(formatDateMMDDYYYY, 1);
```

Quand l'objet date contient `Fri Sep 27 2024 16:16:11 GMT+0500 (Pakistan Standard Time)`, la fonction `formatDateMMDDYYYY` doit renvoyer `Formatted Date (MM/DD/YYYY): 9/27/2024`.

```js
let dateString = "Fri Sep 27 2024 16:21:18 GMT+0500 (Pakistan Standard Time)";
let dateObject = new Date(dateString);

assert.equal(formatDateMMDDYYYY(dateObject).toString(), "Formatted Date (MM/DD/YYYY): 9/27/2024");

// prevent hardcoding

dateString = "Fri Sep 28 2024 18:31:18 GMT+0500 (Pakistan Standard Time)";
dateObject = new Date(dateString);

assert.strictEqual(formatDateMMDDYYYY(dateObject).toString(), "Formatted Date (MM/DD/YYYY): 9/28/2024");
```

Tu dois avoir une fonction `formatDateLong`

```js
assert.isFunction(formatDateLong);
```

La fonction `formatDateLong` doit prendre un seul paramètre.

```js
assert.lengthOf(formatDateLong, 1);
```

Quand l'objet date contient `Fri Sep 27 2024 16:16:11 GMT+0500 (Pakistan Standard Time)`, la fonction `formatDateLong` doit renvoyer `Formatted Date (Month Day, Year): September 27, 2024`.

```js
let testDate = new Date('Fri Sep 27 2024 16:21:18 GMT+0500 (Pakistan Standard Time)');
let result = formatDateLong(testDate);

assert.strictEqual(result,'Formatted Date (Month Day, Year): September 27, 2024', `Test failed: ${result}`);

// prevent hardcoding

testDate = new Date('Fri Sep 28 2024 18:31:18 GMT+0500 (Pakistan Standard Time)');
result = formatDateLong(testDate);
assert.strictEqual(result, 'Formatted Date (Month Day, Year): September 28, 2024');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const currentDate = new Date();
const currentDateFormat = `Current Date and Time: ${currentDate}`;
console.log(currentDateFormat);

function formatDateMMDDYYYY(date) {
  const month = date.getMonth() + 1;
  const day = date.getDate();
  const year = date.getFullYear();
  return `Formatted Date (MM/DD/YYYY): ${month}/${day}/${year}`;
}

function formatDateLong(date) {
  const longDateFormat =  date.toLocaleDateString('en-US', {
    year: 'numeric',
    month: 'long',
    day: 'numeric'
  });
  return `Formatted Date (Month Day, Year): ${longDateFormat}`;
}

// usage
console.log(formatDateMMDDYYYY(currentDate));
console.log(formatDateLong(currentDate));

```
