---
id: 694a9ff65d4472dd451d8255
title: Étape 1
challengeType: 20
dashedName: step-1
---

# --description--

Dans ce workshop, tu vas t'entraîner à utiliser les booléens et les instructions conditionnelles en Python en construisant un calculateur de réservation de billets de cinéma.

Commence par déclarer des variables qui stockent des informations sur l'utilisateur et la séance.

Crée une variable nommée `base_price` pour stocker le prix de base du billet de cinéma et définis sa valeur à `15`. Crée une autre variable nommée `age` pour stocker l'âge de l'utilisateur et définis sa valeur à `21`.

# --hints--

Tu devrais déclarer une variable nommée `base_price`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("base_price")`)) })
```

Tu devrais assigner la valeur entière `15` à `base_price`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("base_price").is_equivalent("base_price = 15")
`)) })
```

Tu devrais déclarer une variable nommée `age`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("age")`)) })
```

Tu devrais assigner la valeur entière `21` à `age`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("age").is_equivalent("age = 21")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
