---
id: 694cc1b515976889afd4cded
title: Étape 2
challengeType: 20
dashedName: step-2
---

# --description--

Maintenant, tu dois stocker quelques valeurs de type chaîne concernant le billet de cinéma.

Crée une variable nommée `seat_type` pour stocker le type de siège choisi par l'utilisateur et définis sa valeur à la chaîne `Gold`. Crée une autre variable nommée `show_time` pour stocker l'horaire de la séance et définis sa valeur à la chaîne `Evening`.

# --hints--

Tu devrais déclarer une variable nommée `seat_type`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("seat_type")`)) })
```

Tu devrais assigner la chaîne `Gold` à `seat_type`. N'oublie pas d'entourer la valeur de quotes simples ou doubles.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("seat_type").is_equivalent("seat_type = 'Gold'")
`)) })
```

Tu devrais déclarer une variable nommée `show_time`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("show_time")`)) })
```

Tu devrais assigner la chaîne `Evening` à `show_time`. N'oublie pas d'entourer la valeur de quotes simples ou doubles.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("show_time").is_equivalent("show_time = 'Evening'")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
--fcc-editable-region--

--fcc-editable-region--
```
