---
id: 6958e0ecc121b0ae3dd9c5e0
title: Étape 9
challengeType: 20
dashedName: step-9
---

# --description--

Tu dois aussi gérer le cas où l'utilisateur n'a pas droit à une réduction membre.

Ajoute une clause `else` à ton instruction `if` et affiche `User does not qualify for membership discount` dans le corps du `else`.

Tu peux afficher un message suivi d'une variable de cette façon :

```py
print('Total:', total)
```

Tu veux aussi afficher la valeur mise à jour de `discount`. Sous l'instruction `if...else`, utilise l'appel `print()` pour afficher un message qui montre `Discount:` suivi de la valeur mise à jour de `discount`. Puis vérifie la sortie dans le terminal.

# --hints--

Tu devrais avoir une clause `else`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[2].find_conditions()[1].is_empty()`)) })
```

Tu devrais afficher `User does not qualify for membership discount` à l'intérieur de ta clause `else`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[1].has_call("print('User does not qualify for membership discount')")
`)) })
```

Tu devrais avoir `print('Discount:', discount)` sous ton instruction `if...else`.

```js
({ test: () => (runPython(`
if_stmt = """if is_member:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')"""
print_call = "print('Discount:', discount)"
assert _Node(_code).is_ordered(if_stmt, print_call)
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
if is_member:
    discount = 3
    print('User qualifies for membership discount')
--fcc-editable-region--

--fcc-editable-region--
```
