---
id: 6958de0d773289a89515a6b8
title: Étape 4
challengeType: 20
dashedName: step-4
---

# --description--

Maintenant, tu vas vérifier si l'utilisateur a le droit de réserver une séance du soir selon son âge.

Crée une instruction `if` pour vérifier si `age` est supérieur ou égal à `21`. Dans le corps de l'instruction `if`, affiche `User is eligible for Evening shows`.

# --hints--

Tu devrais avoir deux instructions `if` dans ton code.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[1]`)) })
```

La condition de ta deuxième instruction `if` devrait être `age >= 21`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[1].find_conditions()[0]
assert _cond.is_equivalent("age >= 21") or _cond.is_equivalent("21 <= age")
`)) })
```

Tu devrais afficher `User is eligible for Evening shows` à l'intérieur de ta nouvelle instruction `if`. N'oublie pas d'entourer le message de quotes simples ou doubles.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[1].find_bodies()[0].has_call("print('User is eligible for Evening shows')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

--fcc-editable-region--

--fcc-editable-region--
```
