---
id: 6958dd85587e10c99162dd81
title: Étape 3
challengeType: 20
dashedName: step-3
---

# --description--

Comme tu l'as appris dans les leçons précédentes, en Python, une instruction `if` permet d'exécuter du code selon qu'une condition est vraie.

Une instruction `if` se compose du mot-clé `if`, suivi d'une condition et d'un deux-points. Le code à exécuter quand la condition est vraie, qui doit être indenté, s'appelle le corps de l'instruction `if`.

```py
if condition:
    # Code to execute if condition is True
```

Dans cette étape, tu vas vérifier si l'utilisateur est éligible à réserver un billet de cinéma selon son âge.

Crée une instruction `if` pour vérifier si `age` est supérieur à `17`. Dans le corps de l'instruction `if`, affiche `User is eligible to book a ticket`. Cela affichera le message uniquement quand l'âge de l'utilisateur est supérieur à `17`.

N'oublie pas d'indenter le corps de l'instruction `if` et d'entourer le message de quotes simples ou doubles dans l'appel `print()`.

# --hints--

Tu devrais avoir une instruction `if`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[0]`)) })
```

La condition de l'instruction `if` devrait être `age > 17`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[0].find_conditions()[0]
assert _cond.is_equivalent("age > 17") or _cond.is_equivalent("17 < age")
`)) })
```

Tu devrais avoir un appel `print()` à l'intérieur de ton instruction `if`. N'oublie pas d'indenter le corps de l'instruction `if`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[0].find_bodies()[0].find_calls("print")[0]
`)) })
```

Tu devrais afficher `User is eligible to book a ticket` à l'intérieur de ton instruction `if`. N'oublie pas d'entourer le message de quotes simples ou doubles.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[0].find_bodies()[0].has_call("print('User is eligible to book a ticket')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

--fcc-editable-region--

--fcc-editable-region--
```
