---
id: 6958dfa0c74b237e1b56ce04
title: Étape 7
challengeType: 20
dashedName: step-7
---

# --description--

L'utilisateur a droit à une réduction membre s'il est membre.

Crée une variable nommée `discount` et définis sa valeur à `0`. Elle stockera la réduction dont l'utilisateur bénéficie sur le billet de cinéma.

# --hints--

Tu devrais déclarer une variable nommée `discount`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("discount")`)) })
```

Tu devrais assigner la valeur entière `0` à `discount`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("discount").is_equivalent("discount = 0")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

--fcc-editable-region--

--fcc-editable-region--
```
