---
id: 6958e59845777391b86d7331
title: Étape 18
challengeType: 20
dashedName: step-18
---

# --description--

En Python, une instruction `if` peut aussi être placée dans le corps d'une autre instruction `if`. C'est ce qu'on appelle une instruction `if` imbriquée.

Une instruction `if` imbriquée te permet de vérifier une condition supplémentaire seulement après que la première condition a déjà été satisfaite. L'instruction `if` intérieure ne s'exécutera que si la condition du `if` extérieur est vraie.

```py
if condition1:
    # Code to execute if condition1 is True
    if condition2:
        # Code to execute if both conditions are True
```

Maintenant, tu vas calculer les frais de service selon le type de siège choisi par l'utilisateur.

Dans le corps de la dernière instruction `if`, sous la ligne `print('Ticket booking condition satisfied')`, crée une variable nommée `service_charges` et définis-la à `0`. Assure-toi d'indenter ton code de quatre espaces pour le garder dans le corps de l'instruction `if` extérieure.

Puis, crée une instruction `if` imbriquée pour vérifier si `seat_type` est égal à `Premium`. Dans le corps de l'instruction `if` imbriquée, mets à jour la valeur de `service_charges` à `5`.

# --hints--

Tu devrais déclarer une variable nommée `service_charges` dans le corps de ta dernière instruction `if`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[0].has_variable("service_charges")
`)) })
```

Tu devrais assigner la valeur entière `0` à `service_charges`.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[4].find_bodies()[0].find_variable("service_charges")
assert _var.is_equivalent("service_charges = 0")
`)) })
```

Tu devrais avoir une instruction `if` dans le corps de ta dernière instruction `if`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4].find_ifs()[0]`)) })
```

La condition de la nouvelle instruction `if` devrait être `seat_type == 'Premium'`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_ifs()[0].find_conditions()[0]
assert _cond.is_equivalent("seat_type == 'Premium'")
`)) })
```

Tu devrais assigner la valeur entière `5` à la variable `service_charges` à l'intérieur de ta nouvelle instruction `if`.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[4].find_ifs()[0].find_bodies()[0].find_variable("service_charges")
assert _var.is_equivalent("service_charges = 5")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

if age >= 21 or age >= 18 and (show_time != 'Evening' or is_member):
    print('Ticket booking condition satisfied')

--fcc-editable-region--

--fcc-editable-region--
else:
    print('Ticket booking failed due to restrictions')
```
