---
id: 6958e2377cacbc46a7f7383a
title: Étape 13
challengeType: 20
dashedName: step-13
---

# --description--

Ajoute maintenant une clause `else` à ton instruction `if` et affiche `No extra charges will be applied` dans le corps du `else`. Cela affichera le message uniquement quand la condition des frais supplémentaires est fausse.

Puis, sous la clause `else`, utilise l'appel `print()` pour afficher un message qui montre `Extra charges:` suivi de la valeur mise à jour de `extra_charges`, et vérifie la sortie dans le terminal.

# --hints--

Tu devrais avoir une clause `else`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[3].find_conditions()[1].is_empty()`)) })
```

Tu devrais afficher `No extra charges will be applied` à l'intérieur de ta clause `else`. N'oublie pas d'entourer le message de quotes simples ou doubles.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_bodies()[1].has_call("print('No extra charges will be applied')")
`)) })
```

Tu devrais avoir `print('Extra charges:', extra_charges)` sous ton instruction `if...else`.

```js
({ test: () => (runPython(`
if_stmt = """if is_weekend:
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')"""
print_call = "print('Extra charges:', extra_charges)"
assert _Node(_code).is_ordered(if_stmt, print_call)
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend:
    extra_charges = 2
    print('Extra charges will be applied')
--fcc-editable-region--

--fcc-editable-region--
```
