---
id: 6958e1875e592b58e8852757
title: Étape 12
challengeType: 20
dashedName: step-12
---

# --description--

Crée maintenant une variable nommée `extra_charges` et définis-la à `0`. Elle représentera les frais supplémentaires à appliquer au billet de cinéma le week-end.

Puis, crée une instruction `if` pour vérifier si `is_weekend` est truthy. Dans le corps de l'instruction `if`, mets à jour la valeur de `extra_charges` à `2` et affiche `Extra charges will be applied` dans le terminal.

# --hints--

Tu devrais déclarer une variable nommée `extra_charges`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("extra_charges")`)) })
```

Tu devrais assigner la valeur entière `0` à `extra_charges`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("extra_charges").is_equivalent("extra_charges = 0")
`)) })
```

Tu devrais avoir une quatrième instruction `if`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[3]`)) })
```

La condition de ta quatrième instruction `if` devrait être `is_weekend`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_conditions()[0].is_equivalent("is_weekend")
`)) })
```

Tu devrais assigner la valeur entière `2` à la variable `extra_charges` à l'intérieur de ta nouvelle instruction `if`.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[3].find_bodies()[0].find_variable("extra_charges")
assert _var.is_equivalent("extra_charges = 2")
`)) })
```

Tu devrais afficher `Extra charges will be applied` à l'intérieur de ta nouvelle instruction `if`. N'oublie pas d'entourer le message de quotes simples ou doubles.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_bodies()[0].has_call("print('Extra charges will be applied')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

--fcc-editable-region--

--fcc-editable-region--
```
