---
id: 6960ba6e6c95ed4ad896c6f2
title: Étape 17
challengeType: 20
dashedName: step-17
---

# --description--

Il y a encore une exception aux règles de réservation. Les utilisateurs entre `18` et `21` ans peuvent réserver des séances du soir s'ils sont membres.

Quand plusieurs opérateurs logiques sont utilisés dans une instruction `if`, les conditions liées par `and` sont évaluées avant celles liées par `or`. Les parenthèses `()` servent en Python à grouper les conditions et à contrôler l'ordre d'évaluation.

```py
if condition1 and (condition2 or condition3):
    # Code to execute if conditions evaluate to True
else:
    # Code to execute if conditions evaluate to False
```

Mets à jour la condition de la ligne `if age >= 21 or age >= 18 and show_time != 'Evening':` pour ajouter une autre condition avec l'opérateur `or` qui vérifie si `is_member` est truthy. Utilise des parenthèses `()` pour grouper ensemble les conditions `show_time != 'Evening'` et `is_member`, comme dans l'exemple ci-dessus.

# --hints--

La condition de ta dernière instruction `if` devrait être `age >= 21 or age >= 18 and (show_time != 'Evening' or is_member)`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_conditions()[0]
assert _cond.is_equivalent("age >= 21 or age >= 18 and (show_time != 'Evening' or is_member)")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

--fcc-editable-region--
if age >= 21 or age >= 18 and show_time != 'Evening':
--fcc-editable-region--
    print('Ticket booking condition satisfied')
else:
    print('Ticket booking failed due to restrictions')
```
