---
id: 6960a5165f7efe52a0c0aa39
title: Étape 10
challengeType: 20
dashedName: step-10
---

# --description--

En Python, l'opérateur `and` sert à vérifier si plusieurs conditions sont vraies. Voici comment l'utiliser pour combiner deux conditions dans une instruction `if` :

```py
if condition1 and condition2:
    # Code to execute if all conditions are True
```

La réduction membre ne s'applique aux membres que si leur `age` est supérieur ou égal à `21`.

Mets à jour la condition de la ligne `if is_member:` en utilisant l'opérateur `and` pour combiner la condition existante avec une autre condition qui vérifie si `age` est supérieur ou égal à `21`.

# --hints--

La condition de ta troisième instruction `if` devrait être `is_member and age >= 21`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[2].find_conditions()[0]
assert _cond.is_equivalent("is_member and age >= 21") or _cond.is_equivalent("is_member and 21 <= age")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
--fcc-editable-region--
if is_member:
--fcc-editable-region--
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)
```
