---
id: 696099c1ac72a4e585de1d5b
title: Étape 6
challengeType: 20
dashedName: step-6
---

# --description--

Certaines informations ne peuvent être que vraies ou fausses. Comme tu l'as appris dans les leçons précédentes, on peut les représenter avec des valeurs booléennes.

Crée une variable nommée `is_member` pour indiquer si l'utilisateur est membre, et définis sa valeur à `True`.

Sous la variable `is_member`, crée une autre variable nommée `is_weekend` pour indiquer si la séance a lieu le week-end, et définis sa valeur à `False`. N'entoure pas la valeur de quotes.

# --hints--

Tu devrais déclarer une variable nommée `is_member`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("is_member")`)) })
```

Tu devrais assigner la valeur booléenne `True` à `is_member`. N'entoure pas la valeur de quotes.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("is_member").is_equivalent("is_member = True")
`)) })
```

Tu devrais déclarer une variable nommée `is_weekend`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("is_weekend")`)) })
```

Tu devrais assigner la valeur booléenne `False` à `is_weekend`. N'entoure pas la valeur de quotes.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("is_weekend").is_equivalent("is_weekend = False")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

--fcc-editable-region--

--fcc-editable-region--
```
