---
id: 696099abb6acc539b27055a1
title: Étape 5
challengeType: 20
dashedName: step-5
---

# --description--

À l'étape précédente, tu as vérifié une condition avec une instruction `if`. La clause `else` sert à traiter le cas où la condition est fausse. Voici la syntaxe d'une instruction `if...else` :

```py
if condition:
    # Code to execute if condition is True
else:
    # Code to execute if condition is False
```

Ajoute maintenant une clause `else` à ton instruction `if` et affiche `User is not eligible for Evening shows` dans le corps du `else`. Cela s'affichera uniquement quand la condition de l'instruction `if` est fausse.

# --hints--

Tu devrais avoir une clause `else`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[1].find_conditions()[1].is_empty()`)) })
```

Tu devrais afficher `User is not eligible for Evening shows` à l'intérieur de ta clause `else`. N'oublie pas d'entourer le message de quotes simples ou doubles.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[1].find_bodies()[1].has_call("print('User is not eligible for Evening shows')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
--fcc-editable-region--

--fcc-editable-region--
```
