---
id: 6960a848aacf859f7bb15447
title: Étape 11
challengeType: 20
dashedName: step-11
---

# --description--

Change maintenant la valeur de la variable `is_member` en `False`, car l'utilisateur n'est pas membre.

Ensuite, tu verras que la valeur de `discount` reste `0`, car les deux conditions doivent être satisfaites pour que la réduction s'applique.

# --hints--

Tu devrais mettre à jour la valeur de la variable `is_member` à `False`. N'entoure pas la valeur de quotes.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("is_member").is_equivalent("is_member = False")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

--fcc-editable-region--
is_member = True
--fcc-editable-region--
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)
```
