---
id: 696cb889fc4fdeb337e52864
title: Étape 8
challengeType: 20
dashedName: step-8
---

# --description--

En Python, les conditions d'un `if` n'ont pas besoin de comparer explicitement les valeurs à `True` ou `False`. Elles s'appuient plutôt sur le caractère truthy des valeurs. Les valeurs truthy s'évaluent à `True` dans un contexte booléen, comme la condition d'une instruction `if`. Cela inclut les nombres non nuls, les chaînes non vides et la valeur booléenne `True`.

```py
if value:
   print('value is truthy')
```

Crée une instruction `if` pour vérifier si `is_member` est truthy. Dans le corps de l'instruction `if`, mets à jour la valeur de `discount` à `3` et affiche `User qualifies for membership discount` dans le terminal.

# --hints--

Tu devrais avoir trois instructions `if` dans ton code.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[2]`)) })
```

La condition de ta troisième instruction `if` devrait être `is_member`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_conditions()[0].is_equivalent("is_member")
`)) })
```

Tu devrais assigner la valeur entière `3` à la variable `discount` à l'intérieur de ta nouvelle instruction `if`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[0].find_variable("discount").is_equivalent("discount = 3")
`)) })
```

Tu devrais afficher `User qualifies for membership discount` à l'intérieur de ta nouvelle instruction `if`. N'oublie pas d'entourer le message de quotes simples ou doubles.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[0].has_call("print('User qualifies for membership discount')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
--fcc-editable-region--

--fcc-editable-region--
```
