---
id: ab306dbdcc907c7ddfc30830
title: Crée un outil d'aplatissement profond
challengeType: 26
dashedName: create-a-deep-flattening-tool
---

# --description--

Dans ce lab, tu vas implémenter un algorithme d'aplatissement de tableau.

Aplatir un tableau signifie transformer un tableau imbriqué de n'importe quelle profondeur en un seul tableau unidimensionnel. Le processus extrait tous les éléments dans l'ordre, en dépliant uniquement les tableaux. Les autres types restent inchangés.

Par exemple :

| Original | Aplati |
| --- | --- |
| `[[1], [], [2, [3]]]` | `[1, 2, 3]` |
| `[1, {"foo": "bar"}, [2]]` | `[1, {"foo": "bar"}, 2]` |
| `["baz", [1, 2], {}]` | `["baz", 1, 2, {}]` |

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois avoir une fonction nommée `steamrollArray`.
1. La fonction `steamrollArray` doit accepter un argument : un tableau imbriqué.
1. La fonction doit aplatir le tableau imbriqué en tenant compte de différents niveaux d'imbrication.
1. Ta solution ne doit pas utiliser les méthodes `Array.prototype.flat()` ou `Array.prototype.flatMap()`.
1. Les variables globales ne doivent pas être utilisées.

# --hints--

Tu dois avoir une fonction `steamrollArray`.

```js
assert.isFunction(steamrollArray);
```

`steamrollArray([[["a"]], [["b"]]])` doit renvoyer `["a", "b"]`.

```js
assert.deepEqual(steamrollArray([[['a']], [['b']]]), ['a', 'b']);
```

`steamrollArray([1, [2], [3, [[4]]]])` doit renvoyer `[1, 2, 3, 4]`.

```js
assert.deepEqual(steamrollArray([1, [2], [3, [[4]]]]), [1, 2, 3, 4]);
```

`steamrollArray([1, [], [3, [[4]]]])` doit renvoyer `[1, 3, 4]`.

```js
assert.deepEqual(steamrollArray([1, [], [3, [[4]]]]), [1, 3, 4]);
```

`steamrollArray([1, {}, [3, [[4]]]])` doit renvoyer `[1, {}, 3, 4]`.

```js
assert.deepEqual(steamrollArray([1, {}, [3, [[4]]]]), [1, {}, 3, 4]);
```

Ta solution ne doit pas utiliser les méthodes `Array.prototype.flat()` ou `Array.prototype.flatMap()`.

```js
const arr = [1,2,3,[1,2,3,[1,2,3]]];

const spyFlat = __helpers.spyOn(Array.prototype, 'flat');
try {
  steamrollArray(arr);
  assert.isEmpty(spyFlat.calls);
} finally {
  spyFlat.restore();
}

const spyFlatMap = __helpers.spyOn(Array.prototype, 'flatMap');
try {
  steamrollArray(arr);
  assert.isEmpty(spyFlatMap.calls);
} finally {
  spyFlatMap.restore();
}
```

Tu ne dois pas utiliser de variables globales.

```js
steamrollArray([1, {}, [3, [[4]]]])
assert.deepEqual(steamrollArray([1, {}, [3, [[4]]]]), [1, {}, 3, 4])
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function steamrollArray(arr) {
  if (!Array.isArray(arr)) {
    return [arr];
  }
  const out = [];
  arr.forEach(function(e) {
    steamrollArray(e).forEach(function(v) {
      out.push(v);
    });
  });
  return out;
}
```
