---
id: a3566b1109230028080c9345
title: Conçois un algorithme de somme de tous les nombres
challengeType: 26
dashedName: lab-sum-all-numbers-algorithm
---

# --description--

Dans ce lab, tu vas devoir concevoir un algorithme qui additionne tous les nombres.

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois avoir une fonction nommée `sumAll` qui accepte un tableau de deux nombres.
1. `sumAll([n, m])` doit renvoyer la somme de `n` et `m` plus la somme de tous les nombres entre eux. Le plus petit nombre ne sera pas toujours en premier. Par exemple, `sumAll([4,1])` doit renvoyer `10`, car la somme de tous les nombres entre `1` et `4` (inclus) est `10`.

# --hints--

Tu dois avoir une fonction nommée `sumAll`.

```js
assert.isFunction(sumAll);
```

`sumAll([1, 4])` doit renvoyer un nombre.

```js
assert.isNumber(sumAll([1, 4]));
```

`sumAll([1, 4])` doit renvoyer `10`.

```js
assert.equal(sumAll([1, 4]), 10);
```

`sumAll([4, 1])` doit renvoyer `10`.

```js
assert.equal(sumAll([4, 1]), 10);
```

`sumAll([5, 10])` doit renvoyer `45`.

```js
assert.equal(sumAll([5, 10]), 45);
```

`sumAll([10, 5])` doit renvoyer `45`.

```js
assert.equal(sumAll([10, 5]), 45);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function sumAll(arr) {
  let sum = 0;

  arr.sort((a, b) => a - b);

  for (var i = arr[0]; i <= arr[1]; i++) {
    sum += i;
  }

  return sum;
}
```
