---
id: 615f38cabc64e3556f98cc1a
title: Étape 8
challengeType: 0
dashedName: step-8
---

# --description--

Les bordures peuvent servir à regrouper et hiérarchiser le contenu.

Crée un sélecteur `.label` et donne-lui une `border` définie sur `2px solid black`.

# --hints--

Tu devrais avoir un sélecteur `.label`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.label'));
```

Ton sélecteur `.label` devrait avoir une propriété `border` définie sur `2px solid black`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.label')?.border === '2px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

--fcc-editable-region--

--fcc-editable-region--
```
