---
id: 615f3b091162165948e1cb33
title: Étape 11
challengeType: 0
dashedName: step-11
---

# --description--

Si tu inspectes ton élément `.label` avec les outils de développement de ton navigateur, tu remarqueras qu'il fait en réalité 288 pixels de large au lieu de 270. C'est parce que, par défaut, le navigateur inclut la bordure et le padding quand il calcule la taille d'un élément.

Pour régler ça, réinitialise le modèle de boîte en créant un sélecteur `*` et en lui donnant une propriété `box-sizing` définie sur `border-box`.

# --hints--

Tu devrais créer un sélecteur `*`.

```js
assert(new __helpers.CSSHelp(document).getStyle('*'));
```

Ton sélecteur `*` devrait avoir une propriété `box-sizing` définie sur `border-box`.

```js
assert(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing === 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}
```
