---
id: 615f3949f58e12577dcefb00
title: Étape 9
challengeType: 0
dashedName: step-9
---

# --description--

Bien utiliser l'espace blanc permet d'attirer l'attention sur les éléments importants de ta page et de guider l'œil de l'utilisateur à travers ton texte.

Donne à ton sélecteur `.label` une propriété `width` définie sur `270px`.

# --hints--

Ton sélecteur `.label` devrait avoir une propriété `width` définie sur `270px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.label')?.width === '270px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

--fcc-editable-region--
.label {
  border: 2px solid black;
}
--fcc-editable-region--
```
