---
id: 615f34948891834dd77655a6
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

Ajoute un deuxième élément `p` avec le texte `Serving size 2/3 cup (55g)`.

# --hints--

Tu devrais avoir un deuxième élément `p`.

```js
assert(document.querySelectorAll('p')?.length === 2);
```

Ton deuxième élément `p` devrait être dans ton élément `body`.

```js
assert(document.querySelectorAll('p')?.[1]?.parentElement?.localName === 'body');
```

Ton deuxième élément `p` devrait venir après ton élément `p` existant.

```js
assert(document.querySelectorAll('p')?.[1]?.previousElementSibling?.localName === 'p');
```

Ton deuxième élément `p` devrait avoir le texte `Serving size 2/3 cup (55g)`.

```js
assert(document.querySelectorAll('p')?.[1]?.innerText === 'Serving size 2/3 cup (55g)');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>

</body>
</html>
--fcc-editable-region--
```

```css

```
