---
id: 615f2abbe7d18d49a1e0e1c8
title: Étape 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Un code HTML de base t'a été fourni.

Crée un élément `h1` dans ton élément `body` et donne-lui le texte `Nutrition Facts`.

# --hints--

Tu devrais ajouter un nouvel élément `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Ton élément `h1` devrait être dans ton élément `body`.

```js
assert.equal(document.querySelector('h1')?.parentElement?.localName, 'body');
```

Ton élément `h1` devrait avoir le texte `Nutrition Facts`.

```js
assert.equal(document.querySelector('h1')?.innerText, 'Nutrition Facts');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>

</body>
</html>
--fcc-editable-region--
```

```css

```
