---
id: 615f3cafd794015aa9547a6d
title: Étape 12
challengeType: 0
dashedName: step-12
---

# --description--

Rappelle-toi que l'utilisation des balises `h1`, `h2` et similaires détermine la structure sémantique de ton HTML. Mais tu peux ajuster le CSS de ces éléments pour contrôler la mise en page visuelle et la hiérarchie.

Crée une règle `h1` et définis la propriété `font-weight` sur `800`. Cela va rendre ton texte `h1` plus gras.

# --hints--

Tu devrais créer un sélecteur `h1`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

Ton sélecteur `h1` devrait avoir une propriété `font-weight` définie sur `800`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.fontWeight === '800');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

--fcc-editable-region--

--fcc-editable-region--
```
