---
id: 615f34ecc1091b4fd5a8a484
title: Étape 4
challengeType: 0
dashedName: step-4
---

# --description--

Dans ton élément `head`, ajoute un élément `link` avec l'attribut `rel` défini sur `stylesheet` et l'attribut `href` défini sur `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

Cela va importer la famille de polices `Open Sans`, avec les valeurs de graisse `400`, `700` et `800`.

Ajoute aussi un élément `link` pour lier ton fichier `styles.css`.

# --hints--

Ton code devrait avoir deux éléments `link`.

```js
assert.strictEqual(document.querySelectorAll('link').length, 2);
```

Tes éléments `link` devraient être des éléments vides, ils ne devraient pas avoir de balise fermante `</link>`.

```js
assert.notMatch(code, /<\/link>/);
```

Tes deux éléments `link` devraient être à l'intérieur de l'élément `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link').length, 2);
```

Tes deux éléments `link` devraient avoir l'attribut `rel` défini sur `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

L'un de tes éléments `link` devrait avoir un attribut `href` défini sur `styles.css`.

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
```

L'un de tes éléments `link` devrait avoir un attribut `href` défini sur `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link?.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Open+Sans:400,700,800'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>

</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
--fcc-editable-region--
```

```css

```
