---
id: 615f2d4150fe0d4cbd0f2628
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

Sous ton élément `h1`, ajoute un élément `p` avec le texte `8 servings per container`.

# --hints--

Tu devrais ajouter un nouvel élément `p`.

```js
assert.exists(document.querySelector('p'));
```

Ton élément `p` devrait être dans ton élément `body`.

```js
assert.equal(document.querySelector('p')?.parentElement?.localName, 'body');
```

Ton élément `p` devrait venir après ton élément `h1`.

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.localName, 'h1');
```

Ton élément `p` devrait avoir le texte `8 servings per container`.

```js
assert.equal(document.querySelector('p')?.innerText, '8 servings per container');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>

</body>
</html>
--fcc-editable-region--
```

```css

```
