---
id: 615f357957e370510f21ea16
title: Étape 5
challengeType: 0
dashedName: step-5
---

# --description--

Crée un sélecteur `body` et donne-lui un `font-family` défini sur `Open Sans` avec `sans-serif` comme valeur de repli.

Rappelle-toi que les polices dont le nom contient des espaces doivent être entourées de guillemets en CSS.

# --hints--

Tu devrais avoir un sélecteur `body`.

```js
assert(new __helpers.CSSHelp(document).getStyle('body'));
```

Ton sélecteur `body` devrait avoir une propriété `font-family` définie sur `"Open Sans", sans-serif`.

```js
assert(new __helpers.CSSHelp(document).getStyle('body')?.fontFamily === '"Open Sans", sans-serif');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
