---
id: 684cff0bf0a4ca0ad7e42990
title: Étape 7
challengeType: 0
dashedName: step-7
---

# --description--

L’étape suivante consiste à arrondir les coins de la bordure. Pour définir la propriété `border-radius` avec Tailwind CSS, tu peux utiliser une classe utilitaire du motif `rounded-size`, où `size` représente le rayon.

Voici un exemple d’application d’un border-radius extra-large à un élément :

```html
<p class="rounded-xl">This is border radius.</p>
```

Dans ta balise ouvrante `div`, applique un border-radius de taille `md`.

# --hints--

Ta balise ouvrante `div` doit avoir la classe utilitaire `rounded-md`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("rounded-md"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
