---
id: 684cc59eed2706fa50e8933c
title: Étape 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Dans cet atelier, tu vas découvrir les bases de Tailwind CSS en construisant un composant de message d’erreur.

Pour les premières étapes, tu te concentres sur le balisage HTML initial.

À l’intérieur de l’élément `body`, ajoute un élément `div`. Puis, à l’intérieur de ce `div`, ajoute un élément `p`.

# --hints--

Tu dois avoir un élément `div`.

```js
assert.exists(document.querySelector("div"));
```

Tu dois avoir un élément `p`.

```js
assert.exists(document.querySelector("p"));
```

Ton élément `p` doit être à l’intérieur de l’élément `div`.

```js
assert.exists(document.querySelector("div p"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>
```
