---
id: 684ccbdc1ef89e042b744e09
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

La prochaine partie de l’atelier consiste à ajouter le bouton de fermeture.

Sous ton élément `p`, ajoute un élément `button` avec un `aria-label` égal à `"Dismiss"`. À l’intérieur de ce `button`, ajoute un élément `span` avec un attribut `aria-hidden`. Le texte du `span` doit être `&times;`.

*NOTE* : Comme cet atelier se concentre uniquement sur les styles Tailwind CSS, le `button` n’aura aucun type d’interactivité.

# --hints--

Tu dois avoir un élément `button`.

```js
assert.exists(document.querySelector("button"));
```

Ton élément `button` doit avoir un attribut `aria-label`.

```js
const buttonEl = document.querySelector("button");
assert.isTrue(buttonEl?.hasAttribute("aria-label"));
```

Ton élément `button` doit avoir un attribut `aria-label` égal à `"Dismiss"`.

```js
const buttonEl = document.querySelector("button");
assert.equal(buttonEl?.getAttribute("aria-label"), "Dismiss");
```

Ton élément `button` doit avoir un élément `span` imbriqué à l’intérieur.

```js
assert.exists(document.querySelector("button span"));
```

Ton élément `span` doit avoir un attribut `aria-hidden`.

```js
const spanEl = document.querySelector("span");
assert.isTrue(spanEl?.hasAttribute("aria-hidden"));
```

Ton élément `span` doit avoir le texte `&times;`.

```js
const spanEl = document.querySelector("span");
assert.equal(spanEl?.innerHTML, "×");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      --fcc-editable-region--

      --fcc-editable-region--
    </div>
  </body>
</html>
```
