---
id: 684cca09aac0fc035db4099f
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

À l’intérieur de ton élément `p`, ajoute le texte `Error! Something went wrong. Please try again.` où `Error!` doit être enveloppé dans un élément `strong`.

# --hints--

Tu dois avoir un élément `strong` à l’intérieur de ton élément `p`.

```js
assert.exists(document.querySelector("p strong"));
```

Ton élément `strong` doit avoir le texte `Error!`.

```js
assert.equal(document.querySelector("strong")?.textContent, "Error!");
```

Ton élément `p` doit avoir le texte correct qui se termine par `Something went wrong. Please try again.`

```js
const remainingText = "Something went wrong. Please try again."
const pElement = document.querySelector("p");
assert.isTrue(pElement?.textContent.trim().endsWith(remainingText));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      --fcc-editable-region--
      <p>

      </p>
      --fcc-editable-region--
    </div>
  </body>
</html>
```
