---
id: 684cd59ec94b5e063bf776f0
title: Étape 5
challengeType: 0
dashedName: step-5
---

# --description--

Ce serait bien d’avoir une bordure autour de ce fond rouge clair.

Pour définir le `border-width` en Tailwind CSS, tu peux utiliser la classe utilitaire `border-number` où `number` représente l’épaisseur de ta bordure.

Voici un exemple d’application d’un `border-width` de `4px` sur tous les côtés :

```html
<p class="border-4">This is how to set the border-width.</p>
```

Dans ta balise ouvrante `div`, ajoute la classe utilitaire correcte pour un `border-width` de `2px` sur tous les côtés.

# --hints--

Ta balise ouvrante `div` doit avoir la classe `border-2`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("border-2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
