---
id: 684cd027744c2a0520ba4d58
title: Étape 4
challengeType: 0
dashedName: step-4
---

# --description--

Tailwind CSS utilise des classes utilitaires appliquées directement au balisage. Ces classes couvrent typographie, couleurs, layouts Flexbox, Grid, et plus encore.

Pour appliquer une couleur de fond en Tailwind CSS, tu peux utiliser le préfixe `bg`, suivi du nom de couleur et d’un nombre. Le nombre représente la nuance de la couleur.

Voici un exemple d’application d’une couleur de fond bleu foncé à un paragraphe.

```html
<p class="bg-blue-900">This has a dark blue background.</p>
```

Dans ta balise ouvrante `div`, ajoute une classe avec la classe utilitaire `bg-red-100` qui représente une couleur de fond rouge clair.

# --hints--

Ta balise ouvrante `div` doit avoir un attribut class.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.hasAttribute("class"));
```

Ton attribut class doit contenir la classe utilitaire `bg-red-100`.

```js
const divEl = document.querySelector("div");
assert.exists(document.querySelector("div.bg-red-100"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div>
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
