---
id: 684d0d369e2b3b0f288d2c20
title: Étape 12
challengeType: 0
dashedName: step-12
---

# --description--

Ce serait bien de centrer horizontalement l’élément `div`. Pour appliquer une `margin` sur deux côtés à la fois, tu peux utiliser les motifs `mx-number` ou `my-number`.

`mx` définit la marge gauche et droite, tandis que `my` définit la marge haut et bas.

Voici un exemple d’application des marges haut et bas à un élément :

```html
<p class="my-8">This is margin.</p>
```

Dans ta balise ouvrante `div`, utilise la classe utilitaire correcte pour définir la marge gauche et droite à `auto`.

# --hints--

Ta balise ouvrante `div` doit avoir la classe utilitaire `mx-auto`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
