---
id: 684d1942a1904a14a3709c6c
title: Étape 20
challengeType: 0
dashedName: step-20
---

# --description--

Le style de ton composant de message d’erreur est presque terminé, mais il reste une chose à corriger. Comme tu as défini des tailles de police différentes pour le paragraphe et le bouton, l’alignement vertical entre ces éléments flex est décalé.

Pour positionner les éléments flex le long de l’axe transversal d’un conteneur avec Tailwind CSS, tu peux utiliser le motif `items-value`. `value` peut être `start`, `end`, `center`, `baseline` et plus encore.

Dans ta balise ouvrante `div`, utilise la classe utilitaire correcte pour définir la propriété `align-items` à `center`.

Et avec ce dernier changement, ton composant de message d’erreur est terminé !

# --hints--

Ta balise ouvrante `div` doit avoir la classe utilitaire `items-center`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("items-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
    --fcc-editable-region--
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss" class="text-red-700 text-2xl">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 mx-auto flex gap-4 justify-center items-center">
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss" class="text-red-700 text-2xl">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
