---
id: 684d0eeb9c81e5109e52fe8b
title: Étape 15
challengeType: 0
dashedName: step-15
---

# --description--

L’élément `div` commence à mieux paraître. Mais maintenant le `button` de fermeture est un peu trop proche du paragraphe.

Pour créer de l’espace entre les éléments flex, tu peux utiliser le motif `gap-number`. `number` représente l’espace fixe entre les lignes et les colonnes dans un layout flex.

Voici un exemple d’application d’un `gap` à un élément :

```html
<p class="flex gap-8">This is gap.</p>
```

Dans ta balise ouvrante `div`, utilise la classe utilitaire correcte pour définir le `gap` à `4`.

# --hints--

Ta balise ouvrante `div` doit avoir la classe utilitaire `gap-4`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("gap-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
