---
id: 684d07cbfa43a20d6a002ed7
title: Étape 10
challengeType: 0
dashedName: step-10
---

# --description--

Les styles de l’élément `div` commencent à se mettre en place. Mais ce serait bien d’avoir un peu d’espace autour du texte.

Pour appliquer du `padding` à un élément avec Tailwind CSS, tu peux utiliser le motif `p-number`, où `number` représente les pixels à appliquer sur les quatre côtés.

Voici un exemple d’application d’un `padding` de `8px` à un élément :

```html
<p class="p-8">This is padding.</p>
```

Dans ta balise ouvrante `div`, utilise la classe utilitaire correcte pour appliquer un `padding` de `4px`.

# --hints--

Ta balise ouvrante `div` doit avoir la classe utilitaire `p-4`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("p-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
