---
id: 684d17a3b350bb13d3c19c0f
title: Étape 19
challengeType: 0
dashedName: step-19
---

# --description--

Il est temps d’appliquer quelques styles à l’élément `button`.

Commence par ajouter un attribut class à ton élément `button`. Ensuite, applique la classe utilitaire correcte pour une couleur de texte rouge avec la nuance `700`.

Puis, dans ta balise ouvrante `button`, utilise la classe utilitaire correcte pour définir le `font-size` à `2xl`, ce qui représente `1.5rem` ou (`24px`).

# --hints--

Ton élément `button` doit avoir un attribut class.

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl?.hasAttribute("class"));
```

Ta balise ouvrante `button` doit avoir la classe utilitaire `text-red-700`.

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl?.classList.contains("text-red-700"));
```

Ta balise ouvrante `button` doit avoir la classe utilitaire `text-2xl`.

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl.classList.contains("text-2xl"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      --fcc-editable-region--
      <button aria-label="Dismiss">
      --fcc-editable-region--
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
