---
id: 684d11ffd15a531225d081e9
title: Étape 17
challengeType: 0
dashedName: step-17
---

# --description--

Maintenant que les styles ont été appliqués à l’élément `div`, il est temps d’en appliquer quelques-uns à l’élément `p`.

Actuellement, la couleur du texte du paragraphe est noire. Mais ce serait bien qu’elle soit d’un rouge foncé.

Dans une étape précédente, tu as appris à travailler avec les couleurs de fond comme `bg-red-100`. Tu peux faire quelque chose de similaire avec la couleur du texte en suivant le motif `text-color-shade`.

Voici un exemple d’application d’une couleur de texte bleue à un élément :

```html
<p class="text-blue-500">This has blue text.</p>
```

Commence par ajouter un attribut class à ton élément `p`. Puis applique la classe utilitaire correcte pour une couleur de texte rouge avec la nuance `700`.

# --hints--

Ton élément `p` doit avoir un attribut class.

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl?.hasAttribute("class"));
```

Ta balise ouvrante `p` doit avoir la classe utilitaire `text-red-700`.

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl?.classList.contains("text-red-700"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      --fcc-editable-region--
      <p>
      --fcc-editable-region--
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
