---
id: 684d105e488dbd1166a717ae
title: Étape 16
challengeType: 0
dashedName: step-16
---

# --description--

Si tu agrandis la fenêtre d’aperçu, tu remarqueras que tous les éléments sont encore alignés à gauche dans le conteneur. Ce serait bien que les éléments flex soient centrés.

Pour positionner les éléments flex le long de l’axe principal d’un conteneur avec Tailwind CSS, tu peux utiliser le motif `justify-value`. `value` peut être `start`, `end`, `center`, `between` et plus encore. Toute valeur valide de `justify-content` en CSS vanilla peut être utilisée ici avec Tailwind CSS.

Dans ta balise ouvrante `div`, utilise la classe utilitaire correcte pour définir la propriété `justify-content` à `center`.

# --hints--

Ta balise ouvrante `div` doit avoir la classe utilitaire `justify-center`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("justify-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
