---
id: 684d09cad4ec900e41db2e74
title: Étape 11
challengeType: 0
dashedName: step-11
---

# --description--

Pour l’instant, l’élément `div` est placé dans le coin supérieur gauche de l’aperçu. Mais ce serait bien d’ajouter un peu d’espace au-dessus de l’élément. Pour appliquer une `margin` à un élément avec Tailwind CSS, tu peux utiliser le motif `m-number`, où `number` représente les pixels fixes à appliquer sur les quatre côtés.

Si tu veux appliquer une marge d’un seul côté, tu peux utiliser l’un des motifs suivants : `mt-number` (margin top), `mb-number` (margin bottom), `mr-number` (margin right) et `ml-number` (margin left).

Voici un exemple d’application de `margin-bottom` à un élément :

```html
<p class="mb-8">This is margin.</p>
```

Dans ta balise ouvrante `div`, utilise la classe utilitaire correcte pour appliquer un `margin-top` de `4px`.

# --hints--

Ta balise ouvrante `div` doit avoir la classe utilitaire `mt-4`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("mt-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
