---
id: 684d03df2a0ddf0c735865f0
title: Étape 9
challengeType: 0
dashedName: step-9
---

# --description--

Tu vois maintenant que la largeur de l’élément fait la moitié de celle du `body`. Cela fonctionne bien sur les grands appareils, mais moins bien sur les petits.

Tailwind CSS suit une approche mobile-first. Donc si tu veux créer des designs responsives pour d’autres appareils, tu peux utiliser les variantes responsives de Tailwind comme `md` et `lg` comme ceci :

```html
<p class="lg:w-2/3">This is a percentage width.</p>
```

L’exemple suivant définit la largeur à `66%` de l’élément parent pour les tailles d’appareil de `1024px` et plus.

Dans ta balise ouvrante `div`, modifie ton `w-1/2` pour qu’il ne s’applique qu’aux tailles d’appareil de `768px` et plus. Tu peux préfixer cette classe avec la variante `md`.

# --hints--

Ta balise ouvrante `div` doit avoir la classe utilitaire `md:w-1/2`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("md:w-1/2"));
```

Tu ne dois plus avoir la classe `w-1/2` dans ta balise ouvrante `div`.

```js
const divEl = document.querySelector("div");
assert.isFalse(divEl.classList.contains("w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md w-1/2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
