---
id: 684d009994a3840b98c9aa3a
title: Étape 8
challengeType: 0
dashedName: step-8
---

# --description--

Pour l’instant, la largeur de l’élément `div` s’étend sur toute la largeur de la page. Pour appliquer des largeurs en Tailwind CSS, tu as plusieurs options, comme le motif `w-number` pour des largeurs fixes ou `w-fraction` pour des pourcentages.

Voici un exemple d’application d’une largeur égale à `66%` de son conteneur parent :

```html
<p class="w-2/3">This is a percentage width.</p>
```

Dans ta balise ouvrante `div`, applique une largeur égale à `50%` de son conteneur parent.

# --hints--

Ta balise ouvrante `div` doit avoir la classe utilitaire `w-1/2`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
