---
id: 68501d32b69f9d7756334ffc
title: Étape 13
challengeType: 0
dashedName: step-13
---

# --description--

Si tu agrandis la fenêtre d’aperçu, tu dois voir l’élément `div` centré horizontalement sur la page. Mais sur les petits appareils, la largeur de l’élément `div` s’étend sur toute la largeur de la page.

Mets à jour la classe `mx-auto` pour qu’elle ne s’applique qu’aux appareils `md` et plus grands.

# --hints--

Ta balise ouvrante `div` doit avoir la classe utilitaire `md:mx-auto`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("md:mx-auto"));
```

Tu ne dois plus avoir la classe `mx-auto` dans ta balise ouvrante `div`.

```js
const divEl = document.querySelector("div");
assert.isFalse(divEl.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 mx-auto">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
