---
id: 66b62d0ad68488dd76228d6c
title: 'Déboguer un changeur de couleur d''arrière-plan aléatoire'
challengeType: 25
dashedName: debug-a-random-background-color-changer
demoType: onClick
---

# --description--

Camperbot apprend JavaScript et a essayé de construire son propre changeur de couleur d'arrière-plan aléatoire. Cependant, il a fait quelques erreurs en chemin.

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour que le lab fonctionne correctement.

**User Stories :**

1. Tu dois corriger les erreurs de syntaxe dans le tableau `darkColorsArr`.
1. Tu dois corriger la `ReferenceError` levée quand l'objet `math` est référencé à l'intérieur de la fonction `getRandomIndex`.
1. Tu dois arrondir `randomIndex` à l'entier inférieur le plus proche.
1. Tu dois mettre à jour la valeur de la variable `body` pour utiliser le nom de méthode correct et corriger la `TypeError` associée.
1. Tu dois corriger `bgHexCodeSpanElement` pour qu'il sélectionne correctement l'élément avec l'`id` `bg-hex-code`.
1. Tu dois modifier la variable `color` à l'intérieur de `changeBackgroundColor` pour qu'elle soit définie à une couleur aléatoire du tableau `darkColorsArr`.
1. Tu dois mettre à jour la variable `btn` pour qu'elle cible l'élément bouton. Ouvre le fichier `index.html` pour voir le nom d'`id` correct de cet élément bouton.
1. Tu dois corriger la ligne qui ajoute un écouteur d'événement pour l'événement `click` à `btn` en passant une référence à la fonction `changeBackgroundColor` à la méthode.
1. Ton arrière-plan doit changer aléatoirement vers l'une des couleurs de `darkColorsArr` à chaque clic sur `#btn`.

# --hints--

Tu dois résoudre les erreurs de syntaxe dans le tableau `darkColorsArr`. Souviens-toi que les éléments d'un tableau doivent être séparés par des virgules.

```js
assert.isArray(darkColorsArr);
```

Tu dois corriger l'erreur de majuscule dans la ligne `math.random()`.

```js
assert.match(getRandomIndex.toString(), /\s*(darkColorsArr\.length\s*\*\s*Math\.random\(\s*\)|Math\.random\(\s*\)\s*\*\s*darkColorsArr\.length)\s*/);
```

Tu dois arrondir `darkColorsArr.length * Math.random()` à l'entier inférieur le plus proche.

```js
assert.match(getRandomIndex.toString(), /Math\.floor\(\s*(darkColorsArr\.length\s*\*\s*Math\.random\(\s*\)|Math\.random\(\s*\)\s*\*\s*darkColorsArr\.length)\s*\)/);
```

Tu dois mettre à jour la valeur de ta variable `body` pour que `queryselector` utilise le camel case.

```js
assert.match(code, /const\s+body\s*=\s*document\.querySelector\s*\(\s*("|')body\1\s*\)/);
```

Tu dois corriger `"bg-hex-code"` pour qu'il référence correctement l'`id` `bg-hex-code`. Souviens-toi que les sélecteurs d'`id` en CSS commencent par un dièse (`#`).

```js
assert.match(bgHexCodeSpanElement.id, /bg-hex-code/);
```

Tu dois appeler la fonction `getRandomIndex` dans la ligne `darkColorsArr[getRandomIndex]`.

```js
assert.match(changeBackgroundColor.toString(), /darkColorsArr\[\s*getRandomIndex\(\s*\)\s*\]/);
```

Tu dois corriger le nom d'`id` de la ligne `"#click-btn"` pour qu'il corresponde au nom d'`id` correct dans le fichier `index.html`.

```js
assert.match(btn?.id, /btn/);
```

Tu ne dois pas appeler la fonction `changeBackgroundColor`. À la place, tu dois passer la référence de la fonction à la méthode `addEventListener`.

```js
assert.match(code, /btn\.addEventListener\s*\(\s*("|')click\1\s*,\s*changeBackgroundColor\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Debug a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
</head>

<body>
    <h1>Random Background Color changer</h1>

    <main>
        <section class="bg-information-container">
            <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
        </section>

        <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
    margin: 0;
    padding: 0;
    box-sizing: border-box;
}

:root {
    --yellow: #f1be32;
    --golden-yellow: #feac32;
    --dark-purple: #110815;
    --light-grey: #efefef;
}

body {
    background-color: var(--dark-purple);
    color: var(--light-grey);
    text-align: center;
}

.bg-information-container {
    margin: 15px 0 25px;
    font-size: 1.2rem;
}

.btn {
    cursor: pointer;
    padding: 10px;
    margin: 10px;
    color: var(--dark-purple);
    background-color: var(--golden-yellow);
    background-image: linear-gradient(#fecc4c, #ffac33);
    border-color: var(--golden-yellow);
    border-width: 3px;
}

.btn:hover {
    background-image: linear-gradient(#ffcc4c, #f89808);
}
```

```js
const darkColorsArr = [
    "#2C3E50",
    "#34495E",
    "#2C2C2C"
    "#616A6B"
    "#4A235A"
    "#2F4F4F",
    "#0E4B5A",
    "#36454F",
    "#2C3E50",
    "#800020",
  ];

  function getRandomIndex() {
    const randomIndex = darkColorsArr.length * math.random();
    return randomIndex;
  }
  console.log(getRandomIndex())

  const body = document.queryselector("body");
  const bgHexCodeSpanElement = document.querySelector("bg-hex-code");
  console.log(bgHexCodeSpanElement)

  function changeBackgroundColor() {
    const color = darkColorsArr[getRandomIndex];

    bgHexCodeSpanElement.innerText = color;
    body.style.backgroundColor = color;
  }
  const btn = document.querySelector("#click-btn");
  console.log(btn)

  btn.addEventListener("click", changeBackgroundColor());
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Debug a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
</head>

<body>
    <h1>Random Background Color changer</h1>

    <main>
        <section class="bg-information-container">
            <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
        </section>

        <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
    margin: 0;
    padding: 0;
    box-sizing: border-box;
}

:root {
    --yellow: #f1be32;
    --golden-yellow: #feac32;
    --dark-purple: #110815;
    --light-grey: #efefef;
}

body {
    background-color: var(--dark-purple);
    color: var(--light-grey);
    text-align: center;
}

.bg-information-container {
    margin: 15px 0 25px;
    font-size: 1.2rem;
}

.btn {
    cursor: pointer;
    padding: 10px;
    margin: 10px;
    color: var(--dark-purple);
    background-color: var(--golden-yellow);
    background-image: linear-gradient(#fecc4c, #ffac33);
    border-color: var(--golden-yellow);
    border-width: 3px;
}

.btn:hover {
    background-image: linear-gradient(#ffcc4c, #f89808);
}
```

```js
const darkColorsArr = [
  "#2C3E50",
  "#34495E",
  "#2C2C2C",
  "#616A6B",
  "#4A235A",
  "#2F4F4F",
  "#0E4B5A",
  "#36454F",
  "#2C3E50",
  "#800020",
];

function getRandomIndex() {
  const randomIndex = Math.floor(darkColorsArr.length * Math.random());
  return randomIndex;
}

const body = document.querySelector("body");
const bgHexCodeSpanElement = document.querySelector("#bg-hex-code");

function changeBackgroundColor() {
  const color = darkColorsArr[getRandomIndex()];

  bgHexCodeSpanElement.innerText = color;
  body.style.backgroundColor = color;
}
const btn = document.querySelector("#btn");

btn.addEventListener("click", changeBackgroundColor);
```
