---
id: a6e40f1041b06c996f7b2406
title: Crée un détecteur de premier élément
challengeType: 26
dashedName: build-a-first-element-finder
---

# --description--

Dans ce lab, tu vas créer une fonction qui parcourt un tableau et renvoie le premier élément qui passe un test (aussi appelé « truth test »).

La fonction parcourrait le tableau et testerait chaque élément à l'aide de la fonction de test fournie. À la fin, elle renverrait le premier élément pour lequel la fonction de test renvoie `true`.

Exemple :

```js
findElement([1, 3, 5, 8], num => num % 2 === 0) // returns 8
findElement([1, 3, 5], num => num % 2 === 0)    // returns undefined
```

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois avoir une fonction nommée `findElement` qui accepte un tableau et une fonction en arguments.
2. La fonction doit renvoyer le premier élément du tableau qui passe un truth test. Cela signifie qu'en appelant la fonction passée `func`, pour un élément `x`, le truth test est réussi si `func(x)` est `true`.
3. Si aucun élément ne passe le test, la fonction doit renvoyer `undefined`.

# --hints--

Tu dois avoir une fonction `findElement`.

```js
assert.isFunction(findElement);
```

`findElement([1, 3, 5, 8, 9, 10], function(num) { return num % 2 === 0; })` doit renvoyer `8`.

```js
assert.strictEqual(
  findElement([1, 3, 5, 8, 9, 10], function (num) {
    return num % 2 === 0;
  }),
  8
);
```

`findElement([1, 3, 5, 9], function(num) { return num % 2 === 0; })` doit renvoyer `undefined`.

```js
assert.strictEqual(
  findElement([1, 3, 5, 9], function (num) {
    return num % 2 === 0;
  }),
  undefined
);
```

`findElement([1, 2, 3, 4], function(num) { return num > 2; })` doit renvoyer `3`.

```js
assert.strictEqual(
  findElement([1, 2, 3, 4], function (num) {
    return num > 2;
  }),
  3
);
```

`findElement(["hello", "world", "javascript"], function(str) { return str.length > 5; })` doit renvoyer `"javascript"`.

```js
assert.strictEqual(
  findElement(["hello", "world", "javascript"], function (str) {
    return str.length > 5;
  }),
  "javascript"
);
```

`findElement(["cat", "dog", "bird"], function(str) { return str.length > 10; })` doit renvoyer `undefined`.

```js
assert.strictEqual(
  findElement(["cat", "dog", "bird"], function (str) {
    return str.length > 10;
  }),
  undefined
);
```

`findElement([2, 4, 6, 8], function(num) { return num % 2 === 0; })` doit renvoyer `2`.

```js
assert.strictEqual(
  findElement([2, 4, 6, 8], function (num) {
    return num % 2 === 0;
  }),
  2
);
```

`findElement([], function(num) { return num > 0; })` doit renvoyer `undefined`.

```js
assert.strictEqual(
  findElement([], function (num) {
    return num > 0;
  }),
  undefined
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function findElement(arr, func) {
  return arr.filter(func)[0];
}

findElement([1, 2, 3, 4], num => num % 2 === 0);
```
