---
id: 698367b9d1c6914a22095aec
title: Implémenter l'algorithme de tri à bulles
challengeType: 26
dashedName: implement-the-bubble-sort-algorithm
---

# --description--

Pour ce lab, tu vas implémenter l'algorithme de tri à bulles. Il commence au début d'un tableau non trié et fait « remonter comme des bulles » les valeurs non triées vers la fin, en parcourant le tableau jusqu'à ce qu'il soit complètement trié. Il le fait en comparant les éléments adjacents et en les échangeant s'ils sont dans le désordre. La méthode continue à boucler sur le tableau jusqu'à ce qu'aucun échange ne se produise, moment auquel le tableau est trié.

Cette méthode nécessite plusieurs passages sur le tableau et, pour les cas moyen et pire, a une complexité temporelle quadratique. Bien que simple, elle est généralement peu pratique dans la plupart des situations.

**Objectif :** Remplir les user stories ci-dessous et faire passer tous les tests pour terminer le lab.

**User Stories :**

1. Écris une fonction `bubbleSort` qui prend un tableau d'entiers en entrée et renvoie un tableau de ces entiers triés du plus petit au plus grand.

# --hints--

Tu dois avoir une fonction nommée `bubbleSort`.

```js
assert.isFunction(bubbleSort);
```

`bubbleSort` doit renvoyer un tableau trié (du plus petit au plus grand).

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
assert.isTrue(
  isSorted(
    bubbleSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`bubbleSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` doit renvoyer un tableau inchangé sauf pour l'ordre.

```js
assert.sameMembers(
  bubbleSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`bubbleSort` ne doit pas utiliser la méthode intégrée `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    bubbleSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function bubbleSort(array) {
  for (let i = 0; i < array.length; i++) {
    let swapped = false;
    for (let j = 1; j < array.length; j++) {
      if (array[j - 1] > array[j]) {
        let temp = array[j-1];
        array[j-1] =  array[j];
        array[j] = temp;
        swapped = true;
      }
    }
    if (swapped === false) {
      break;
    }
  }
  return array;
}
```
