---
id: a6b0bb188d873cb2c8729495
title: Implémente un convertisseur d'entités HTML
challengeType: 26
dashedName: implement-an-html-entity-converter
---

# --description--

Ce lab consiste à convertir les caractères spéciaux d'une chaîne en leurs entités HTML correspondantes.

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois avoir une fonction `convertHTML` qui accepte une chaîne en argument.
1. La fonction `convertHTML` doit renvoyer une nouvelle chaîne en convertissant les caractères spéciaux de la chaîne d'argument en leurs entités HTML correspondantes.

   - `&` doit être converti en `&amp;`.
   - `<` doit être converti en `&lt;`.
   - `>` doit être converti en `&gt;`.
   - `"` doit être converti en `&quot;`.
   - `'` doit être converti en `&apos;`.

# --hints--

Tu dois avoir une fonction `convertHTML`.

```js
assert.isFunction(convertHTML);
```

`convertHTML("Dolce & Gabbana")` doit renvoyer la chaîne `Dolce &amp; Gabbana`.

```js
assert.match(convertHTML('Dolce & Gabbana'), /Dolce &amp; Gabbana/);
```

`convertHTML("Hamburgers < Pizza < Tacos")` doit renvoyer la chaîne `Hamburgers &lt; Pizza &lt; Tacos`.

```js
assert.match(
  convertHTML('Hamburgers < Pizza < Tacos'),
  /Hamburgers &lt; Pizza &lt; Tacos/
);
```

`convertHTML("Sixty > twelve")` doit renvoyer la chaîne `Sixty &gt; twelve`.

```js
assert.match(convertHTML('Sixty > twelve'), /Sixty &gt; twelve/);
```

`convertHTML('Stuff in "quotation marks"')` doit renvoyer la chaîne `Stuff in &quot;quotation marks&quot;`.

```js
assert.match(
  convertHTML('Stuff in "quotation marks"'),
  /Stuff in &quot;quotation marks&quot;/
);
```

`convertHTML("Schindler's List")` doit renvoyer la chaîne `Schindler&apos;s List`.

```js
assert.match(convertHTML("Schindler's List"), /Schindler&apos;s List/);
```

`convertHTML("<>")` doit renvoyer la chaîne `&lt;&gt;`.

```js
assert.match(convertHTML('<>'), /&lt;&gt;/);
```

`convertHTML("abc")` doit renvoyer la chaîne `abc`.

```js
assert.strictEqual(convertHTML('abc'), 'abc');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const map = {
  '&': '&amp;',
  '<': '&lt;',
  '>': '&gt;',
  '"': '&quot;',
  "'": '&apos;'
};

function convertHTML(str) {
  return str.replace(/[&<>"']/g, function(char) {
    return map[char];
  });
}
```
