---
id: 68e51d8f163d3f8b70f9c4e5
title: Étape 1
challengeType: 1
dashedName: step-1
---

# --description--

Dans ce workshop, tu vas t'entraîner à utiliser les méthodes `includes()` et `slice()`.

Pour commencer, crée une variable nommée `fccSentence` et affecte-lui la chaîne `"freeCodeCamp is a great place to learn web development."`.

# --hints--

Tu devrais déclarer une variable nommée `fccSentence`.

```js
assert.exists(fccSentence);
```

Ta variable `fccSentence` devrait être une chaîne.

```js
assert.isString(fccSentence);
```

Ta variable `fccSentence` devrait recevoir la chaîne `"freeCodeCamp is a great place to learn web development."`.

```js
assert.strictEqual(fccSentence, "freeCodeCamp is a great place to learn web development.");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
