---
id: 68e6172975008fea742a63ce
title: Étape 8
challengeType: 1
dashedName: step-8
---

# --description--

La méthode `slice()` extrait une section d'une chaîne. Elle prend deux arguments : l'indice de début (inclus) et l'indice de fin (non inclus).

Voici un exemple :

```js
const greeting = "Hello World";
const firstWord = greeting.slice(0, 5);
console.log(firstWord); // "Hello"
```

Les indices de chaîne commencent à 0. Dans la chaîne `message`, le caractère à l'indice 11 est `f` (le début de "freeCodeCamp").

Crée une variable nommée `platform`. Puis affecte-lui le résultat de l'utilisation de `message.slice(11, 23)` pour extraire `"freeCodeCamp"`.

# --hints--

Tu devrais déclarer une variable nommée `platform`.

```js
assert.exists(platform);
```

Ta variable `platform` devrait être une chaîne.

```js
assert.isString(platform);
```

Tu devrais utiliser `message.slice(11, 23)`.

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*11\s*,\s*23\s*\)/);
```

Tu devrais affecter le résultat de `message.slice(11, 23)` à ta variable `platform`.

```js
assert.match(__helpers.removeJSComments(code), /platform\s*=\s*message\.slice\(\s*11\s*,\s*23\s*\)/);
```

Ta variable `platform` devrait être égale à `"freeCodeCamp"`.

```js
assert.strictEqual(platform, "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

--fcc-editable-region--

--fcc-editable-region--
```
