---
id: 68e6172975008fea742a63cb
title: Étape 5
challengeType: 1
dashedName: step-5
---

# --description--

Il est maintenant temps de voir ce qui se passe quand la sous-chaîne n'est PAS trouvée.

Crée une variable nommée `hasJavaScript`. Puis affecte-lui le résultat de l'utilisation de la méthode `includes()` pour vérifier si `fccSentence` contient `"JavaScript"`.

Ensuite, utilise soit un littéral de modèle, soit la concaténation de chaînes pour afficher le message `fccSentence.includes("JavaScript") returns <hasJavaScript> because the word "JavaScript" is not in the sentence.` dans la console. Remplace `<hasJavaScript>` par la valeur réelle de la variable.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Tu devrais déclarer une variable nommée `hasJavaScript`.

```js
assert.exists(hasJavaScript);
```

Tu devrais utiliser `fccSentence.includes("JavaScript")`.

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])JavaScript\1\s*\)/);
```

Tu devrais affecter le résultat de `fccSentence.includes("JavaScript")` à ta variable `hasJavaScript`.

```js
assert.match(__helpers.removeJSComments(code), /hasJavaScript\s*=\s*fccSentence\.includes\(\s*(['"`])JavaScript\1\s*\)/);
```

Le résultat de `fccSentence.includes("JavaScript")` devrait être `false`.

```js
assert.isFalse(hasJavaScript);
```

Tu devrais afficher le message `fccSentence.includes("JavaScript") returns <hasJavaScript> because the word "JavaScript" is not in the sentence.`, où `<hasJavaScript>` doit être remplacé par la valeur réelle de la variable.

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("JavaScript") returns false because the word "JavaScript" is not in the sentence.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

--fcc-editable-region--

--fcc-editable-region--
```
