---
id: 68e6172975008fea742a63d1
title: Étape 11
challengeType: 1
dashedName: step-11
---

# --description--

La méthode `slice()` peut utiliser des indices négatifs pour compter depuis la fin de la chaîne. `-1` fait référence au dernier caractère.

Voici un exemple :

```js
const text = "JavaScript";
const lastThree = text.slice(-3);
console.log(lastThree); // "ipt"
```

Crée une variable nommée `endPunctuation`. Puis affecte-lui le résultat de l'utilisation de `message.slice(-1)` pour extraire le dernier caractère.

Ensuite, utilise soit un littéral de modèle, soit la concaténation de chaînes pour afficher le message `The ending punctuation mark is a "<endPunctuation>"` dans la console. Remplace `<endPunctuation>` par la valeur réelle de la variable.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Tu devrais déclarer une variable nommée `endPunctuation`.

```js
assert.exists(endPunctuation);
```

Ta variable `endPunctuation` devrait être une chaîne.

```js
assert.isString(endPunctuation);
```

Tu devrais utiliser `message.slice(-1)`.

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*-1\s*\)/);
```

Tu devrais affecter le résultat de `message.slice(-1)` à ta variable `endPunctuation`.

```js
assert.match(__helpers.removeJSComments(code), /endPunctuation\s*=\s*message\.slice\(\s*-1\s*\)/);
```

Ta variable `endPunctuation` devrait être égale à `"!"`.

```js
assert.strictEqual(endPunctuation, "!");
```

Tu devrais afficher le message `The ending punctuation mark is a "<endPunctuation>"`, où `<endPunctuation>` doit être remplacé par la valeur réelle de la variable.

```js
assert.deepInclude(spy.calls, ['The ending punctuation mark is a "!"']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
console.log(`The word "${platform}" was sliced from the message.`);

const greetingWord = message.slice(0, 7);
console.log(`The first word is "${greetingWord}".`);

--fcc-editable-region--

--fcc-editable-region--
```
