---
id: 68e6172975008fea742a63cc
title: Étape 6
challengeType: 1
dashedName: step-6
---

# --description--

La méthode `includes()` est sensible à la casse, ce qui signifie qu'elle distingue les lettres majuscules des minuscules.

Crée une variable nommée `hasLowercaseFCC`. Puis affecte-lui le résultat de l'utilisation de la méthode `includes()` pour vérifier si `fccSentence` contient `"freecodecamp"` (tout en minuscules).

Ensuite, utilise soit un littéral de modèle, soit la concaténation de chaînes pour afficher le message `fccSentence.includes("freecodecamp") returns <hasLowercaseFCC> because includes is case-sensitive.` dans la console. Remplace `<hasLowercaseFCC>` par la valeur réelle de la variable.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Tu devrais déclarer une variable nommée `hasLowercaseFCC`.

```js
assert.exists(hasLowercaseFCC);
```

Tu devrais utiliser `fccSentence.includes("freecodecamp")`.

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])freecodecamp\1\s*\)/);
```

Tu devrais affecter le résultat de `fccSentence.includes("freecodecamp")` à ta variable `hasLowercaseFCC`.

```js
assert.match(__helpers.removeJSComments(code), /hasLowercaseFCC\s*=\s*fccSentence\.includes\(\s*(['"`])freecodecamp\1\s*\)/);
```

Le résultat de `fccSentence.includes("freecodecamp")` devrait être `false`.

```js
assert.isFalse(hasLowercaseFCC);
```

Tu devrais afficher le message qui contient `fccSentence.includes("freecodecamp") returns <hasLowercaseFCC> because includes is case-sensitive.`, où `<hasLowercaseFCC>` doit être remplacé par la valeur réelle de la variable.

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("freecodecamp") returns false because includes is case-sensitive.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

--fcc-editable-region--

--fcc-editable-region--
```
