---
id: 68e6172975008fea742a63ca
title: Étape 4
challengeType: 1
dashedName: step-4
---

# --description--

Maintenant, utilise soit un littéral de modèle, soit la concaténation de chaînes pour afficher le message `fccSentence.includes("freeCodeCamp") returns <hasFreeCodeCamp> because the word "freeCodeCamp" is in the sentence.` dans la console. Remplace `<hasFreeCodeCamp>` par la valeur réelle de la variable.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Tu devrais afficher le message `fccSentence.includes("freeCodeCamp") returns <hasFreeCodeCamp> because the word "freeCodeCamp" is in the sentence.`, où `<hasFreeCodeCamp>` doit être remplacé par la valeur réelle de la variable.

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("freeCodeCamp") returns true because the word "freeCodeCamp" is in the sentence.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
--fcc-editable-region--

--fcc-editable-region--
```
