---
id: 68e6172975008fea742a63d0
title: Étape 10
challengeType: 1
dashedName: step-10
---

# --description--

Tu peux utiliser `slice()` pour extraire depuis le début d'une chaîne.

Crée une variable nommée `greetingWord`. Puis affecte-lui le résultat de l'utilisation de `message.slice(0, 7)` pour extraire le premier mot `"Welcome"`.

Ensuite, utilise soit un littéral de modèle, soit la concaténation de chaînes pour afficher le message `The first word is "<greetingWord>".` dans la console. Remplace `<greetingWord>` par la valeur réelle de la variable.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Tu devrais déclarer une variable nommée `greetingWord`.

```js
assert.exists(greetingWord);
```

Ta variable `greetingWord` devrait être une chaîne.

```js
assert.isString(greetingWord);
```

Tu devrais utiliser `message.slice(0, 7)`.

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*0\s*,\s*7\s*\)/);
```

Tu devrais affecter le résultat de `message.slice(0, 7)` à ta variable `greetingWord`.

```js
assert.match(__helpers.removeJSComments(code), /greetingWord\s*=\s*message\.slice\(\s*0\s*,\s*7\s*\)/);
```

Ta variable `greetingWord` devrait être égale à `"Welcome"`.

```js
assert.strictEqual(greetingWord, "Welcome");
```

Tu devrais afficher le message `The first word is "<greetingWord>".`, où `<greetingWord>` doit être remplacé par la valeur réelle de la variable.

```js
assert.deepInclude(spy.calls, ['The first word is "Welcome".']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
console.log(`The word "${platform}" was sliced from the message.`);

--fcc-editable-region--

--fcc-editable-region--
```
