---
id: 68e6172975008fea742a63cd
title: Étape 7
challengeType: 1
dashedName: step-7
---

# --description--

Maintenant, tu vas explorer la méthode `slice()`, qui extrait une portion d'une chaîne.

Crée une variable nommée `message` et affecte-lui la chaîne `"Welcome to freeCodeCamp!"`.

Puis utilise `console.log()` pour afficher `"Here are some examples of the slice() method:"` dans la console.

# --hints--

Tu devrais déclarer une variable nommée `message`.

```js
assert.exists(message);
```

Ta variable `message` devrait être une chaîne.

```js
assert.isString(message);
```

Ta variable `message` devrait être `"Welcome to freeCodeCamp!"`.

```js
assert.strictEqual(message, "Welcome to freeCodeCamp!");
```

Tu devrais afficher `"Here are some examples of the slice() method:"` dans la console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*('|"|`)Here are some examples of the slice\(\) method:\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

--fcc-editable-region--

--fcc-editable-region--
```
