---
id: 68e6172875008fea742a63c9
title: Étape 3
challengeType: 1
dashedName: step-3
---

# --description--

Comme tu t'en souviens des leçons précédentes, la méthode `includes()` vérifie si une chaîne contient une sous-chaîne précise et renvoie `true` ou `false`.

Voici un exemple utilisant la méthode `includes()` :

```js
const text = "Hello World";
const hasHello = text.includes("Hello");
console.log(hasHello); // true
```

Crée une variable nommée `hasFreeCodeCamp`. Puis affecte-lui le résultat de l'utilisation de la méthode `includes()` sur `fccSentence` pour vérifier si elle contient `"freeCodeCamp"`.

# --hints--

Tu devrais déclarer une variable nommée `hasFreeCodeCamp`.

```js
assert.exists(hasFreeCodeCamp);
```

Tu devrais utiliser `fccSentence.includes("freeCodeCamp")`.

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

Tu devrais affecter le résultat de `fccSentence.includes("freeCodeCamp")` à ta variable `hasFreeCodeCamp`.

```js
assert.match(__helpers.removeJSComments(code), /hasFreeCodeCamp\s*=\s*fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

Le résultat de `fccSentence.includes("freeCodeCamp")` devrait être `true`.

```js
assert.isTrue(hasFreeCodeCamp);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

--fcc-editable-region--

--fcc-editable-region--
```
