---
id: 66a9b48a211a73155621d0df
title: Étape 1
challengeType: 0
dashedName: step-1
---

# --description--

Dans cet atelier, tu vas t'entraîner à travailler avec les tableaux HTML en construisant un tableau de notes d'examen final pour un groupe d'élèves.

Dans les leçons précédentes, tu as appris à travailler avec l'élément `table` pour représenter des données tabulaires.

Dans ton élément `body`, imbrique un élément `table`.

# --hints--

Tu dois avoir une balise ouvrante `table`.

```js
assert.match(code, /<table>/i);
```

Tu dois avoir une balise fermante `table`.

```js
assert.match(code, /<\/table>/i);
```

Ton élément `table` doit être dans ton élément `body`.

```js
const tableEl = document.querySelector('table');
assert.strictEqual(tableEl?.parentElement?.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
