---
id: 66a9bb8578459a175432e7d0
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

Pour la première section du tableau, tu dois regrouper le contenu d'en-tête qui représente les libellés de colonnes pour le prénom, le nom et la note finale de l'élève.

L'élément <dfn>table head</dfn>, `thead`, sert à regrouper le contenu d'en-tête dans un tableau.

Voici un exemple utilisant l'élément `thead` :

```html
<table>
  <thead>
    <!-- header content goes here -->
  </thead>
</table>
```

Sous ton élément `caption`, ajoute un élément d'en-tête de tableau.

# --hints--

Tu dois avoir une balise ouvrante `thead`.

```js
assert.match(code, /<thead>/i);
```

Tu dois avoir une balise fermante `thead`.

```js
assert.match(code, /<\/thead>/i);
```

Ton élément `thead` doit être dans ton élément `table`.

```js
const theadEl = document.querySelector('thead');
assert.strictEqual(theadEl?.parentElement?.tagName, 'TABLE');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
