---
id: 66a9b8f14b963916a3baa732
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

Pour ajouter une légende à un `table`, tu peux utiliser l'élément <dfn>table caption</dfn>.

Voici un exemple utilisant l'élément `caption` :

```html
<table>
  <caption>Football Scores</caption>
</table>
```

Dans ton élément `table`, imbrique un élément `caption` avec le texte `Calculus Final Exam Grades`.

# --hints--

Tu dois avoir une balise ouvrante `caption`.

```js
assert.match(code, /<caption>/i);
```

Tu dois avoir une balise fermante `caption`.

```js
assert.match(code, /<\/caption>/i);
```

Ton élément `caption` doit être dans ton élément `table`.

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.parentElement.tagName, 'TABLE');
```

Ton `caption` doit avoir le texte `Calculus Final Exam Grades`.

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.innerText.trim(), 'Calculus Final Exam Grades');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
    --fcc-editable-region--
      
    --fcc-editable-region--      
    </table>
  </body>
</html>
```
