---
id: 66a9c14e3b34c719e34bae20
title: Étape 6
challengeType: 0
dashedName: step-6
---

# --description--

Pour ajouter des données d’élèves au tableau, tu dois utiliser les éléments de ligne de tableau et de données de tableau.

L'élément <dfn>table data</dfn>, `td`, sert à créer une cellule dans le tableau.

Voici un exemple d'utilisation de l'élément `td` pour un tableau de joueurs de sport :

```html
<tr>
  <td>1</td>
  <td>John Doe</td>
  <td>USA</td>
</tr>
```

Dans ton élément de corps de tableau, ajoute un élément de ligne de tableau, `tr`, avec trois éléments de données de tableau, `td`.

Le premier élément de données de tableau doit contenir le nom `Davis`.

Le deuxième élément de données de tableau doit contenir le prénom `Alex`.

Le troisième élément de données de tableau doit contenir la note `54`.

# --hints--

Tu dois avoir un élément `tr` dans ton élément `tbody`.

```js
assert.isNotNull(document.querySelector('tbody tr'));
```

Tu dois avoir un élément `td` avec le texte `Davis` dans l’élément `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr td')?.textContent.trim(), 'Davis');
```

Tu dois avoir un élément `td` avec le texte `Alex` dans l’élément `tr`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[1]?.textContent.trim(), 'Alex');
```

Tu dois avoir un élément `td` avec le texte `54` dans l’élément `tr`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[2]?.textContent.trim(), '54');
```

Tu dois avoir les trois éléments `td` dans l’élément `tr`.

```js
assert.match(code, /<tr>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<\/tr>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
