---
id: 66a9ca099e3ecb1be10a2696
title: Étape 8
challengeType: 0
dashedName: step-8
---

# --description--

Il est maintenant temps d’ajouter deux élèves de plus au tableau.

En suivant le même modèle que les étapes précédentes, ajoute une quatrième ligne d’élève au tableau. Utilise les données suivantes pour les éléments de données du tableau :

- Nom : `Thompson`
- Prénom : `Jane`
- Note : `77`

Pour la cinquième ligne d’élève, utilise les données suivantes pour les éléments de données du tableau :

- Nom : `Williams`
- Prénom : `Natalie`
- Note : `83`

# --hints--

Tu dois avoir un quatrième `tr` dans ton élément `tbody`.

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(4)'));
```

Tu dois avoir un élément `td` avec le texte `Thompson` dans le quatrième élément `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:first-child')?.textContent.trim(), 'Thompson');
```

Tu dois avoir un élément `td` avec le texte `Jane` dans le quatrième élément `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(2)')?.textContent.trim(), 'Jane');
```

Tu dois avoir un élément `td` avec le texte `77` dans le quatrième élément `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(3)')?.textContent.trim(), '77');
```

Tu dois avoir un cinquième `tr` dans ton élément `tbody`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 5);
```

Tu dois avoir un élément `td` avec le texte `Williams` dans le cinquième élément `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:first-child')?.textContent.trim(), 'Williams');
```

Tu dois avoir un élément `td` avec le texte `Natalie` dans le cinquième élément `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(2)')?.textContent.trim(), 'Natalie');
```

Tu dois avoir un élément `td` avec le texte `83` dans le cinquième élément `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(3)')?.textContent.trim(), '83');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>
        
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
