---
id: 66a9c006e4f58819396605b3
title: Étape 5
challengeType: 0
dashedName: step-5
---

# --description--

Maintenant que tu as terminé la section d'en-tête, il est temps d'ajouter le <dfn>table body</dfn>, `tbody`. Le corps du tableau représentera tous les noms des élèves et leurs notes.

Ajoute un élément de corps de tableau à ton tableau.

# --hints--

Tu dois avoir une balise ouvrante `tbody`.

```js
assert.match(code, /<tbody>/i);
```

Tu dois avoir une balise fermante `tbody`.

```js
assert.match(code, /<\/tbody>/i);
```

Ta balise ouvrante `tbody` doit venir avant la balise fermante `table`.

```js
assert.match(code, /<tbody>[\s\S]*<\/tbody>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>
      
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
