---
id: 66a9cb26ec6bd41cf6c82bc5
title: Étape 9
challengeType: 0
dashedName: step-9
---

# --description--

La dernière section à ajouter au tableau est l’élément <dfn>table foot</dfn>, `tfoot`. Le pied du tableau servira à afficher la moyenne de tous les élèves du tableau.

Ajoute un élément `tfoot` au tableau.

Dans l’élément `tfoot`, ajoute un élément `tr`.

Dans l’élément `tr`, ajoute deux éléments `td`.

Le premier élément `td` doit contenir le texte `Average Grade`. Le deuxième élément `td` doit contenir le texte `78.8`.

# --hints--

Tu dois avoir une balise ouvrante `tfoot`.

```js
assert.match(code, /<tfoot>/);
```

Tu dois avoir une balise fermante `tfoot`.

```js
assert.match(code, /<\/tfoot>/);
```

Tu dois avoir un élément `tr` dans ton élément `tfoot`.

```js
assert.isNotNull(document.querySelector('tfoot tr'));
```

Ton premier élément `td` doit contenir le texte `Average Grade`. Assure-toi qu’il est dans l’élément `tr`.

```js
assert.strictEqual(document.querySelector('tfoot tr td:first-child')?.textContent.trim(), 'Average Grade');
```

Ton deuxième élément `td` doit contenir le texte `78.8`. Assure-toi qu’il est dans l’élément `tr`.

```js
assert.strictEqual(document.querySelector('tfoot tr td:nth-of-type(2)')?.textContent.trim(), '78.8');
```

Ton élément `tr` doit contenir deux éléments `td`.

```js
assert.strictEqual(document.querySelectorAll('tfoot tr td').length, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>

        <tr>
          <td>Thompson</td>
          <td>Jane</td>
          <td>77</td>
        </tr>

        <tr>
          <td>Williams</td>
          <td>Natalie</td>
          <td>83</td>
        </tr>
      </tbody>
      
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
