---
id: 66a9c7a4521d2b1b1ec6dcf0
title: Étape 7
challengeType: 0
dashedName: step-7
---

# --description--

Il est maintenant temps d’ajouter deux élèves de plus au tableau.

En suivant le même modèle que l’étape précédente, ajoute une deuxième ligne d’élève au tableau. Utilise les données suivantes pour les éléments de données du tableau :

- Nom : `Doe`
- Prénom : `Samantha`
- Note : `92`

Pour la troisième ligne d’élève, utilise les données suivantes pour les éléments de données du tableau :

- Nom : `Rodriguez`
- Prénom : `Marcus`
- Note : `88`

# --hints--

Tu dois avoir un deuxième `tr` dans ton élément `tbody`.

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(2)'));
```

Tu dois avoir un élément `td` avec le texte `Doe` dans le deuxième élément `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:first-child')?.textContent.trim(), 'Doe');
```

Tu dois avoir un élément `td` avec le texte `Samantha` dans le deuxième élément `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:nth-child(2)')?.textContent.trim(), 'Samantha');
```

Tu dois avoir un élément `td` avec le texte `92` dans le deuxième élément `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:last-child')?.textContent.trim(), '92');
```

Tu dois avoir un troisième `tr` dans ton élément `tbody`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 3);
```

Tu dois avoir un élément `td` avec le texte `Rodriguez` dans le troisième élément `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:first-child')?.textContent.trim(), 'Rodriguez');
```

Tu dois avoir un élément `td` avec le texte `Marcus` dans le troisième élément `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:nth-child(2)')?.textContent.trim(), 'Marcus');
```

Tu dois avoir un élément `td` avec le texte `88` dans le troisième élément `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:last-child')?.textContent.trim(), '88');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>
        
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
