---
id: 66a9bcf00f13a418368a272e
title: Étape 4
challengeType: 0
dashedName: step-4
---

# --description--

L'élément d'en-tête de tableau se compose d'un élément <dfn>table row</dfn>, `tr`, qui contient les éléments <dfn>table header cell</dfn>, `th`.

Voici un exemple utilisant les éléments `tr` et `th` pour un tableau de sport :

```html
<table>
  <caption>Football Scores</caption>
  <thead>
    <tr>
      <th>Team</th>
      <th>Wins</th>
      <th>Losses</th>
    </tr>
  </thead>
</table>
```

Dans ton élément `thead`, ajoute un élément `tr`.

Dans ton élément `tr`, ajoute trois éléments `th`.

Le premier élément `th` doit contenir le texte `Last Name`. Le deuxième élément `th` doit contenir le texte `First Name`. Le troisième élément `th` doit contenir le texte `Grade`.

# --hints--

Tu dois avoir une balise ouvrante `tr`.

```js
assert.match(code, /<tr>/i);
```

Tu dois avoir une balise fermante `tr`.

```js
assert.match(code, /<\/tr>/i);
```

Ton élément `tr` doit être dans ton élément `thead`.

```js
const trEl = document.querySelector('tr');
assert.strictEqual(trEl?.parentElement?.tagName, 'THEAD');
```

Ton premier élément `th` doit avoir le texte `Last Name`.

```js
const thEl = document.querySelectorAll('th')[0];
assert.strictEqual(thEl?.textContent.trim(), 'Last Name');
```

Ton deuxième élément `th` doit avoir le texte `First Name`.

```js
const thEl = document.querySelectorAll('th')[1];
assert.strictEqual(thEl?.textContent.trim(), 'First Name');
```

Ton troisième élément `th` doit avoir le texte `Grade`.

```js
const thEl = document.querySelectorAll('th')[2];
assert.strictEqual(thEl?.textContent.trim(), 'Grade');
```

Tous tes éléments `th` doivent être dans ton élément `tr`.

```js
assert.match(code, /<tr>.*<th>.*<\/th>.*<th>.*<\/th>.*<th>.*<\/th>.*<\/tr>/is);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
      --fcc-editable-region--
        
      --fcc-editable-region--
      </thead>
    </table>
  </body>
</html>
```
