---
id: 66bb6467b4b7381178420970
title: Étape 17
challengeType: 0
dashedName: step-17
---

# --description--

Bien que cette expression corresponde à `1 hundred dollars`, elle ne correspondra pas à `1  hundred  dollars`, ni à `10 dollars`.

Les messages de spam peuvent et vont trouver un moyen d'exploiter les failles de ta détection. Il est temps d'améliorer ta regex.

Remplace le premier espace littéral par l'expression `\s*`. La classe de caractères `\s` correspond aux espaces blancs, comme les espaces, les tabulations et les sauts de ligne. Le quantificateur `*` signifie « correspondre au caractère précédent 0 fois ou plus ».

Remplace le second espace littéral par `\s+`. Le quantificateur `+` signifie « correspondre au caractère précédent au moins une fois ».

# --hints--

Ton `dollarRegex` ne doit pas utiliser d'espaces littéraux.

```js
assert.notMatch(dollarRegex.source, / /);
```

Ton `dollarRegex` doit autoriser n'importe quel nombre d'espaces avant la quantité en dollars, et un ou plusieurs espaces après.

```js
let dollarStr = '10hundred dollars'
assert.deepEqual(dollarStr.match(dollarRegex), dollarStr.match(/[0-9]+\s*(hundred|thousand|million|billion)?\s+dollars/i));
dollarStr = '10 hundred   dollars'
assert.deepEqual(dollarStr.match(dollarRegex), dollarStr.match(/[0-9]+\s*(hundred|thousand|million|billion)?\s+dollars/i));
dollarStr = '10    hundred dollars'
assert.deepEqual(dollarStr.match(dollarRegex), dollarStr.match(/[0-9]+\s*(hundred|thousand|million|billion)?\s+dollars/i));
dollarStr = '10 hundreddollars'
assert.notMatch(dollarStr, dollarRegex)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width" />
    <title>Learn Regular Expressions by Building a Spam Filter</title>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <header class="main-text">
      <h1 class="title">Is this Spam?</h1>
      <p class="description">
        Enter a phrase to check if it would be marked as spam or not.
      </p>
    </header>

    <main>
      <label class="message-label" for="message-input">Message: </label>
      <textarea
        placeholder="Enter message here"
        value=""
        type="text"
        name="message"
        id="message-input"
        rows="10"
        cols="40"
      ></textarea>
      <button class="btn" id="check-message-btn" type="button">
        Check message
      </button>
      <p id="result-message"></p>
    </main>

    <footer class="footer">&copy; freeCodeCamp</footer>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --golden-yellow: #fecc4c;
  --yellow: #ffcc4c;
  --gold: #feac32;
  --orange: #ffac33;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
  color: var(--light-grey);
}

body,
#message-input:placeholder-shown {
  text-align: center;
}

textarea {
  max-width: 90%;
}

.main-text {
  margin: 25px 0;
}

.title {
  font-size: 2.5rem;
}

.description {
  margin-top: 15px;
  font-size: 1.4rem;
}

.message-label {
  display: block;
  margin-bottom: 20px;
  font-size: 1.5rem;
}

#message-input:placeholder-shown,
textarea {
  font-size: 1.1rem;
}

.btn {
  display: block;
  cursor: pointer;
  width: 200px;
  margin: 10px auto;
  color: var(--dark-grey);
  background-color: var(--gold);
  background-image: linear-gradient(var(--golden-yellow), var(--orange));
  border-color: var(--gold);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

#result {
  font-size: 2rem;
  margin: 20px 0;
}

.footer {
  margin-top: 10px;
}
```

```js
const messageInput = document.getElementById("message-input");
const result = document.getElementById("result-message");
const checkMessageButton = document.getElementById("check-message-btn");

const helpRegex = /please help|assist me/i;
--fcc-editable-region--
const dollarRegex = /[0-9]+ (hundred|thousand|million|billion)? dollars/i;
--fcc-editable-region--

const denyList = [helpRegex, dollarRegex];

const isSpam = (msg) => denyList.some((regex) => regex.test(msg));

checkMessageButton.addEventListener("click", () => {
  if (messageInput.value === "") {
    alert("Please enter a message.");
    return;
  }

  result.textContent = isSpam(messageInput.value)
    ? "Oh no! This looks like a spam message."
    : "This message does not seem to contain any spam.";
  messageInput.value = "";
});
```
