---
id: 69a92f0b9cc04bb0d5327bb5
title: Implémenter une pile
challengeType: 26
dashedName: implement-a-stack
---

# --description--

Dans ce lab, tu vas implémenter une structure de données pile à l'aide de fonctions. Une pile est une structure de données Last-In-First-Out (LIFO) où les éléments sont ajoutés et retirés du sommet.

**Objectif :** Remplir les user stories ci-dessous et faire passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu dois avoir une fonction `initStack` qui renvoie un objet avec une propriété `collection` définie à un tableau vide.
1. Tu dois avoir une fonction `push` qui ajoute un élément au sommet de la pile.
1. Tu dois avoir une fonction `pop` qui retire et renvoie l'élément au sommet de la pile, ou `undefined` s'il n'y en a pas.
1. Tu dois avoir une fonction `peek` qui renvoie l'élément au sommet de la pile sans le retirer, ou `undefined` s'il n'y en a pas.
1. Tu dois avoir une fonction `isEmpty` qui renvoie `true` si la pile ne contient aucun élément, et `false` sinon.
1. Tu dois avoir une fonction `clear` qui retire tous les éléments de la pile.

**Note :** La plupart des tests dépendent des fonctions `initStack` et `push`. Implémente-les d'abord, car les tests pour `pop`, `peek`, `isEmpty` et `clear` nécessitent d'ajouter des éléments à la pile.

# --hints--

Tu dois avoir une fonction `push`.

```js
assert.isFunction(push);
```

La fonction `push` doit ajouter un élément au sommet de la pile.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.deepEqual(stack.collection, ['first', 'second']);
```

Tu dois avoir une fonction `pop`.

```js
assert.isFunction(pop);
```

La fonction `pop` doit retirer et renvoyer l'élément au sommet de la pile.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
assert.strictEqual(pop(stack), 'third');
assert.strictEqual(pop(stack), 'second');
```

La fonction `pop` doit renvoyer `undefined` si la pile est vide.

```js
const stack = initStack();
assert.isUndefined(pop(stack));
```

La fonction `pop` doit renvoyer correctement les valeurs falsy.

```js
const stack = initStack();
push(stack, '');
assert.strictEqual(pop(stack), '');
assert.lengthOf(stack.collection, 0);
```

Tu dois avoir une fonction `peek`.

```js
assert.isFunction(peek);
```

La fonction `peek` doit renvoyer l'élément au sommet de la pile sans le retirer.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.strictEqual(peek(stack), 'second');
assert.strictEqual(peek(stack), 'second');
```

La fonction `peek` doit renvoyer `undefined` si la pile est vide.

```js
const stack = initStack();
assert.isUndefined(peek(stack));
```

La fonction `peek` doit renvoyer correctement les valeurs falsy.

```js
const stack = initStack();
push(stack, 0);
assert.strictEqual(peek(stack), 0);
```

Tu dois avoir une fonction `isEmpty`.

```js
assert.isFunction(isEmpty);
```

La fonction `isEmpty` doit renvoyer `true` pour une pile vide.

```js
const stack = initStack();
assert.isTrue(isEmpty(stack));
```

La fonction `isEmpty` doit renvoyer `false` pour une pile non vide.

```js
const stack = initStack();
push(stack, 'element');
assert.isFalse(isEmpty(stack));
```

La fonction `isEmpty` doit renvoyer `false` quand l'élément au sommet est falsy.

```js
const stack = initStack();
push(stack, false);
assert.isFalse(isEmpty(stack));
```

Tu dois avoir une fonction `clear`.

```js
assert.isFunction(clear);
```

La fonction `clear` doit retirer tous les éléments de la pile.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
clear(stack);
assert.lengthOf(stack.collection, 0);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function initStack() {
  return {
    collection: []
  };
}

function print(stack) {
  console.log(stack.collection);
}

function push(stack, element) {
  stack.collection.push(element);
}

function pop(stack) {
  return stack.collection.pop();
}

function peek(stack) {
  return stack.collection[stack.collection.length - 1];
}

function isEmpty(stack) {
  return stack.collection.length === 0;
}

function clear(stack) {
  stack.collection = [];
}
```
