---
id: 6989ef885db3a3488a13b426
title: Étape 2
challengeType: 1
dashedName: step-2
---

# --description--

Avant d'implémenter l'algorithme principal, tu dois valider l'entrée. Le paramètre `pairs` doit être un entier, car il représente le nombre de paires de parenthèses à générer.

Ajoute une instruction `if` au début de ta fonction pour vérifier si `pairs` n'est pas un entier. Utilise la méthode `Number.isInteger()` pour cela.

Si la condition est vraie, renvoie la chaîne `The number of pairs should be an integer`.

# --hints--

Tu dois avoir une instruction if avec la condition `!Number.isInteger(pairs)` à l'intérieur de la fonction `genParentheses`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*!\s*Number\.isInteger\s*\(\s*pairs\s*\)\s*\)/);
```

Tu dois renvoyer une chaîne quand l'entrée n'est pas un entier.

```js
let result = genParentheses("i");
assert.isString(result);
```

Tu dois renvoyer la chaîne `The number of pairs should be an integer` quand l'entrée est une chaîne.

```js
let result = genParentheses("i");
assert.strictEqual(result, 'The number of pairs should be an integer');
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
--fcc-editable-region--

--fcc-editable-region--
  return [];
}
```
