---
id: 6989f3632f8b8bdf168aa321
title: Étape 7
challengeType: 1
dashedName: step-7
---

# --description--

À l'intérieur de ta boucle `while`, utilise `queue.shift()` pour retirer et obtenir le premier élément de la file. Cela implémente le comportement premier-entré-premier-sorti (FIFO) caractéristique du BFS.

Déstructure ce tableau en trois variables : `current`, `opensUsed`, et `closesUsed`.

# --hints--

Tu dois utiliser `queue.shift()` pour obtenir le premier élément de la file.

```js
assert.match(genParentheses.toString(), /queue\.shift\s*\(\s*\)/);
```

Tu dois déstructurer le résultat en trois variables : `current`, `opensUsed`, et `closesUsed`.

```js
assert.match(genParentheses.toString(), /opensUsed/);
assert.match(genParentheses.toString(), /closesUsed/);
assert.match(genParentheses.toString(), /current/);
assert.match(genParentheses.toString(), /queue\.shift\s*\(\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
--fcc-editable-region--

--fcc-editable-region--
  }

  return result;
}
```
