---
id: 6989f48c3f0292c84f8d8938
title: Étape 9
challengeType: 1
dashedName: step-9
---

# --description--

Maintenant tu dois identifier quand tu as construit une combinaison complète de parenthèses. Une combinaison complète a une longueur égale à deux fois le nombre de paires (car chaque paire apporte une parenthèse ouvrante et une fermante).

À l'intérieur de ta boucle `while`, ajoute une instruction `if` pour vérifier si `current.length === 2 * pairs`. Quand cette condition est vraie, pousse `current` dans le tableau `result`.

# --hints--

Tu dois avoir une instruction `if` avec la condition `current.length === 2 * pairs` à l'intérieur de ta boucle `while`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)/);
```

Tu dois pousser `current` dans le tableau `result` à l'intérieur de ta nouvelle instruction `if`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)\s*\{[\s\S]*?result\.push\s*\(\s*current\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();
--fcc-editable-region--

--fcc-editable-region--
  }
  return result;
}

console.log(genParentheses(1));
```
