---
id: 6989f1ba358acba651d5ee6b
title: Étape 4
challengeType: 1
dashedName: step-4
---

# --description--

Maintenant tu vas préparer la structure de données pour stocker tes résultats. Crée une variable nommée `result` et initialise-la à un tableau vide. Ce tableau stockera toutes les combinaisons valides de parenthèses que tu génères.

Mets à jour ton instruction return pour renvoyer `result` au lieu d'un tableau vide.

# --hints--

Tu dois déclarer une variable nommée `result` dans ta fonction `genParentheses`.

```js
assert.match(genParentheses.toString(), /result/);
```

Tu dois initialiser `result` à un tableau vide.

```js
assert.match(genParentheses.toString(), /result\s*=\s*\[\]/);
```

Tu dois renvoyer `result` à la fin de ta fonction.

```js
assert.match(genParentheses.toString(), /return\s*result/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
--fcc-editable-region--

  return [];
--fcc-editable-region--
}
```
