---
id: 6989f555728e63ad8b349d84
title: Étape 10
challengeType: 1
dashedName: step-10
---

# --description--

Si la chaîne courante n'est pas encore complète, tu dois explorer les prochains états possibles. Ajoute une clause `else` pour l'instruction `if` actuelle.

À l'intérieur du bloc `else`, tu vas gérer l'ajout de parenthèses ouvrantes. Ajoute une instruction `if` pour vérifier si `opensUsed` est inférieur à `pairs`. Cela garantit que tu n'utilises pas plus de parenthèses ouvrantes que permis.

Si cette condition est `true`, pousse un nouveau tableau dans la `queue` : `[current + '(', opensUsed + 1, closesUsed]`. Cela représente l'état après l'ajout d'une parenthèse ouvrante.

# --hints--

Tu dois avoir une clause `else` pour ton instruction `if` à l'intérieur de la boucle `while`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)\s*\{[\s\S]*?\}\s*else\s*\{/);
```

Tu dois avoir une instruction `if` à l'intérieur du bloc `else` avec la condition `opensUsed < pairs`.

```js
assert.match(genParentheses.toString(), /else\s*\{[\s\S]*?if\s*\(\s*opensUsed\s*<\s*pairs\s*\)/);
```

Tu dois pousser `[current + '(', opensUsed + 1, closesUsed]` dans la file quand la condition est `true`.

```js
assert.match(genParentheses.toString(), /queue\.push\s*\(\s*\[\s*current\s*\+\s*['"`]\(['"`]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();
    if (current.length === 2 * pairs) {
      result.push(current);
--fcc-editable-region--
    }
--fcc-editable-region--
  }
  return result;
}

console.log(genParentheses(1));
```
