---
id: 6989f5bf7f1a1d0fc8ebd1e8
title: Étape 11
challengeType: 1
dashedName: step-11
---

# --description--

Maintenant tu dois gérer l'ajout de parenthèses fermantes. La règle clé est que tu ne peux ajouter une parenthèse fermante que si cela maintient l'équilibre, c'est-à-dire qu'il doit y avoir plus de parenthèses ouvrantes utilisées que de parenthèses fermantes.

Toujours dans le bloc `else`, ajoute une deuxième instruction `if` pour vérifier si `closesUsed` est inférieur à `opensUsed`.

Si cette condition est vraie, pousse un autre nouveau tableau dans la `queue` : `[current + ')', opensUsed, closesUsed + 1]`. Cela représente l'état après l'ajout d'une parenthèse fermante.

# --hints--

Tu dois avoir une deuxième instruction `if` avec la condition `closesUsed < opensUsed` à l'intérieur de ton bloc `else`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*closesUsed\s*<\s*opensUsed\s*\)/);
```

Tu dois pousser `[current + ')', opensUsed, closesUsed + 1]` dans la file quand la condition est vraie.

```js
assert.match(genParentheses.toString(), /queue\.push\s*\(\s*\[\s*current\s*\+\s*['"`]\)['"`]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();

    if (current.length === 2 * pairs) {
      result.push(current);
    } else {
      if (opensUsed < pairs) {
        queue.push([current + '(', opensUsed + 1, closesUsed]);
      }
--fcc-editable-region--

--fcc-editable-region--
    }
  }
  return result;
}

console.log(genParentheses(1));
```
