---
id: 6989f207306332c06b0add39
title: Étape 5
challengeType: 1
dashedName: step-5
---

# --description--

Pour l'approche de parcours en largeur, tu vas utiliser une file pour suivre les différents états pendant que tu construis les combinaisons de parenthèses. Chaque état sera représenté comme un tableau contenant trois éléments :

- La chaîne courante en cours de construction
- Le nombre de parenthèses ouvrantes utilisées jusqu'ici
- Le nombre de parenthèses fermantes utilisées jusqu'ici

Crée une variable nommée `queue` et initialise-la à un tableau contenant un tableau : `['', 0, 0]`. Cela représente l'état de départ avec une chaîne vide et zéro parenthèse utilisée.

# --hints--

Tu dois déclarer une variable nommée `queue` dans ta fonction `genParentheses`.

```js
assert.match(genParentheses.toString(), /queue/);
```

Tu dois initialiser `queue` avec un tableau contenant le tableau `['', 0, 0]`.

```js
assert.match(genParentheses.toString(), /queue\s*=\s*\[\s*\['',\s*0,\s*0\]\s*\]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
--fcc-editable-region--

--fcc-editable-region--
  let result = [];
  return result;
}
```
