---
id: 6989f2e35e6941e13e445156
title: Étape 6
challengeType: 1
dashedName: step-6
---

# --description--

Maintenant tu vas implémenter la boucle BFS principale. Crée une boucle `while` qui continue tant que la `queue` n'est pas vide, c'est-à-dire que `queue.length` est supérieur à `0`.

À l'intérieur de la boucle, affiche `queue` dans la console.

# --hints--

Tu dois créer une boucle `while` avec la condition `queue.length > 0`.

```js
assert.match(genParentheses.toString(), /while\s*\(\s*queue\.length\s*>\s*0\s*\)/);
```

Tu dois afficher `queue` dans la console à l'intérieur de ta boucle `while`.

```js
assert.match(genParentheses.toString(), /while\s*\(\s*queue\.length\s*>\s*0\s*\)/);
assert.match(genParentheses.toString(), /console\.log\s*\(\s*queue\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  let queue = [['', 0, 0]];
  let result = [];
--fcc-editable-region--

--fcc-editable-region--

  return result;
}
```
