---
id: 6989f10a72ad424d4ad6d514
title: Étape 3
challengeType: 1
dashedName: step-3
---

# --description--

Ensuite, tu dois valider que le nombre de paires est au moins un, car tu ne peux pas générer de combinaisons de parenthèses avec zéro ou des paires négatives.

Ajoute une autre instruction `if` pour vérifier si `pairs` est inférieur à `1`. Si cette condition est vraie, renvoie la chaîne `The number of pairs should be at least 1`.

# --hints--

Tu dois avoir une deuxième instruction `if` avec la condition `pairs < 1`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*pairs\s*<\s*1\s*\)/);
```

Tu dois renvoyer la chaîne `The number of pairs should be at least 1` quand `pairs` est inférieur à `1`.

```js
let result = genParentheses(0);
assert.strictEqual(result, 'The number of pairs should be at least 1');
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
--fcc-editable-region--

--fcc-editable-region--
  return [];
}
```
