---
id: 6989eb65324d97775c4965fe
title: Étape 1
challengeType: 1
dashedName: step-1
---

# --description--

Dans cet atelier, tu vas implémenter une fonction qui génère toutes les combinaisons valides de parenthèses avec une approche de parcours en largeur (BFS). Par exemple, les combinaisons valides de deux paires de parenthèses sont `(())` et `()()`.

Commence par créer une fonction nommée `genParentheses` avec un seul paramètre `pairs`. Pour l'instant, renvoie un tableau vide depuis la fonction.

# --hints--

Tu dois définir une fonction nommée `genParentheses`.

```js
assert.isFunction(genParentheses);
```

Ta fonction `genParentheses` doit avoir un seul paramètre nommé `pairs`.

```js
const genParenthesesParams = __helpers.getFunctionParams(genParentheses.toString());
assert.equal(genParenthesesParams[0].name, 'pairs')
```

Ta fonction doit renvoyer un tableau vide.

```js
assert.isArray(genParentheses(1));
assert.lengthOf(genParentheses(1), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
