---
id: 616968c2c94c8071b349c146
title: Étape 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Dans le `body`, imbrique un élément `h1` avec le texte `CSS Color Markers`.

# --hints--

Ton code devrait avoir un élément `h1`.

```js
const title = document.querySelector('h1');
assert.exists(title);
```

Tu devrais avoir une balise ouvrante `h1`.

```js
assert.match(code, /<h1\s*>/i);
```

Tu devrais avoir une balise fermante `h1`.

```js
assert.match(code, /<\/h1\s*>/i);
```

Ton élément `h1` devrait être dans l'élément `body`.

```js
assert.strictEqual(document.querySelector('h1')?.parentElement?.localName, 'body');
```

Ton élément `h1` devrait avoir le texte `CSS Color Markers`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent?.trim()?.toLowerCase(), 'css color markers')
```

Souviens-toi que la casse et l'orthographe comptent pour le texte de `h1`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent?.trim(), 'CSS Color Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
