---
id: 616971b3cd990186b66c99a1
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

Maintenant que ton fichier CSS externe est configuré, tu peux commencer à styliser la page.

Pour rappel, voici comment cibler un élément paragraphe et l’aligner à droite :

```css
p {
  text-align: right;
}
```

Crée une nouvelle règle CSS qui cible l'élément `h1`, et définis sa propriété `text-align` sur `center`.

# --hints--

Tu devrais utiliser un sélecteur `h1`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('h1'));
```

Ta règle CSS `h1` devrait avoir une propriété `text-align` définie sur `center`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('h1')?.textAlign, 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
