---
id: 61696ef7ac756c829f9e4048
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

Dans ce projet, tu vas travailler avec un fichier CSS externe pour styliser la page. Un fichier `styles.css` a déjà été créé pour toi. Mais avant de pouvoir l'utiliser, tu dois le lier à la page.

Imbrique un élément `link` dans l'élément `head`. Donne-lui un attribut `rel` défini sur `"stylesheet"` et un attribut `href` défini sur `"styles.css"`.

# --hints--

Tu devrais avoir un seul élément `link`.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

Ton élément `link` devrait être un élément vide, il ne devrait pas avoir de balise fermante `</link>`.

```js
assert.notMatch(code, /<\/link>/i);
```

Ton élément `link` devrait avoir un attribut `rel` avec la valeur `"stylesheet"`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel")?.toLowerCase();
assert.strictEqual(rel, "stylesheet");
```

Ton élément `link` devrait avoir un attribut `href` avec la valeur `"styles.css"`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

Ton élément `link` devrait être dans l'élément `head`.

```js
const linkElementRegex = /<head\s*>(?:.|\r|\n)*?<link(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, linkElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  --fcc-editable-region--

  --fcc-editable-region--
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
