---
id: 616d55bd160a95e22453a081
title: Étape 9
challengeType: 0
dashedName: step-9
---

# --description--

Maintenant que tu as un marqueur coloré centré, il est temps d'ajouter les autres marqueurs.

Dans l'élément `.container`, ajoute deux autres éléments `div` et donne à chacun la classe `marker`.

# --hints--

Ton premier nouvel élément `div` devrait avoir une balise ouvrante.

```js
assert.exists([...code.matchAll(/<div.*?>/gi)][2]);
```

Ton premier nouvel élément `div` devrait avoir une balise fermante.

```js
assert.exists([...code.matchAll(/<\/div\s*>/gi)][2]);
```

Ton deuxième nouvel élément `div` devrait avoir une balise ouvrante.

```js
assert.exists([...code.matchAll(/<div.*?>/gi)][3]);
```

Ton deuxième nouvel élément `div` devrait avoir une balise fermante.

```js
assert.exists([...code.matchAll(/<\/div\s*>/gi)][3]);
```

Tes nouveaux éléments `div` devraient être dans le `div` avec la classe `container`.

```js
const containerChildren = [...document.querySelector('.container')?.children];
containerChildren.forEach(child => assert.strictEqual(child?.localName, 'div'));
assert.lengthOf(containerChildren, 3);
```

Tes nouveaux éléments `div` devraient tous les deux avoir une classe `marker`.

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert.isNotEmpty(containerChildren);
containerChildren.forEach(child => assert.isTrue(child?.classList?.contains('marker')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
--fcc-editable-region--

--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}

```
