---
id: 616d50b93ba424d6282c99cf
title: Étape 7
challengeType: 0
dashedName: step-7
---

# --description--

La couleur d'arrière-plan a été appliquée, mais comme l'élément `div` du marqueur n'a aucun contenu, il n'a pas de hauteur par défaut.

Dans ta règle CSS `.marker`, définis la propriété `height` sur `25px` et la propriété `width` sur `200px`.

# --hints--

Ta règle CSS `.marker` devrait avoir une propriété `width` définie sur `200px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.width, '200px');
```

Ta règle CSS `.marker` devrait avoir une propriété `height` définie sur `25px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.height, '25px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  background-color: red;
--fcc-editable-region--

--fcc-editable-region--
}

```
