---
id: 616d4a84b756d9c4b8255093
title: Étape 6
challengeType: 0
dashedName: step-6
---

# --description--

Il est temps d'ajouter de la couleur au marqueur. Souviens-toi qu'une façon d'ajouter de la couleur à un élément est d'utiliser un <dfn>mot-clé de couleur</dfn> comme `black`, `cyan` ou `yellow`.

Pour rappel, voici comment cibler la classe `freecodecamp` :

```css
.freecodecamp {

}
```

Crée une nouvelle règle CSS qui cible la classe `marker`, et définis sa propriété `background-color` sur `red`.

**Note :** Tu ne verras aucun changement après avoir ajouté le CSS.

# --hints--

Tu devrais créer un sélecteur de classe pour cibler la classe `marker`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.marker'));
```

Ta règle CSS `.marker` devrait avoir une propriété `background-color` définie sur `red`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.backgroundColor, 'red');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

```
