---
id: 616d595270d902f0e7086e18
title: Étape 10
challengeType: 0
dashedName: step-10
---

# --description--

Même si tu as trois éléments `div` de marqueur séparés, ils ressemblent à un seul grand rectangle. Tu devrais ajouter de l'espace entre eux pour mieux voir chaque élément.

Quand la propriété raccourcie `margin` a deux valeurs, elle définit `margin-top` et `margin-bottom` sur la première valeur, et `margin-left` et `margin-right` sur la deuxième valeur.

Dans ta règle CSS `.marker`, définis la propriété `margin` sur `10px auto`.

# --hints--

Ta règle CSS `.marker` devrait avoir une propriété `margin` définie sur `10px auto`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.margin, '10px auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
--fcc-editable-region--
  margin: auto;
--fcc-editable-region--
}

```
