---
id: 616d3a67ccf800ad94ec89ae
title: Étape 4
challengeType: 0
dashedName: step-4
---

# --description--

Tu vas maintenant ajouter des éléments que tu transformeras ensuite en marqueurs colorés.

D'abord, dans l'élément `body`, ajoute un élément `div` avec un attribut `class` défini sur `container`. Assure-toi que l'élément `div` est placé sous l'élément `h1`.

# --hints--

Tu devrais avoir une balise ouvrante `div`.

```js
assert.match(code, /<div.*>/i);
```

Tu devrais avoir une balise fermante `div`.

```js
assert.match(code, /<\/div\s*>/i);
```

Ton élément `div` devrait être dans ton élément `body`.

```js
assert.strictEqual(document.querySelector('div')?.parentElement?.localName, 'body');
```

Ton élément `div` devrait être après l'élément `h1`.

```js
assert.exists(document.querySelector('h1 + div'));
```

Ton élément `div` devrait avoir un attribut `class` défini sur `container`.

```js
assert.isTrue(document.querySelector('div')?.classList?.contains('container'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```

```css
h1 {
  text-align: center;
}
```
