---
id: 616d525007b8c5d8b3308b1c
title: Étape 8
challengeType: 0
dashedName: step-8
---

# --description--

Ton marqueur aurait meilleure allure s'il était centré sur la page. Une façon simple de le faire est d'utiliser la <dfn>propriété raccourcie</dfn> `margin`.

Tu peux définir séparément la zone de marge des éléments avec des propriétés comme `margin-top` et `margin-left`, mais la propriété raccourcie `margin` permet de définir facilement plusieurs zones de marge en même temps.

Pour centrer ton marqueur sur la page, définis sa propriété `margin` sur `auto`. Cela définit `margin-top`, `margin-right`, `margin-bottom` et `margin-left` sur `auto`.

# --hints--

Ta règle CSS `.marker` devrait avoir une propriété `margin` définie sur `auto`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.margin, 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
--fcc-editable-region--

--fcc-editable-region--
}

```
