---
id: 6176519636a76b810ab1219a
title: Étape 15
challengeType: 0
dashedName: step-15
---

# --description--

Crée une règle CSS qui cible la classe `two` et définis sa propriété `background-color` sur `green`.

Crée aussi une règle CSS séparée qui cible la classe `three` et définis son `background-color` sur `blue`.

# --hints--

Tu devrais utiliser un sélecteur de classe pour cibler la classe `two`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.two'));
```

Ta règle CSS `.two` devrait avoir une propriété `background-color` définie sur `green`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'green');
```

Tu devrais utiliser un sélecteur de classe pour cibler la classe `three`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.three'));
```

Ta règle CSS `.three` devrait avoir une propriété `background-color` définie sur `blue`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

--fcc-editable-region--

--fcc-editable-region--

```
