---
id: 61764c602bee6974e7790f35
title: Étape 11
challengeType: 0
dashedName: step-11
---

# --description--

Pour donner des couleurs différentes aux marqueurs, tu devras ajouter une classe unique à chacun. Tu peux ajouter plusieurs classes à un élément en les listant dans l'attribut `class` et en les séparant par un espace. Par exemple, le code suivant ajoute les classes `animal` et `dog` à un élément `div` :

```html
<div class="animal dog">
```

Si tu ajoutes plusieurs classes à un élément HTML, les styles des premières classes listées dans la feuille de style peuvent être remplacés par les classes suivantes.

Pour commencer, ajoute la classe `one` au premier élément `div` de marqueur.


# --hints--

Tu devrais ajouter la classe `one` au premier élément `div` de marqueur.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert.isTrue(containerFirstChild?.classList?.contains('one'));
```

Ton premier `div` de marqueur devrait avoir les classes `marker` et `one`.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert.isTrue(containerFirstChild?.classList?.contains('marker'));
assert.isTrue(containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker">
--fcc-editable-region--
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```
