---
id: 617b65579ce424bf5f02ca73
title: Étape 16
challengeType: 0
dashedName: step-16
---

# --description--

Comme tu l'as appris dans les leçons, il existe deux modèles de couleur principaux : le modèle RGB <dfn>additif</dfn> (rouge, vert, bleu), utilisé dans les appareils électroniques, et le modèle CMYK <dfn>soustractif</dfn> (cyan, magenta, jaune, noir), utilisé pour l'impression.

Dans ce projet, tu vas travailler avec le modèle RGB. Cela signifie que les couleurs commencent en noir et changent à mesure que différents niveaux de rouge, de vert et de bleu sont introduits. Une façon simple de voir cela est d'utiliser la fonction CSS `rgb`.

Crée une nouvelle règle CSS qui cible la classe `container` et définis son `background-color` sur noir avec `rgb(0, 0, 0)`.

# --hints--

Tu devrais utiliser un sélecteur de classe pour cibler la classe `container`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.container'));
```

Ta règle CSS `.container` devrait avoir une propriété `background-color` définie sur `rgb(0, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor, 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
