---
id: 617b8b38f32bf2080a140675
title: Étape 17
challengeType: 0
dashedName: step-17
---

# --description--

Comme tu l'as appris dans les leçons précédentes, une fonction est un morceau de code qui peut recevoir une entrée et effectuer une action précise. La fonction CSS `rgb` accepte des valeurs, ou <dfn>arguments</dfn>, pour le rouge, le vert et le bleu, et produit une couleur :

```css
rgb(red, green, blue);
```

Chaque valeur de rouge, de vert et de bleu est un nombre de `0` à `255`. `0` signifie qu'il y a 0 % de cette couleur, ce qui donne du noir. `255` signifie qu'il y a 100 % de cette couleur.

Dans la règle CSS `.one`, remplace le mot-clé de couleur `red` par la fonction `rgb`. Pour la fonction `rgb`, définis la valeur du rouge sur `255`, la valeur du vert sur `0` et la valeur du bleu sur `0`.

# --hints--

Ta règle CSS `.one` ne devrait pas utiliser le mot-clé de couleur `red` pour définir la propriété `background-color`.

```js
assert.notStrictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'red');
```

Ta règle CSS `.one` devrait avoir une propriété `background-color` définie sur `rgb(255, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
--fcc-editable-region--
  background-color: red;
--fcc-editable-region--
}

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
