---
id: 617bc3386dc7d07d6469bf20
title: Étape 27
challengeType: 0
dashedName: step-27
---

# --description--

Remarque que, pour créer l'orange, tu as augmenté l'intensité du rouge et diminué l'intensité du vert dans les valeurs `rgb`. C'est parce que l'orange est créé en mélangeant du rouge avec du jaune.

Pour créer la couleur tertiaire vert printemps, combine le cyan avec le vert. Mets à jour la fonction `rgb` dans la règle CSS `.two` en définissant le vert sur sa valeur maximale et le bleu sur `127`.

# --hints--

Ta règle CSS `.two` devrait avoir une propriété `background-color` définie sur `rgb(0, 255, 127)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 255, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 127, 0);
}

.two {
--fcc-editable-region--
  background-color: rgb(0, 255, 255);
--fcc-editable-region--
}

.three {
  background-color: rgb(255, 0, 255);
}

```
