---
id: 617bbb6b97a83f6d1f7d6e38
title: Étape 26
challengeType: 0
dashedName: step-26
---

# --description--

Maintenant que tu t'es entraîné avec les couleurs secondaires, voyons comment créer des <dfn>couleurs tertiaires</dfn>. Les couleurs tertiaires sont créées en combinant une couleur primaire avec une couleur secondaire voisine.

Pour créer la couleur tertiaire orange, mets à jour la fonction `rgb` dans la règle CSS `.one` en définissant le rouge sur sa valeur maximale (`255`) et le vert sur `127`.

# --hints--

Ta règle CSS `.one` devrait avoir une propriété `background-color` définie sur `rgb(255, 127, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(255, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
--fcc-editable-region--
  background-color: rgb(255, 255, 0);
--fcc-editable-region--
}

.two {
  background-color: rgb(0, 255, 255);
}

.three {
  background-color: rgb(255, 0, 255);
}

```
