---
id: 617bd6ec666b1da2587e4e37
title: Étape 29
challengeType: 0
dashedName: step-29
---

# --description--

Il existe trois autres couleurs tertiaires : le vert chartreuse (vert + jaune), l’azur (bleu + cyan) et le rose (rouge + magenta).

Pour créer du vert chartreuse, mets à jour la fonction `rgb` dans la règle CSS `.one` afin que le rouge soit à `127`, et définis le vert sur la valeur maximale.

Pour l'azur, mets à jour la fonction `rgb` dans la règle CSS `.two` afin que le vert soit à `127` et que le bleu soit à la valeur maximale.

Et pour le rose, parfois appelé rose vif, mets à jour la fonction `rgb` dans la règle CSS `.three` afin que le bleu soit à `127` et que le rouge soit à la valeur maximale.

# --hints--

Ta règle CSS `.one` devrait avoir une propriété `background-color` définie sur `rgb(127, 255, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(127, 255, 0)');
```

Ta règle CSS `.two` devrait avoir une propriété `background-color` définie sur `rgb(0, 127, 255)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 127, 255)');
```

Ta règle CSS `.three` devrait avoir une propriété `background-color` définie sur `rgb(255, 0, 127)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'rgb(255, 0, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 127, 0);
}

.two {
  background-color: rgb(0, 255, 127);
}

.three {
  background-color: rgb(127, 0, 255);
}
--fcc-editable-region--

```
