---
id: 617b9ad735109e217284e095
title: Étape 22
challengeType: 0
dashedName: step-22
---

# --description--

Maintenant que tu as les couleurs primaires RGB, il est temps de les combiner.

Pour la fonction `rgb` dans la règle `.container`, définis les valeurs rouge, verte et bleue au maximum de `255`.

# --hints--

Ta règle CSS `.container` devrait avoir une propriété `background-color` définie sur `rgb(255, 255, 255)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor, 'rgb(255, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
--fcc-editable-region--
  background-color: rgb(0, 0, 0);
--fcc-editable-region--
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 255, 0);
}

.three {
  background-color: rgb(0, 0, 255);
}

```
