---
id: 617b954d9f4f6217a749380e
title: Étape 20
challengeType: 0
dashedName: step-20
---

# --description--

Ajoute maintenant un peu plus d'espace vertical entre tes marqueurs et le bord de l'élément `container` qui les contient.

Dans la règle CSS `.container`, utilise la propriété raccourcie `padding` pour ajouter `10px` de remplissage en haut et en bas, et définis le remplissage gauche et droit sur `0`. Cela fonctionne de manière similaire à la propriété raccourcie `margin` que tu as utilisée plus tôt.

# --hints--

Tu ne devrais pas supprimer la propriété `background-color` et sa valeur de la règle CSS `.container`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor, 'rgb(0, 0, 0)');
```

Ta règle CSS `.container` devrait avoir une propriété `padding` définie sur `10px 0`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.container')?.padding, '10px 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
--fcc-editable-region--

--fcc-editable-region--
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 127, 0);
}

.three {
  background-color: rgb(0, 0, 255);
}

```
