---
id: 617b97abd9f3f61d1590b815
title: Étape 21
challengeType: 0
dashedName: step-21
---

# --description--

Dans le modèle de couleur RGB additif, les <dfn>couleurs primaires</dfn> sont des couleurs qui, lorsqu'elles sont combinées, créent du blanc pur. Mais pour que cela arrive, chaque couleur doit être à son intensité maximale.

Avant de combiner les couleurs, remets ton marqueur vert en vert pur. Pour la fonction `rgb` dans la règle CSS `.two`, remets le vert à sa valeur maximale de `255`.

# --hints--

Ta règle CSS `.two` devrait avoir une propriété `background-color` définie sur `rgb(0, 255, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
--fcc-editable-region--
  background-color: rgb(0, 127, 0);
--fcc-editable-region--
}

.three {
  background-color: rgb(0, 0, 255);
}

```
