---
id: 617bb5624a75e86463b7e638
title: Étape 23
challengeType: 0
dashedName: step-23
---

# --description--

Les <dfn>couleurs secondaires</dfn> sont les couleurs que tu obtiens lorsque tu combines des couleurs primaires. Tu as peut-être remarqué certaines couleurs secondaires à l'étape précédente en modifiant les valeurs rouge, verte et bleue.

Pour créer la première couleur secondaire, le jaune, mets à jour la fonction `rgb` dans la règle CSS `.one` afin de combiner le rouge pur et le vert pur.

# --hints--

Ta règle CSS `.one` devrait avoir une propriété `background-color` définie sur `rgb(255, 255, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(255, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
--fcc-editable-region--
  background-color: rgb(255, 0, 0);
--fcc-editable-region--
}

.two {
  background-color: rgb(0, 255, 0);
}

.three {
  background-color: rgb(0, 0, 255);
}

```
