---
id: 617bb77353188166af43f3ac
title: Étape 24
challengeType: 0
dashedName: step-24
---

# --description--

Pour créer la couleur secondaire suivante, le cyan, mets à jour la fonction `rgb` dans la règle CSS `.two` afin de combiner le vert pur et le bleu pur.

# --hints--

Ta règle CSS `.two` devrait avoir une propriété `background-color` définie sur `rgb(0, 255, 255)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 255, 0);
}

.two {
--fcc-editable-region--
  background-color: rgb(0, 255, 0);
--fcc-editable-region--
}

.three {
  background-color: rgb(0, 0, 255);
}

```
