---
id: 617b8e0d93a8d10d9a90e720
title: Étape 18
challengeType: 0
dashedName: step-18
---

# --description--

Remarque que le `background-color` de ton marqueur est toujours rouge. C'est parce que tu as défini la valeur rouge de la fonction `rgb` au maximum de `255`, soit 100 % de rouge, et que tu as défini les valeurs verte et bleue sur `0`.

Utilise maintenant la fonction `rgb` pour définir les autres couleurs.

Dans la règle CSS `.two`, utilise la fonction `rgb` pour définir le `background-color` sur la valeur maximale pour le vert, et sur `0` pour les autres valeurs. Et dans la règle CSS `.three`, utilise la fonction `rgb` pour définir le `background-color` sur la valeur maximale pour le bleu, et sur `0` pour les autres valeurs.

# --hints--

Ta règle CSS `.two` ne devrait pas utiliser le mot-clé de couleur `green` pour définir la propriété `background-color`.

```js
assert.notStrictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'green');
```

Ta règle CSS `.two` devrait avoir une propriété `background-color` définie sur `rgb(0, 255, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 255, 0)');
```

Ta règle CSS `.three` ne devrait pas utiliser le mot-clé de couleur `blue` pour définir la propriété `background-color`.

```js
assert.notStrictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'blue');
```

Ta règle CSS `.three` devrait avoir une propriété `background-color` définie sur `rgb(0, 0, 255)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: green;
}

.three {
  background-color: blue;
}
--fcc-editable-region--

```
