---
id: 617b92b9de349513466f6156
title: Étape 19
challengeType: 0
dashedName: step-19
---

# --description--

Même si les marqueurs rouge et bleu ont le même aspect, le marqueur vert est beaucoup plus clair qu'avant. C'est parce que le mot-clé de couleur `green` est en réalité une nuance plus sombre, à peu près à mi-chemin entre le noir et la valeur maximale du vert.

Dans la règle CSS `.two`, définis la valeur verte dans la fonction `rgb` sur `127` pour réduire son intensité.

# --hints--

Ta règle CSS `.two` devrait avoir une propriété `background-color` définie sur `rgb(0, 127, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
--fcc-editable-region--
  background-color: rgb(0, 255, 0);
--fcc-editable-region--
}

.three {
  background-color: rgb(0, 0, 255);
}

```
