---
id: 618a00ed1ca871a2b3aca0cb
title: Étape 30
challengeType: 0
dashedName: step-30
---

# --description--

Maintenant que tu as parcouru toutes les couleurs primaires, secondaires et tertiaires sur une roue chromatique, il sera plus facile de comprendre d'autres concepts de théorie des couleurs et leur impact sur le design.

D'abord, dans les règles CSS `.one`, `.two` et `.three`, ajuste les valeurs dans la fonction `rgb` afin que le `background-color` de chaque élément soit défini sur noir pur. Souviens-toi que la fonction `rgb` utilise le modèle de couleur additif, où les couleurs commencent en noir et changent à mesure que les valeurs de rouge, de vert et de bleu augmentent.

# --hints--

Ta règle CSS `.one` devrait avoir une propriété `background-color` définie sur `rgb(0, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(0, 0, 0)');
```

Ta règle CSS `.two` devrait avoir une propriété `background-color` définie sur `rgb(0, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 0, 0)');
```

Ta règle CSS `.three` devrait avoir une propriété `background-color` définie sur `rgb(0, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(127, 255, 0);
}

.two {
  background-color: rgb(0, 127, 255);
}

.three {
  background-color: rgb(255, 0, 127);
}
--fcc-editable-region--

```
