---
id: 618a0927005553b74bfa05ae
title: Étape 31
challengeType: 0
dashedName: step-31
---

# --description--

Une roue chromatique est un cercle où les couleurs similaires, ou <dfn>teintes</dfn>, sont proches les unes des autres, tandis que les couleurs différentes sont plus éloignées. Par exemple, le rouge pur se trouve entre les teintes rose et orange.

Deux couleurs opposées l'une à l'autre sur la roue chromatique sont appelées <dfn>couleurs complémentaires</dfn>. Si deux couleurs complémentaires sont combinées, elles produisent du gris. Mais lorsqu'elles sont placées côte à côte, ces couleurs produisent un fort contraste visuel et paraissent plus lumineuses.

Dans la fonction `rgb` de la règle CSS `.one`, définis la valeur du rouge au maximum de `255` pour produire du rouge pur. Dans la fonction `rgb` de la règle CSS `.two`, définis les valeurs du vert et du bleu au maximum de `255` pour produire du cyan.

# --hints--

Ta règle CSS `.one` devrait avoir une propriété `background-color` définie sur `rgb(255, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(255, 0, 0)');
```

Ta règle CSS `.two` devrait avoir une propriété `background-color` définie sur `rgb(0, 255, 255)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(0, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 0);
}

```
