---
id: 618a132676346ac9f7fd59dd
title: Étape 34
challengeType: 0
dashedName: step-34
---

# --description--

Remarque comme tes yeux sont naturellement attirés par la couleur rouge au centre. Quand tu conçois un site, tu peux utiliser cet effet pour attirer l'attention sur des titres, boutons ou liens importants.

Il existe plusieurs autres combinaisons de couleurs importantes en dehors des couleurs complémentaires, mais tu les apprendras un peu plus tard.

Pour l’instant, utilise la fonction `rgb` dans la règle CSS `.two` pour définir le `background-color` sur noir.

# --hints--

Ta règle CSS `.two` devrait avoir une propriété `background-color` définie sur `rgb(0, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
  background-color: rgb(0, 255, 255);
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(0, 0, 0);
}

.two {
--fcc-editable-region--
  background-color: rgb(255, 0, 0);
--fcc-editable-region--
}

.three {
  background-color: rgb(0, 0, 0);
}

```
