---
id: 619b7fd56aa2253778bcf5f7
title: Étape 43
challengeType: 0
dashedName: step-43
---

# --description--

Le modèle de couleur <dfn>HSL</dfn>, ou teinte, saturation et luminosité, est une autre façon de représenter les couleurs.

La fonction CSS `hsl` accepte 3 valeurs : un nombre de 0 à 360 pour la teinte, un pourcentage de 0 à 100 pour la saturation, et un pourcentage de 0 à 100 pour la luminosité.

Si tu imagines une roue chromatique, la teinte rouge est à 0 degré, le vert est à 120 degrés et le bleu est à 240 degrés.

La saturation est l'intensité d'une couleur, de 0 %, ou gris, à 100 % pour une couleur pure. Tu dois ajouter le signe de pourcentage `%` aux valeurs de saturation et de luminosité.

La luminosité indique à quel point une couleur paraît claire, de 0 %, ou noir complet, à 100 %, blanc complet, avec 50 % comme valeur neutre.

Dans la règle CSS `.blue`, utilise la fonction `hsl` pour changer la propriété `background-color` en bleu pur. Définis la teinte sur `240`, la saturation sur `100%` et la luminosité sur `50%`.

# --hints--

Ta règle CSS `.blue` devrait avoir une propriété `background-color` définie sur `hsl(240, 100%, 50%)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.green {
  background-color: #007F00;
}

.blue {
--fcc-editable-region--
  background-color: rgb(0, 0, 0);
--fcc-editable-region--
}

```
