---
id: 619b74fa777a2b2473c94f82
title: Étape 40
challengeType: 0
dashedName: step-40
---

# --description--

Mets à jour le sélecteur de classe CSS `.two` pour qu’il cible la nouvelle classe `green`. Et mets à jour le sélecteur de classe `.three` pour qu’il cible la nouvelle classe `blue`.

# --hints--

Ton code ne devrait plus avoir de sélecteur de classe `.two`.

```js
assert.notExists(new __helpers.CSSHelp(document).getStyle('.two'));
```

Tu devrais utiliser un sélecteur de classe pour cibler la classe `green`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.green'));
```

Ta règle CSS `.green` devrait avoir une propriété `background-color` définie sur `rgb(0, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor, 'rgb(0, 0, 0)');
```

Ton code ne devrait plus avoir de sélecteur de classe `.three`.

```js
assert.notExists(new __helpers.CSSHelp(document).getStyle('.three'));
```

Tu devrais utiliser un sélecteur de classe pour cibler la classe `blue`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.blue'));
```

Ta règle CSS `.blue` devrait avoir une propriété `background-color` définie sur `rgb(0, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor, 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(0, 0, 0);
}

.three {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

```
