---
id: 619b761916dac02643940022
title: Étape 41
challengeType: 0
dashedName: step-41
---

# --description--

Une façon très courante d'appliquer une couleur à un élément avec CSS est d'utiliser des valeurs <dfn>hexadécimales</dfn>, ou valeurs hex. Même si les valeurs hex semblent compliquées, ce ne sont en réalité qu'une autre forme de valeurs RGB.

Les valeurs de couleur hex commencent par le caractère `#` et utilisent six caractères de 0 à 9 et de A à F. La première paire de caractères représente le rouge, la deuxième le vert et la troisième le bleu. Par exemple, `#4B5320`.

Dans le sélecteur de classe `.green`, définis la propriété `background-color` sur un code couleur hex avec les valeurs `00` pour le rouge, `FF` pour le vert et `00` pour le bleu.

# --hints--

Ta règle CSS `.green` devrait avoir une propriété `background-color` définie sur `#00FF00`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor, 'rgb(0, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.green {
--fcc-editable-region--
  background-color: rgb(0, 0, 0);
--fcc-editable-region--
}

.blue {
  background-color: rgb(0, 0, 0);
}

```
