---
id: 619b7c3c83de403126b69c1e
title: Étape 42
challengeType: 0
dashedName: step-42
---

# --description--

Tu connais peut-être déjà les valeurs décimales, ou en base 10, qui vont de 0 à 9. Les valeurs hexadécimales, ou en base 16, vont de 0 à 9, puis de A à F :

```js
0, 1, 2, 3, 4, 5, 6, 7, 8, 9, A, B, C, D, E, F
```

Avec les couleurs hex, `00` représente 0 % de cette couleur, et `FF` représente 100 %. Ainsi, `#00FF00` correspond à 0 % de rouge, 100 % de vert et 0 % de bleu, et c’est la même chose que `rgb(0, 255, 0)`.

Réduis l'intensité du vert en définissant la valeur verte de la couleur hex sur `7F`.

# --hints--

Ta règle CSS `.green` devrait avoir une propriété `background-color` définie sur `#007F00`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor, 'rgb(0, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.green {
--fcc-editable-region--
  background-color: #00FF00;
--fcc-editable-region--
}

.blue {
  background-color: rgb(0, 0, 0);
}

```
