---
id: 619b72a0db211f1c29afb906
title: Étape 37
challengeType: 0
dashedName: step-37
---

# --description--

Mets à jour la règle CSS `.one` pour cibler la nouvelle classe `red`.

# --hints--

Ton code ne devrait plus avoir de sélecteur de classe `.one`.

```js
assert.notExists(new __helpers.CSSHelp(document).getStyle('.one'));
```

Tu devrais utiliser un sélecteur de classe pour cibler la classe `red`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.red'));
```

Ta règle CSS `.red` devrait avoir une propriété `background-color` définie sur `rgb(0, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.red')?.backgroundColor, 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
--fcc-editable-region--
  background-color: rgb(0, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}

.three {
  background-color: rgb(0, 0, 0);
}
```
