---
id: 61a493ead935c148d2b76312
title: Étape 45
challengeType: 0
dashedName: step-45
---

# --description--

La fonction `linear-gradient` est très flexible — voici la syntaxe de base que tu vas revoir dans ce tutoriel :

```css
linear-gradient(gradientDirection, color1, color2, ...);
```

`gradientDirection` est la direction de la ligne utilisée pour la transition. `color1` et `color2` sont des arguments de couleur, c’est-à-dire les couleurs utilisées dans la transition elle-même. Elles peuvent être de n’importe quel type, y compris des mots-clés de couleur, des valeurs hex, `rgb` ou `hsl`.

Tu vas maintenant appliquer un dégradé du rouge vers le vert le long d’une ligne de 90 degrés au premier marqueur.

D'abord, dans la règle CSS `.red`, définis la propriété `background` sur `linear-gradient()` et passe-lui la valeur `90deg` comme `gradientDirection`.

# --hints--

Ta règle CSS `.red` devrait avoir une propriété `background` avec la valeur `linear-gradient(90deg)`.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:linear-gradient\(90deg\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
--fcc-editable-region--
  background: rgb(255, 0, 0);
--fcc-editable-region--
}

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
