---
id: 61a489b8579e87364b2d2cdb
title: Étape 44
challengeType: 0
dashedName: step-44
---

# --description--

Tu t'es entraîné avec quelques façons de définir des couleurs plates en CSS, mais tu peux aussi utiliser une transition de couleur, ou <dfn>dégradé</dfn>, sur un élément.

Un dégradé se produit lorsqu’une couleur se transforme en une autre. La fonction CSS `linear-gradient` te permet de contrôler la direction de la transition le long d’une ligne, ainsi que les couleurs utilisées.

Une chose à retenir est que la fonction `linear-gradient` crée en réalité un élément `image`, et qu'elle est généralement associée à la propriété `background`, qui peut accepter une image comme valeur.

Dans la règle CSS `.red`, remplace la propriété `background-color` par `background`.

# --hints--

Ta règle CSS `.red` devrait avoir une propriété `background` avec la valeur `rgb(255, 0, 0)`.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:rgb\(255,0,0\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
--fcc-editable-region--
  background-color: rgb(255, 0, 0);
--fcc-editable-region--
}

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
