---
id: 61a49d15bdbb5e57cc6fd280
title: Étape 49
challengeType: 0
dashedName: step-49
---

# --description--

Les arrêts de couleur te permettent d’ajuster précisément l’endroit où les couleurs sont placées le long de la ligne du dégradé. Ils sont utilisés dans la fonction `linear-gradient` pour préciser où une couleur se termine et où la transition vers la couleur suivante commence. Une valeur en pourcentage ou une valeur de longueur peut être utilisée pour définir la position de l’arrêt de couleur.

Par exemple, dans ce dégradé rouge-noir, la transition du rouge au noir se produit au point 90 % le long de la ligne du dégradé, donc le rouge occupe la majeure partie de l’espace disponible :

```css
linear-gradient(90deg, red 90%, black);
```

Dans la fonction `linear-gradient`, ajoute un arrêt de couleur `75%` après le premier argument de couleur rouge. N’ajoute pas d’arrêts de couleur aux autres arguments de couleur.

# --hints--

Ta règle CSS `.red` devrait avoir une propriété `background` définie sur `linear-gradient(90deg, rgb(255, 0, 0) 75%, rgb(0, 255, 0), rgb(0, 0, 255))`.

```js
assert.include(['linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))', 'rgba(0,0,0,0)linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
--fcc-editable-region--
  background: linear-gradient(90deg, rgb(255, 0, 0), rgb(0, 255, 0), rgb(0, 0, 255));
--fcc-editable-region--
}

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
