---
id: 61a5bfe091060f1d6a629dd0
title: Étape 59
challengeType: 0
dashedName: step-59
---

# --description--

Même sans les arrêts de couleur, tu as peut-être remarqué que les couleurs du marqueur vert changent aux mêmes endroits que celles du marqueur rouge. La première couleur est au début (0 %), la deuxième au milieu (50 %) et la dernière à la fin (100 %) de la ligne du dégradé.

La fonction `linear-gradient` calcule automatiquement ces valeurs pour toi et place les couleurs uniformément le long de la ligne du dégradé par défaut.

Dans la règle CSS `.red`, supprime les trois arrêts de couleur de la fonction `linear-gradient` pour nettoyer un peu ton code.

# --hints--

Ta règle CSS `.red` devrait avoir une propriété `background` définie sur `linear-gradient(180deg, rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27))`.

```js
assert.include(['linear-gradient(rgb(122,74,14),rgb(245,62,113),rgb(162,27,27))', 'rgba(0,0,0,0)linear-gradient(rgb(122,74,14),rgb(245,62,113),rgb(162,27,27))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
--fcc-editable-region--
  background: linear-gradient(180deg, rgb(122, 74, 14) 0%, rgb(245, 62, 113) 50%, rgb(162, 27, 27) 100%);
--fcc-editable-region--
}

.green {
  background: linear-gradient(180deg, #55680D, #71F53E, #116C31);
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
