---
id: 61adc60b69cd4b87739d2174
title: Étape 66
challengeType: 0
dashedName: step-66
---

# --description--

Crée une nouvelle règle CSS qui cible la classe `sleeve`. Définis la propriété `width` sur `110px` et la propriété `height` sur `25px`.

# --hints--

Tu devrais utiliser un sélecteur de classe pour cibler la classe `sleeve`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.sleeve'));
```

Ta règle CSS `.sleeve` devrait avoir une propriété `width` définie sur `110px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.sleeve')?.width, '110px');
```

Ta règle CSS `.sleeve` devrait avoir une propriété `height` définie sur `25px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.sleeve')?.height, '25px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--

--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
