---
id: 61b093429e7fc020b43b1bb6
title: Étape 71
challengeType: 0
dashedName: step-71
---

# --description--

Ton manchon a belle allure, mais il serait encore mieux s'il était positionné davantage vers le côté droit du marqueur. Une façon de le faire est d'ajouter un autre élément avant le manchon pour le pousser vers la droite.

Ajoute un nouveau `div` avec la classe `cap` avant l'élément `div` du manchon.

# --hints--

Ton nouvel élément `div` devrait être dans le `div` du marqueur rouge.

```js
const redMarkerChildren = [...document.querySelector('div.red')?.children];
redMarkerChildren.forEach(child => {
  assert.strictEqual(child?.localName, 'div')
})
assert.lengthOf(redMarkerChildren, 2);
```

Ton nouvel élément `div` devrait avoir une classe `cap`.

```js
const redMarkerChildren = [...document.querySelector('div.red')?.children];
assert.isTrue(redMarkerChildren.some(child => child?.classList?.contains('cap')));
```

Ton `div` de capuchon devrait être avant le `div` du manchon.

```js
const redMarkerChildren = [...document.querySelector('div.red')?.children];
const cap = document.querySelector('div.cap');
const sleeve = document.querySelector('div.sleeve');
assert.isBelow(redMarkerChildren.indexOf(cap), redMarkerChildren.indexOf(sleeve));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      --fcc-editable-region--

      --fcc-editable-region--
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
