---
id: 61b093179e7fc020b43b1bb3
title: Étape 68
challengeType: 0
dashedName: step-68
---

# --description--

L’<dfn>opacité</dfn> décrit à quel point quelque chose est opaque, ou non transparent. Par exemple, un mur plein est opaque, et aucune lumière ne peut le traverser. Mais un verre à boire est beaucoup plus transparent, et tu peux voir à travers le verre jusqu'à l'autre côté.

Avec la propriété CSS `opacity`, tu peux contrôler à quel point un élément est opaque ou transparent. Avec la valeur `0`, ou 0 %, l'élément sera complètement transparent, et avec `1.0`, ou 100 %, l'élément sera complètement opaque comme il l'est par défaut.

Dans la règle CSS `.sleeve`, définis la propriété `opacity` sur `0.5`.

# --hints--

Ta règle CSS `.sleeve` devrait avoir une propriété `opacity` définie sur `0.5`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.sleeve')?.opacity, '0.5');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: white;
--fcc-editable-region--

--fcc-editable-region--
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
