---
id: 61b093219e7fc020b43b1bb4
title: Étape 69
challengeType: 0
dashedName: step-69
---

# --description--

Une autre façon de définir l'opacité d'un élément est d'utiliser le <dfn>canal alpha</dfn>. Comme la propriété `opacity`, le canal alpha contrôle à quel point une couleur est transparente ou opaque.

Tu as déjà défini l’opacité du manchon avec une couleur nommée et la propriété `opacity`, mais tu peux ajouter un canal alpha aux autres propriétés de couleur CSS.

Dans la règle `.sleeve`, supprime la propriété `opacity` et sa valeur.

# --hints--

Ta règle CSS `.sleeve` ne devrait pas avoir de propriété `opacity` ni de valeur.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.sleeve')?.opacity);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: white;
--fcc-editable-region--
  opacity: 0.5;
--fcc-editable-region--
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
