---
id: 61b0a44a6b865738ba49b9fb
title: Étape 81
challengeType: 0
dashedName: step-81
---

# --description--

La dernière chose que tu vas faire est d’ajouter une légère ombre à chaque marqueur pour les rendre encore plus réalistes.

La propriété `box-shadow` te permet d’appliquer une ou plusieurs ombres autour d’un élément. Voici la syntaxe de base :

```css
box-shadow: offsetX offsetY color;
```

Voici comment fonctionnent les valeurs `offsetX` et `offsetY` :

* `offsetX` et `offsetY` acceptent tous les deux des valeurs numériques en `px` et dans d’autres unités CSS
* une valeur positive de `offsetX` déplace l’ombre vers la droite, et une valeur négative la déplace vers la gauche
* une valeur positive de `offsetY` déplace l’ombre vers le bas, et une valeur négative la déplace vers le haut
* si tu veux une valeur de zéro (`0`) pour `offsetX`, `offsetY`, ou les deux, tu n’as pas besoin d’ajouter d’unité. Tous les navigateurs comprennent que zéro signifie aucun changement.

La hauteur et la largeur de l'ombre sont déterminées par la hauteur et la largeur de l'élément auquel elle est appliquée. Tu peux aussi utiliser une valeur optionnelle `spreadRadius` pour étendre la portée de l'ombre. Plus d'informations à ce sujet plus tard.

Commence par ajouter une ombre simple au marqueur rouge.

Dans la règle CSS `.red`, ajoute la propriété `box-shadow` avec les valeurs `5px` pour `offsetX`, `5px` pour `offsetY` et `red` pour `color`.

# --hints--

Ta règle CSS `.red` devrait avoir une propriété raccourcie `box-shadow` avec la valeur `5px 5px red`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.red')?.boxShadow, 'red 5px 5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker blue">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left: 10px double rgba(0, 0, 0, 0.75);
}

.cap, .sleeve {
  display: inline-block;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
--fcc-editable-region--

--fcc-editable-region--
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
