---
id: 56533eb9ac21ba0edf2244b3
title: Crée un convertisseur de Celsius en Fahrenheit
challengeType: 26
dashedName: build-a-celsius-to-fahrenheit-converter
---

# --description--

Dans ce lab, tu vas écrire une fonction qui convertit la température de Celsius en Fahrenheit. La formule pour convertir de Celsius en Fahrenheit est :

```js
fahrenheit = celsius * (9/5) + 32
```

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois créer une fonction nommée `convertCtoF`.
2. `convertCtoF` doit prendre un seul argument numérique, qui est la température en Celsius.
3. `convertCtoF` doit retourner un nombre.

# --hints--

Tu dois créer une fonction nommée `convertCtoF`.

```js
assert.isFunction(convertCtoF);
```

`convertCtoF` doit prendre un seul paramètre.

```js
assert.lengthOf(convertCtoF, 1);
```

`convertCtoF(0)` doit retourner un nombre.

```js
assert.isNumber(convertCtoF(0));
```

`convertCtoF(-30)` doit retourner la valeur `-22`.

```js
assert.strictEqual(convertCtoF(-30), -22);
```

`convertCtoF(-10)` doit retourner la valeur `14`.

```js
assert.strictEqual(convertCtoF(-10), 14);
```

`convertCtoF(0)` doit retourner la valeur `32`.

```js
assert.strictEqual(convertCtoF(0), 32);
```

`convertCtoF(20)` doit retourner la valeur `68`.

```js
assert.strictEqual(convertCtoF(20), 68);
```

`convertCtoF(30)` doit retourner la valeur `86`.

```js
assert.strictEqual(convertCtoF(30), 86);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function convertCtoF(celsius) {
  let fahrenheit = celsius * (9 / 5) + 32;
  return fahrenheit;
}

convertCtoF(30);
```
