---
id: 0095d0567b6817114989a1dd
title: Étape 1
challengeType: 1
dashedName: step-1
---

# --description--

Dans cet atelier, tu vas implémenter l'algorithme du plus court chemin en JavaScript. Tu vas écrire une fonction qui calcule le plus court chemin entre des nœuds dans un graphe pondéré et renvoie à la fois les distances et les chemins empruntés.

Par exemple, pour un graphe où des villes sont reliées par des routes de distances différentes, l'algorithme trouvera la route la plus courte d'une ville à une autre. Si tu veux voyager de la Ville A à la Ville D, l'algorithme pourrait trouver que le trajet A ⇨ B ⇨ C ⇨ D (total : 15 km) est plus court que le trajet direct A ⇨ D (20 km).

Pour commencer, définis une variable nommée `INF` et assigne-lui la valeur `Infinity`. Plus tard, tu l'utiliseras pour indiquer qu'il n'y a pas de connexion directe entre deux nœuds.

# --hints--

Tu dois définir une variable nommée `INF`.

```js
assert.isDefined(INF);
```

Ta variable `INF` doit avoir la valeur `Infinity`.

```js
assert.strictEqual(INF, Infinity);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
