---
id: 69c937c73f23b1140bcd297a
title: Étape 4
challengeType: 1
dashedName: step-4
---

# --description--

Le paramètre `targetNode` est optionnel. Quand il n'est pas fourni, la fonction calculera les plus courts chemins de `startNode` vers tous les autres nœuds du graphe.

Donne au paramètre `targetNode` une valeur par défaut de `null`.

# --hints--

Le paramètre `targetNode` de ta fonction `shortestPath` doit valoir `null` par défaut.

```js
assert.match(
  __helpers.removeJSComments(code),
  /targetNode\s*=\s*null/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode) {

}
--fcc-editable-region--
```
