---
id: 69c918841987bf16c1d1bac5
title: Étape 14
challengeType: 1
dashedName: step-14
---

# --description--

Après avoir parcouru tous les nœuds pour trouver le non visité le plus proche, tu dois gérer le cas où aucun nœud valide n'est trouvé. Cela arrive quand tous les nœuds restants sont inatteignables.

Après la boucle `for` interne, mais toujours dans la boucle externe, ajoute une instruction `if` qui vérifie si `current` est strictement égal à `-1`. Si c'est vrai, utilise `break` dans le bloc `if` pour arrêter tôt la boucle externe.

# --hints--

Tu dois ajouter une instruction `if` qui vérifie si `current === -1`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*current\s*===\s*-1\s*\)/
);
```

Tu dois utiliser `break` à l'intérieur du bloc `if`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*current\s*===\s*-1\s*\)[\s\S]*?break/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
