---
id: 69c918841987bf16c1d1bac8
title: Étape 17
challengeType: 1
dashedName: step-17
---

# --description--

Avant de mettre à jour les distances, tu dois vérifier que le voisin vaut la peine d'être considéré. Il doit y avoir une arête réelle vers lui (`distance !== INF`) et il ne doit pas encore avoir été visité.

À l'intérieur de la boucle `for` des voisins, ajoute une instruction `if` qui vérifie si `distance` n'est pas strictement égal à `INF` **et** si `nodeNo` n'est pas visité.

À l'intérieur du bloc `if`, déclare une variable `newDistance` assignée à `distances[current] + distance`. C'est le coût total pour atteindre le voisin via le nœud courant.

# --hints--

Tu dois ajouter une instruction `if` avec la condition `distance !== INF && !visited[nodeNo]`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*distance\s*!==\s*INF\s*&&\s*!\s*visited\s*\[\s*\w+\s*\]\s*\)/
);
```

Tu dois déclarer une `newDistance` avec la valeur `distances[current] + distance` à l'intérieur du bloc `if`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+newDistance\s*=\s*distances\s*\[\s*current\s*\]\s*\+\s*distance/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      --fcc-editable-region--

      --fcc-editable-region--
    }
  }
}
```
