---
id: 69c918841987bf16c1d1bac3
title: Étape 12
challengeType: 1
dashedName: step-12
---

# --description--

À l'intérieur de la boucle `for` interne, tu dois vérifier si le nœud courant est non visité et plus proche que le meilleur trouvé jusqu'ici.

Ajoute une instruction `if` qui vérifie si `nodeNo` n'a pas été visité et si `distances[nodeNo]` est inférieur à `minDistance`.

Laisse le corps vide pour l'instant.

# --hints--

Tu dois avoir une instruction `if` avec la condition `!visited[nodeNo] && distances[nodeNo] < minDistance` à l'intérieur de la boucle `for` interne.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*!visited\s*\[\s*nodeNo\s*\]\s*&&\s*distances\s*\[\s*nodeNo\s*\]\s*<\s*minDistance\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      --fcc-editable-region--

      --fcc-editable-region--
    }
  }
}
```
