---
id: 69c918841987bf16c1d1baca
title: Étape 19
challengeType: 1
dashedName: step-19
---

# --description--

Quand tu trouves un chemin plus court vers un voisin, tu dois aussi mettre à jour le chemin enregistré pour l'atteindre.

En JavaScript, tu peux créer un nouveau tableau combinant un tableau existant et un nouvel élément avec l'opérateur spread :

```js
const newPath = [...existingArray, newElement];
```

À l'intérieur du bloc `if` imbriqué, mets à jour `paths[nodeNo]` pour qu'il soit le chemin vers le nœud courant avec `nodeNo` ajouté à la fin avec la syntaxe spread comme suit : `[...paths[current], nodeNo]`.

# --hints--

Tu dois mettre à jour `paths[nodeNo]` à l'intérieur du bloc `if` imbriqué.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /paths\s*\[\s*\w+\s*\]\s*=/
);
```

`paths[nodeNo]` doit être défini à `[...paths[current], nodeNo]`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /paths\s*\[\s*\w+\s*\]\s*=\s*\[\s*\.\.\.\s*paths\s*\[\s*current\s*\]\s*,\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        if (newDistance < distances[nodeNo]) {
          distances[nodeNo] = newDistance;
          --fcc-editable-region--

          --fcc-editable-region--
        }
      }
    }
  }
}
```
