---
id: 69c918841987bf16c1d1babb
title: Étape 3
challengeType: 1
dashedName: step-3
---

# --description--

Maintenant tu vas créer la fonction principale qui fait tout le travail.

Crée une fonction nommée `shortestPath` qui prend trois paramètres : `matrix`, `startNode`, et `targetNode`.

Au cours des prochaines étapes, tu vas construire la logique de cette fonction.

# --hints--

Tu dois créer une fonction nommée `shortestPath`.

```js
assert.isFunction(shortestPath);
```

Ta fonction `shortestPath` doit avoir les paramètres `matrix`, `startNode`, et `targetNode`.

```js
const params = __helpers.getFunctionParams(shortestPath.toString());
const names = params.map(p => p.name);
assert.include(names, 'matrix');
assert.include(names, 'startNode');
assert.include(names, 'targetNode');
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--

--fcc-editable-region--
```
