---
id: 69c918841987bf16c1d1baba
title: Étape 2
challengeType: 1
dashedName: step-2
---

# --description--

Tu auras besoin d'un tableau 2D pour représenter la matrice d'adjacence du graphe. Chaque valeur représente le poids (distance) de l'arête entre deux nœuds. Une valeur de `INF` signifie qu'il n'y a pas d'arête directe.

Crée une autre variable nommée `adjMatrix` et assigne-lui un tableau 2D avec les valeurs suivantes :

```js
[0, 5, 3, INF, 11, INF],
[5, 0, 1, INF, INF, 2],
[3, 1, 0, 1, 5, INF],
[INF, INF, 1, 0, 9, 3],
[11, INF, 5, 9, 0, INF],
[INF, 2, INF, 3, INF, 0]
```

# --hints--

Tu dois créer une variable nommée `adjMatrix`.

```js
assert.isDefined(adjMatrix);
```

`adjMatrix` doit être un tableau 2D avec 6 lignes et 6 colonnes.

```js
assert.isArray(adjMatrix);
assert.lengthOf(adjMatrix, 6);
adjMatrix.forEach(row => assert.lengthOf(row, 6));
```

`adjMatrix` doit contenir les valeurs correctes.

```js
assert.deepEqual(adjMatrix[0], [0, 5, 3, INF, 11, INF]);
assert.deepEqual(adjMatrix[1], [5, 0, 1, INF, INF, 2]);
assert.deepEqual(adjMatrix[2], [3, 1, 0, 1, 5, INF]);
assert.deepEqual(adjMatrix[3], [INF, INF, 1, 0, 9, 3]);
assert.deepEqual(adjMatrix[4], [11, INF, 5, 9, 0, INF]);
assert.deepEqual(adjMatrix[5], [INF, 2, INF, 3, INF, 0]);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
--fcc-editable-region--

--fcc-editable-region--
```
