---
id: 69c918841987bf16c1d1bacb
title: Étape 20
challengeType: 1
dashedName: step-20
---

# --description--

Une fois la boucle principale terminée, tu dois décider pour quel(s) nœud(s) afficher les résultats.

Si un `targetNode` spécifique a été fourni, n'affiche les résultats que pour ce nœud. Sinon, affiche les résultats pour tous les nœuds.

En JavaScript, tu peux obtenir tous les indices d'un tableau comme itérable avec `[...Array(n).keys()]`.

Après la boucle `for` externe, crée une variable nommée `targets`. Avec une expression ternaire, si `targetNode` n'est strictement pas `null`, assigne `[targetNode]`, sinon assigne `[...Array(n).keys()]`.

# --hints--

Tu dois créer une variable `targets` après la boucle principale.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+targets\s*=/
);
```

Ta variable `targets` doit utiliser un ternaire vérifiant `targetNode !== null`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /targets\s*=\s*targetNode\s*!==\s*null\s*\?/
);
```

Quand une cible est fournie, `targets` doit être `[targetNode]`, sinon ce doit être `[...Array(n).keys()]`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /\?\s*\[\s*targetNode\s*\]\s*:\s*\[\s*\.\.\.\s*Array\s*\(\s*n\s*\)\s*\.keys\s*\(\s*\)\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        if (newDistance < distances[nodeNo]) {
          distances[nodeNo] = newDistance;
          paths[nodeNo] = [...paths[current], nodeNo];
        }
      }
    }
  }
--fcc-editable-region--

--fcc-editable-region--
}
```
