---
id: 69c918841987bf16c1d1bac9
title: Étape 18
challengeType: 1
dashedName: step-18
---

# --description--

Maintenant tu dois vérifier si passer par le nœud courant donne une route plus courte vers le voisin. Si `newDistance` est meilleur que ce qui est déjà stocké, alors tu dois le mettre à jour.

À l'intérieur du bloc `if` existant, ajoute une instruction `if` imbriquée qui vérifie si `newDistance` est inférieur à `distances[nodeNo]`. Dans ce bloc `if` imbriqué, mets à jour `distances[nodeNo]` à `newDistance`.

# --hints--

Tu dois ajouter une instruction `if` imbriquée avec la condition `newDistance < distances[nodeNo]`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*newDistance\s*<\s*distances\s*\[\s*\w+\s*\]\s*\)/
);
```

Tu dois mettre à jour `distances[nodeNo]` à `newDistance` à l'intérieur du bloc `if` imbriqué.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /distances\s*\[\s*\w+\s*\]\s*=\s*newDistance/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        --fcc-editable-region--

        --fcc-editable-region--
      }
    }
  }
}
```
