---
id: 69c918841987bf16c1d1bac7
title: Étape 16
challengeType: 1
dashedName: step-16
---

# --description--

Maintenant que le nœud courant est marqué comme visité, tu dois regarder tous ses voisins pour voir si tu peux les atteindre plus efficacement via le nœud courant.

Après `visited[current] = true`, ajoute une boucle `for` de voisins qui itère sur chaque `nodeNo` de `0` à `n`. Note que c'est une autre boucle interne au même niveau que la première.

À l'intérieur de la boucle, déclare une variable nommée `distance` et assigne-lui `matrix[current][nodeNo]`. Cela te donne le poids de l'arête entre le nœud courant et son voisin.

# --hints--

Tu dois ajouter une boucle `for` de voisins après avoir marqué le nœud courant comme visité.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /visited\s*\[\s*current\s*\]\s*=\s*true[\s\S]*?for\s*\(/
);
```

La boucle `for` de voisins doit itérer de `0` à `n`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /visited\s*\[\s*current\s*\]\s*=\s*true[\s\S]*?for\s*\(\s*(var|let)\s+\w+\s*=\s*0\s*;\s*\w+\s*<\s*n\s*;\s*\w+\+\+\s*\)/
);
```

Tu dois avoir une variable `distance` assignée à `matrix[current][nodeNo]` à l'intérieur de la boucle des voisins.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+distance\s*=\s*matrix\s*\[\s*current\s*\]\s*\[\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
