---
id: 69c918841987bf16c1d1babf
title: Étape 8
challengeType: 1
dashedName: step-8
---

# --description--

Tu dois aussi suivre le chemin réel emprunté pour atteindre chaque nœud. Tu stockeras une liste de chemins où chaque entrée est un tableau d'indices de nœuds représentant la route empruntée.

Au départ, le chemin de chaque nœud ne contiendra que lui-même. Tu peux créer cette structure avec `Array.from()` :

```js
Array.from({ length: n }, (_, i) => [i]);
```

Cela crée un tableau de `n` éléments, où l'élément à l'indice `i` est `[i]`.

Toujours dans la fonction `shortestPath`, crée une variable nommée `paths` et assigne-lui `Array.from({ length: n }, (_, i) => [i])`.

# --hints--

Tu dois créer une variable `paths` à l'intérieur de la fonction `shortestPath`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+paths\s*=/
);
```

Ta variable `paths` doit être initialisée avec `Array.from`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /paths\s*=\s*Array\.from/
);
```

Ta variable `paths` doit utiliser `{ length: n }` comme premier argument de `Array.from`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /Array\.from\s*\(\s*\{\s*length\s*:\s*n\s*\}/
);
```

Le second argument de `Array.from` doit être un callback `(_, i) => [i]`, où `_` est un placeholder pour l'élément inutilisé et `i` est l'indice. Il doit renvoyer `[i]`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /Array\.from\s*\([\s\S]*?,\s*function\s*\(\s*\w+\s*,\s*(\w+)\s*\)\s*\{[\s\S]*?return\s*\[\s*\1\s*\]|Array\.from\s*\([\s\S]*?,\s*\(\s*\w+\s*,\s*\w+\s*\)\s*=>\s*\[\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  --fcc-editable-region--

  --fcc-editable-region--
}
```
