---
id: 69c918841987bf16c1d1bac2
title: Étape 11
challengeType: 1
dashedName: step-11
---

# --description--

À l'intérieur de la boucle, tu dois parcourir chaque nœud pour trouver celui avec la plus petite distance connue qui n'a pas encore été visité.

Ajoute une boucle `for` à l'intérieur de la boucle principale qui itère sur chaque nœud. Utilise `nodeNo` comme variable de boucle allant de `0` à `n`.

Laisse le corps vide pour l'instant.

# --hints--

Tu dois avoir une seconde boucle `for` (interne) à l'intérieur de la boucle principale avec `nodeNo` comme variable de boucle.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /for\s*\(\s*(let|var)\s+nodeNo\s*=\s*0\s*;\s*nodeNo\s*<\s*n\s*;\s*nodeNo\+\+\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

--fcc-editable-region--
  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;

  }
--fcc-editable-region--
}
```
