---
id: 69c918841987bf16c1d1bace
title: Étape 23
challengeType: 1
dashedName: step-23
---

# --description--

Maintenant que tu sais que le nœud est atteignable et n'est pas le départ, formate son chemin en une chaîne lisible.

Le tableau `paths[nodeNo]` contient la séquence d'indices de nœuds visités pour atteindre `nodeNo`. Tu peux en faire une chaîne lisible comme `0 -> 2 -> 3` avec la méthode `.join()`.

À l'intérieur de la boucle `for...of`, après l'instruction `if` avec `continue`, déclare une variable nommée `path` et assigne-lui la valeur `paths[nodeNo].join(' -> ')`.

# --hints--

Tu dois avoir une variable nommée `path` à l'intérieur de la boucle `for...of`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /for\s*\(\s*(const|let)\s+nodeNo\s+of\s+targets\s*\)[\s\S]*?(const|let)\s+path\s*=/
);
```

Ta variable `path` doit avoir la valeur `paths[nodeNo].join(' -> ')`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /path\s*=\s*paths\s*\[\s*\w+\s*\]\s*\.\s*join\s*\(\s*['"]\s*->\s*['"]\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        if (newDistance < distances[nodeNo]) {
          distances[nodeNo] = newDistance;
          paths[nodeNo] = [...paths[current], nodeNo];
        }
      }
    }
  }

  const targets = targetNode !== null ? [targetNode] : [...Array(n).keys()];
  for (const nodeNo of targets) {
    if (nodeNo === startNode || distances[nodeNo] === INF) {
      continue;
    }
    --fcc-editable-region--

    --fcc-editable-region--
  }
}
```
