---
id: 69c918841987bf16c1d1babd
title: Étape 6
challengeType: 1
dashedName: step-6
---

# --description--

Tu dois suivre la plus courte distance connue du nœud de départ vers chaque autre nœud. Pour commencer, tu vas supposer que chaque nœud est infiniment loin.

En JavaScript, tu peux créer un tableau prérempli avec une valeur en utilisant `new Array(n).fill(value)`. Par exemple, `new Array(3).fill(0)` crée `[0, 0, 0]`.

À l'intérieur de la fonction `shortestPath`, crée une variable nommée `distances` et assigne-lui `new Array(n).fill(INF)`.

# --hints--

Tu dois créer une variable nommée `distances` à l'intérieur de `shortestPath`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+distances\s*=/
);
```

Ta variable `distances` doit être définie à `new Array(n).fill(INF)`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /new\s+Array\s*\(\s*n\s*\)\s*\.fill\s*\(\s*INF\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;

}
--fcc-editable-region--
```
