---
id: 69c918841987bf16c1d1bac1
title: Étape 10
challengeType: 1
dashedName: step-10
---

# --description--

Maintenant tu vas ajouter la boucle principale qui pilote l'algorithme. Elle s'exécute une fois pour chaque nœud du graphe, en sélectionnant à chaque fois le nœud non visité le plus proche.

À l'intérieur de la fonction `shortestPath`, ajoute une boucle `for` qui s'exécute `n` fois en utilisant `i` comme variable de boucle.

Au début du corps de la boucle, avec `let`, déclare deux variables `minDistance` et `current` définies respectivement à `INF` et `-1`. Elles suivront le nœud non visité le plus proche trouvé à chaque itération.

# --hints--

Tu dois avoir une boucle `for` à l'intérieur de `shortestPath` qui s'exécute `n` fois. N'oublie pas de définir `i` avec `let` à l'intérieur.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /for\s*\(\s*(let|var)\s+i\s*=\s*0\s*;\s*i\s*<\s*n\s*;\s*i\+\+\s*\)/
);
```

Tu dois avoir une variable `minDistance` définie à `INF` à l'intérieur de la boucle.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(let|var)\s+minDistance\s*=\s*INF/
);
```

Tu dois avoir une variable `current` définie à `-1` à l'intérieur de la boucle.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(let|var)\s+current\s*=\s*-1/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);
--fcc-editable-region--

--fcc-editable-region--
}
```
