---
id: 69446f18fdcb11aa516f7c5c
title: Étape 1
challengeType: 1
dashedName: step-1
---

# --description--

Dans cet atelier, tu vas créer une fonction `countup` qui crée un tableau avec les nombres de 1 jusqu'à un nombre donné en utilisant la récursion.

Commence par créer une fonction vide nommée `countup` qui prend un seul paramètre.

# --hints--

Tu dois avoir une fonction nommée `countup`.

```js
assert.isFunction(countup);
```

Ta fonction `countup` doit accepter un seul paramètre.

```js
assert.strictEqual(countup.length, 1);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
