---
id: 694470b9a65b278f981ec62e
title: Étape 5
challengeType: 1
dashedName: step-5
---

# --description--

Enfin, teste ta fonction `countup` en l'appelant avec l'argument `5` à l'intérieur d'une instruction `console.log()`.

La sortie devrait être `[1, 2, 3, 4, 5]`.

# --hints--

Tu dois appeler la fonction `countup` avec l'argument `5` à l'intérieur d'une instruction `console.log()`

```js
assert.deepStrictEqual(countup(5), [1, 2, 3, 4, 5]);
assert.match(code, /console\.log\s*\(\s*countup\s*\(\s*5\s*\)\s*\)/)
```

# --seed--

## --seed-contents--

```js
function countup(number) {
  let countArray = [];
  if (number < 1) {
    return [];
  } else {
    countArray = countup(number - 1);
    countArray.push(number);
    return countArray;
  }
}
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
function countup(number) {
  let countArray = [];
  if (number < 1) {
    return [];
  } else {
    countArray = countup(number - 1);
    countArray.push(number);
    return countArray;
  }
}
console.log(countup(5));
```
