---
id: 69447054012a22f38c944ff0
title: Étape 2
challengeType: 1
dashedName: step-2
---

# --description--

Ajoute une instruction `if` à l'intérieur de ta fonction `countup`.

Vérifie si `number` est inférieur à `1` et renvoie un tableau vide.

Cela servira de cas de base pour ta récursion.

# --hints--

Quand `number` est inférieur à `1`, la fonction doit renvoyer un tableau vide.

```js
assert.deepStrictEqual(countup(0), []);
```

Quand `number` n'est **pas** inférieur à `1`, la fonction **ne doit rien** renvoyer.

```js
assert.isUndefined(countup(1));
```

# --seed--

## --seed-contents--

```js
function countup(number) {
--fcc-editable-region--

--fcc-editable-region--
}
```
