---
id: 694870bcf3c6874112b6f951
title: Étape 4
challengeType: 1
dashedName: step-4
---

# --description--

Pour tous les nombres supérieurs ou égaux à `1`, crée un bloc `else` et, à l'intérieur, appelle la fonction `countup` avec `number - 1` et assigne le résultat à `countArray`.

Ensuite, utilise `.push()` pour ajouter `number` à `countArray`, puis renvoie `countArray`.

Cela garantit que la récursion atteint finalement le cas de base et se termine.

# --hints--

Tu dois assigner le résultat de l'appel à `countup` avec `number - 1` à `countArray`.

```js
assert.match(countup.toString(), /countArray\s*=\s*countup\(\s*number\s*-\s*1\s*\)/);
```

Tu dois pousser `number` dans `countArray`.

```js
assert.match(countup.toString(), /countArray\.push\(\s*number\s*\)/);
```

Tu dois renvoyer le `countArray`.

```js
assert.deepStrictEqual(countup(3), [1, 2, 3]);
```

# --seed--

## --seed-contents--

```js
function countup(number) {
  let countArray = [];
  if (number < 1) {
    return [];
--fcc-editable-region--
  }

--fcc-editable-region--
}
```
