---
id: 6948705160b66cf93a916e29
title: Étape 3
challengeType: 1
dashedName: step-3
---

# --description--

Déclare un tableau vide nommé `countArray` à l'intérieur de la fonction `countup`, juste au-dessus de ton instruction `if`.

Ce tableau stockera la valeur de `number` à chaque appel de fonction.

# --hints--

Tu dois déclarer un tableau vide nommé `countArray`.

```js
assert.match(countup.toString(), /(let|const|var)\s+countArray\s*=\s*\[\s*\]/);
```

# --seed--

## --seed-contents--

```js
function countup(number) {
--fcc-editable-region--

--fcc-editable-region--
  if (number < 1) {
    return [];
  }
}
```
