---
id: 646c48df8674cf2b91020ecc
title: Étape 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Un code de base t'a été fourni. Commence par ajouter un élément `link` dans ton élément `head`. Pour cet élément `link`, définis l'attribut `rel` sur `stylesheet` et l'attribut `href` sur `./styles.css`.

# --hints--

Tu devrais avoir un élément `link`.

```js
assert.exists(document.querySelector("head>link"));
```

Ton élément `link` devrait avoir un attribut `rel` et un attribut `href` définis sur `stylesheet` et `./styles.css`.

```js
const link = document.querySelector("head>link");
assert.equal(link?.getAttribute("rel"), "stylesheet");
const href = link?.getAttribute("data-href");
console.log(href)
assert.oneOf(href, ["./styles.css", "styles.css"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
</head>
--fcc-editable-region--
<body></body>
</html>
```

```css

```
