---
id: 646c5ffef5598d449b52ec12
title: Étape 17
challengeType: 0
dashedName: step-17
---

# --description--

Tu devrais maintenant travailler sur les oreilles du chat. Il y aura une oreille droite et une oreille gauche, et chacune contiendra une oreille interne.

Dans ton élément `.cat-head`, crée un élément `div` avec la classe `cat-ears`.

# --hints--

Tu ne devrais pas modifier l'élément `div` existant avec la classe `cat-head`.

```js
assert.lengthOf(document.querySelectorAll('div.cat-head'), 1);
```

Tu devrais créer un élément `div` dans ton élément `.cat-head`.

```js
assert.lengthOf(document.querySelectorAll('.cat-head div'), 1);
```

Ton élément `div` devrait avoir la classe `cat-ears`.

```js
assert.equal(document.querySelectorAll('.cat-head div')[0]?.getAttribute('class'),'cat-ears');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      --fcc-editable-region--
      <div class="cat-head">

      </div>
      --fcc-editable-region--
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
