---
id: 646c5e727e56e743c9aed4a1
title: Étape 7
challengeType: 0
dashedName: step-7
---

# --description--

Le positionnement CSS te permet de définir comment un élément doit être placé dans le navigateur. Il utilise une propriété `position` que tu peux définir sur `static`, `absolute`, `relative`, `sticky` ou `fixed`.

Une fois la propriété `position` de l'élément définie, tu peux déplacer l'élément en donnant une valeur en pixels ou en pourcentage à une ou plusieurs des propriétés `top`, `right`, `left` ou `bottom`.

`static` est le positionnement par défaut de tous les éléments. Si tu l'attribues à un élément, tu ne pourras pas le déplacer avec `top`, `right`, `left` ou `bottom`.

Donne à `.cat-head` une propriété `position` de `static`, puis définis les propriétés `top` et `left` sur `100px` chacune.


# --hints--

Ton sélecteur `.cat-head` devrait avoir une propriété `position` définie sur `static`. Assure-toi d'ajouter un point-virgule.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'static')
```

Ton sélecteur `.cat-head` devrait avoir une propriété `top` définie sur `100px`. Assure-toi d'ajouter un point-virgule.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '100px')
```

Ton sélecteur `.cat-head` devrait avoir une propriété `left` définie sur `100px`. Assure-toi d'ajouter un point-virgule.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left, '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {



  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
--fcc-editable-region--
```
