---
id: 646c5d7057c45f432fcdd46c
title: Étape 5
challengeType: 0
dashedName: step-5
---

# --description--

À l'aide d'un sélecteur de classe, donne à l'élément `.cat-head` une largeur de `205px` et une hauteur de `180px`. Donne-lui aussi une bordure de `1px solid #000` et un `border-radius` de `46%`.

# --hints--

Tu devrais avoir un sélecteur `.cat-head`.

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.cat-head'))
```

Ton sélecteur `.cat-head` devrait avoir un `width` défini sur `205px`.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.width, '205px');
```

Ton sélecteur `.cat-head` devrait avoir un `height` défini sur `180px`.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.height,'180px')
```

Ton sélecteur `.cat-head` devrait avoir un `border` défini sur `1px solid #000`.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.border, '1px solid rgb(0, 0, 0)')
```

Ton sélecteur `.cat-head` devrait avoir un `border-radius` défini sur `46%`.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle(".cat-head")?.borderRadius, '46%')
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  <main>
    <div class='cat-head'></div>
  </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--

--fcc-editable-region--
```
