---
id: 646c59e942f35541923104bf
title: Étape 4
challengeType: 0
dashedName: step-4
---

# --description--

De retour dans ton HTML, crée un élément `main`. Dans cet élément `main`, ajoute un élément `div` avec la classe `cat-head`.

# --hints--

Tu devrais avoir un élément `main`.

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

Tu devrais avoir un élément `div`.

```js
assert.lengthOf(document.querySelectorAll('div'), 1);
```

Ton élément `div` devrait avoir la classe `cat-head`.

```js
assert.equal(document.querySelector('div')?.getAttribute('class'), 'cat-head');
```

Ton élément `div` devrait être dans ta balise `main`.

```js
assert.lengthOf(document.querySelectorAll('main div'), 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  --fcc-editable-region--

  --fcc-editable-region--
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}
```
