---
id: 646cf1206cac5f51804f49cf
title: Étape 25
challengeType: 0
dashedName: step-25
---

# --description--

Pour positionner correctement l'oreille gauche, tu peux utiliser la transformation CSS pour la faire pivoter d'un certain angle.

La propriété `transform` te permet de modifier la forme, la position et la taille d'un élément sans changer la mise en page ni affecter les éléments autour. Elle propose des fonctions comme `translate()`, `rotate()`, `scale()`, `skew()` et `matrix()`.

Définis la propriété `transform` sur `rotate(-45deg)` et observe ce qui se passe.

# --hints--

Tu devrais définir la propriété `transform` de ton élément `.cat-left-ear` sur `rotate(-45deg)`. N'oublie pas d'ajouter un point-virgule.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.transform, 'rotate(-45deg)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.cat-left-ear {
  position: absolute;
  top: -26px;
  left: -31px;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  --fcc-editable-region--

  --fcc-editable-region--
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}
```
