---
id: 646c5ace05e4be4211407935
title: Étape 6
challengeType: 0
dashedName: step-6
---

# --description--

Pour voir l'élément `.cat-head`, donne-lui un arrière-plan en dégradé linéaire avec `#5e5e5e` a `85%` et `#45454f` a `100%`.

Tu ne remarqueras peut-être pas la difference entre ces deux couleurs, mais elles sont bien présentes.

# --hints--

Ton sélecteur `.cat-head` devrait avoir une propriété `background`.

```js
assert.match(code, /background:/)
```

Ta propriété `background` devrait utiliser la fonction `linear-gradient`.

```js
assert.match(code, /background:\s*linear-gradient\(/)
```

Ta fonction `linear-gradient` devrait définir la premiere couleur sur `#5e5e5e` a `85%`.

```js
const gradientBackgroundImage = new __helpers.CSSHelp(document).getStyle('.cat-head')?.getPropVal('background-image', true);
const firstGradient = gradientBackgroundImage.split(/\s*rgb\s*/)[1];
assert.equal(firstGradient,'(94,94,94)85%,');
```

Ta fonction `linear-gradient` devrait définir la deuxième couleur sur `#45454f` a `100%`.

```js
const gradientBackgroundImage = new __helpers.CSSHelp(document).getStyle('.cat-head')?.getPropVal('background-image', true);
const secondGradient = gradientBackgroundImage.split(/\s*rgb\s*/)[2];
assert.equal(secondGradient,'(69,69,79)100%)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;

}
--fcc-editable-region--
```
