---
id: 646ceb843412c74edee27a79
title: Étape 20
challengeType: 0
dashedName: step-20
---

# --description--

À l'aide d'un sélecteur de classe, donne à l'élément `.cat-right-ear` des propriétés `height` et `width` définies sur `100px`. Définis `background-color` sur `white`. Définis les bordures gauche et droite sur `35px solid blue`. Définis les bordures du haut et du bas sur `35px solid red`.

# --hints--

Tu devrais avoir un sélecteur `.cat-right-ear`.

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

Ton sélecteur `.cat-right-ear` devrait avoir une propriété `height` définie sur `100px`. N'oublie pas d'ajouter un point-virgule.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height, '100px')
```

Ton sélecteur `.cat-right-ear` devrait avoir une propriété `width` définie sur `100px`. N'oublie pas d'ajouter un point-virgule.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width, '100px')
```

Ton sélecteur `.cat-right-ear` devrait avoir une propriété `background-color` définie sur `white`. N'oublie pas d'ajouter un point-virgule.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor, 'white')
```

Ton sélecteur `.cat-right-ear` devrait avoir une propriété `border-left` définie sur `35px solid blue`. N'oublie pas d'ajouter un point-virgule.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft, '35px solid blue')
```

Ton sélecteur `.cat-right-ear` devrait avoir une propriété `border-right` définie sur `35px solid blue`. N'oublie pas d'ajouter un point-virgule.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight, '35px solid blue')
```

Ton sélecteur `.cat-right-ear` devrait avoir une propriété `border-top` définie sur `35px solid red`. N'oublie pas d'ajouter un point-virgule.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop, '35px solid red')
```

Ton sélecteur `.cat-right-ear` devrait avoir une propriété `border-bottom` définie sur `35px solid red`. N'oublie pas d'ajouter un point-virgule.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom, '35px solid red')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```
