---
id: 646ce9d790d2a44de5f99e04
title: Étape 19
challengeType: 0
dashedName: step-19
---

# --description--

Crée deux éléments `div` : le premier dans l'élément `.cat-left-ear` avec une classe `cat-left-inner-ear`, et le second dans l'élément `.cat-right-ear` avec une classe `cat-right-inner-ear`.

# --hints--

Tu ne devrais pas modifier l'élément `div` existant avec la classe `cat-left-ear`.

```js
assert.lengthOf(document.querySelectorAll('div.cat-left-ear'), 1);
```

Tu ne devrais pas modifier l'élément `div` existant avec la classe `cat-right-ear`.

```js
assert.lengthOf(document.querySelectorAll('div.cat-right-ear'), 1);
```

Tu devrais avoir un élément `div` dans ton élément `.cat-left-ear`.

```js
assert.lengthOf(document.querySelectorAll('.cat-left-ear div'), 1);
```

Tu devrais avoir un élément `div` dans ton élément `.cat-right-ear`.

```js
assert.lengthOf(document.querySelectorAll('.cat-right-ear div'), 1);
```

Le nouvel élément `div` dans `.cat-left-ear` devrait avoir la classe `cat-left-inner-ear`.

```js
assert.isTrue(document.querySelectorAll('.cat-left-ear div')[0]?.classList.contains('cat-left-inner-ear'));
```

Le nouvel élément `div` dans `.cat-right-ear` devrait avoir la classe `cat-right-inner-ear`.

```js
assert.isTrue(document.querySelectorAll('.cat-right-ear div')[0]?.classList.contains('cat-right-inner-ear'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          --fcc-editable-region--
          <div class="cat-left-ear">

          </div>
          <div class="cat-right-ear">

          </div>
          --fcc-editable-region--
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
