---
id: 646ddab8afd73764f5241bbf
title: Étape 38
challengeType: 0
dashedName: step-38
---

# --description--

Comme tu l'as fait pour l'oreille interne gauche, supprime les bords pointus de l'oreille interne droite en définissant un rayon de bordure inférieur droit et inférieur gauche de `40%`, un rayon supérieur gauche de `90px` et un rayon supérieur droit de `10px`.

# --hints--

Ton sélecteur `.cat-right-inner-ear` devrait avoir une propriété `border-bottom-right-radius` définie sur `40%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-right-inner-ear')?.borderBottomRightRadius, '40%');
```

Ton sélecteur `.cat-right-inner-ear` devrait avoir une propriété `border-bottom-left-radius` définie sur `40%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-right-inner-ear')?.borderBottomLeftRadius, '40%');
```

Ton sélecteur `.cat-right-inner-ear` devrait avoir une propriété `border-top-left-radius` définie sur `90px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-right-inner-ear')?.borderTopLeftRadius, '90px');
```

Ton sélecteur `.cat-right-inner-ear` devrait avoir une propriété `border-top-right-radius` définie sur `10px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-right-inner-ear')?.borderTopRightRadius, '10px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.cat-left-ear {
  position: absolute;
  top: -26px;
  left: -31px;
  z-index: 1;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  transform: rotate(-45deg);
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

.cat-right-ear {
  position: absolute;
  top: -26px;
  left: 163px;
  z-index: 1;
  transform: rotate(45deg);
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

.cat-left-inner-ear {
  position: absolute;
  top: 22px;
  left: -20px;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-bottom-right-radius: 40%;
  border-bottom-left-radius: 40%;
  border-left: 20px solid transparent;
  border-right: 20px solid transparent;
  border-bottom: 40px solid #3b3b4f;
}

.cat-right-inner-ear {
  position: absolute;
  top: 22px;
  left: -20px;
  --fcc-editable-region--

  --fcc-editable-region--
  border-left: 20px solid transparent;
  border-right: 20px solid transparent;
  border-bottom: 40px solid #3b3b4f;
}
```
