---
id: 646dec359bef3b7811fba5a6
title: Étape 50
challengeType: 0
dashedName: step-50
---

# --description--

Déplace la partie interne de l'œil droit à sa place avec une `position` de `absolute`, un `top` de `8px` et un `left` de `18px`. Donne-lui aussi un rayon de bordure de `60%` et fais-la pivoter de `-5deg`.

# --hints--

Ton sélecteur `.cat-right-inner-eye` devrait avoir une propriété `position` définie sur `absolute`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-right-inner-eye')?.position, 'absolute');
```

Ton sélecteur `.cat-right-inner-eye` devrait avoir une propriété `top` définie sur ` 8px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-right-inner-eye')?.top, '8px');
```

Ton sélecteur `.cat-right-inner-eye` devrait avoir une propriété `left` définie sur `18px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-right-inner-eye')?.left, '18px');
```

Ton sélecteur `.cat-right-inner-eye` devrait avoir une propriété `border-radius` définie sur `60%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-right-inner-eye')?.borderRadius, '60%');
```

Ton sélecteur `.cat-right-inner-eye` devrait avoir une propriété `transform` définie sur `-5deg`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-right-inner-eye')?.transform, 'rotate(-5deg)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>

        <div class="cat-eyes">
          <div class="cat-left-eye">
            <div class="cat-left-inner-eye"></div>
          </div>
          <div class="cat-right-eye">
            <div class="cat-right-inner-eye"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.cat-left-ear {
  position: absolute;
  top: -26px;
  left: -31px;
  z-index: 1;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  transform: rotate(-45deg);
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

.cat-right-ear {
  position: absolute;
  top: -26px;
  left: 163px;
  z-index: 1;
  transform: rotate(45deg);
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

.cat-left-inner-ear {
  position: absolute;
  top: 22px;
  left: -20px;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-bottom-right-radius: 40%;
  border-bottom-left-radius: 40%;
  border-left: 20px solid transparent;
  border-right: 20px solid transparent;
  border-bottom: 40px solid #3b3b4f;
}

.cat-right-inner-ear {
  position: absolute;
  top: 22px;
  left: -20px;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-bottom-right-radius: 40%;
  border-bottom-left-radius: 40%;
  border-left: 20px solid transparent;
  border-right: 20px solid transparent;
  border-bottom: 40px solid #3b3b4f;
}

.cat-left-eye {
  position: absolute;
  top: 54px;
  left: 39px;
  border-radius: 60%;
  transform: rotate(25deg);
  width: 30px;
  height: 40px;
  background-color: #000;
}

.cat-right-eye {
  position: absolute;
  top: 54px;
  left: 134px;
  border-radius: 60%;
  transform: rotate(-25deg);
  width: 30px;
  height: 40px;
  background-color: #000;
}

.cat-left-inner-eye {
  position: absolute;
  top: 8px;
  left: 2px;
  width: 10px;
  height: 20px;
  transform: rotate(10deg);
  background-color: #fff;
  border-radius: 60%;
}

.cat-right-inner-eye {
  --fcc-editable-region--

  --fcc-editable-region--
  width: 10px;
  height: 20px;
  background-color: #fff;
}
```
