---
id: 6476fc5cf14b276e6d04e82a
title: Étape 11
challengeType: 0
dashedName: step-11
---

# --description--

Utilise un sélecteur de classe pour donner à ton élément `.box` une largeur de `200px`, une hauteur de `600px` et une couleur d'arrière-plan de `#000`. Donne-lui aussi une `position` de `absolute`, un `top` de `800px` et un `left` de `650px`.

# --hints--

Tu devrais avoir un sélecteur `.box`.

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.box'));
```

Ton sélecteur `.box` devrait avoir une propriété `width` définie sur `200px`.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.box')?.width, '200px');
```

Ton sélecteur `.box` devrait avoir une propriété `height` définie sur `600px`.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.box')?.height, '600px');
```

Ton sélecteur `.box` devrait avoir une propriété `background-color` définie sur `#000`.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.box')?.backgroundColor, 'rgb(0, 0, 0)');
```

Ton sélecteur `.box` devrait avoir une propriété `position` définie sur `absolute`.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.box')?.position, 'absolute');
```

Ton sélecteur `.box` devrait avoir une propriété `left` définie sur `650px`.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.box')?.left, '650px');
```

Ton sélecteur `.box` devrait avoir une propriété `top` définie sur `800px`.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.box')?.top, '800px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: relative;
  top: 100px;
  left: 100px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```
