---
id: 6476f5c17f99146071ee884c
title: Étape 8
challengeType: 0
dashedName: step-8
---

# --description--

Tu as pu voir que rien ne se passe. L'élément `.cat-head` n'a pas bouge même avec un `top` et un `left` de `100px` chacun. Mais ce n'est pas le cas avec le positionnement `relative`.

Quand tu utilises la valeur `relative`, l'élément reste positionne selon le flux normal du document, mais les valeurs `top`, `left`, `bottom` et `right` deviennent actives.

Au lieu de `static`, donne à ton `.cat-head` une position de `relative`, et laisse les propriétés `top` et `left` telles quelles.

# --hints--

Ton sélecteur `.cat-head` devrait avoir une propriété `position` définie sur `relative`. Assure-toi d'ajouter un point-virgule.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'relative');
```

Ton sélecteur `.cat-head` devrait avoir une propriété `top` définie sur `100px`. Assure-toi d'ajouter un point-virgule.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '100px');
```

Ton sélecteur `.cat-head` devrait avoir une propriété `left` définie sur `100px`. Assure-toi d'ajouter un point-virgule.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left, '100px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: static;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
