---
id: 6476f7a4827bcc61682f2347
title: Étape 9
challengeType: 0
dashedName: step-9
---

# --description--

Quand tu utilises la valeur `absolute` pour ta propriété `position`, l'élément est retire du flux normal du document, puis sa position est determinee par les propriétés `top`, `right`, `bottom` et `left`.

Définis la propriété `position` de ton élément `.cat-head` sur `absolute`, puis définis les propriétés `top` et `left` sur n'importe quelle valeur positive en pixels.

<!-- **Note** : Tu peux experimenter ici avec les propriétés `top`, `left`, `bottom` et `right`, mais le test ne passera qu'avec un `top` de `300px` et un `left` de `400px`. -->

# --hints--

Ton sélecteur `.cat-head` devrait avoir une propriété `position` définie sur `absolute`. Assure-toi d'ajouter un point-virgule.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'absolute');
```

Ton sélecteur `.cat-head` devrait avoir une propriété `top` définie sur n'importe quel nombre en `px`. Assure-toi d'ajouter un point-virgule.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '300px')
assert.match(code, /top:\s*\d{1,4}px;/);
```

Ton sélecteur `.cat-head` devrait avoir une propriété `left` définie sur n'importe quel nombre en `px`. Assure-toi d'ajouter un point-virgule.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '400px')
assert.match(code, /left:\s*\d{1,4}px;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
