---
id: 6476fb7cbaafa36d65e9cf35
title: Étape 10
challengeType: 0
dashedName: step-10
---

# --description--

`fixed` est une valeur de la propriété `position` qui te permet de fixer un élément sur la page, peu importe l'endroit ou l'utilisateur fait defiler la page.

Tu vas devoir ajouter un peu plus de balisage pour voir comment fonctionne le positionnement `fixed`. Dans ton HTML, crée un élément `div` avec la classe `box`.

# --hints--

Tu devrais avoir un élément `div`.

```js
assert.lengthOf(document.querySelectorAll('div'), 2);
```

Ton élément `div` devrait avoir la classe `box`.

```js
assert.equal(document.querySelectorAll('div')[1]?.getAttribute('class'), 'box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      --fcc-editable-region--

      --fcc-editable-region--
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  top: 300px;
  left: 400px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
