---
id: 64770351e8586671ec0911f0
title: Étape 13
challengeType: 0
dashedName: step-13
---

# --description--

Le positionnement `sticky` est un melange entre le positionnement `relative` et le positionnement `fixed`. Il permet à un élément de **rester collé** à une position précise dans son élément parent ou dans la fenêtre d'affichage, selon la position de défilement.

Change la valeur de la propriété `position` de `.cat-head` en `sticky`, définis `top` sur `0`, puis supprime `left` et sa valeur.

**Note** : Pour voir comment `sticky` fonctionne, tu dois placer quelques textes avant et apres ton élément `.cat-head`. Si tu fais ensuite defiler vers le bas, tu verras que `.cat-head` reste collé en haut.

# --hints--

Ton sélecteur `.cat-head` devrait avoir une propriété `position` définie sur `sticky`. Assure-toi d'ajouter un point-virgule.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'sticky');
```

Ton sélecteur `.cat-head` devrait avoir une propriété `top` définie sur `0`. Assure-toi d'ajouter un point-virgule.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '0px');
```

Tu ne devrais pas avoir la propriété `left` ni sa valeur dans ton code.

```js
assert.notMatch(code, /left:\s*100px;?/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: fixed;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 200px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
