---
id: 6477062778c85972eb648030
title: Étape 14
challengeType: 0
dashedName: step-14
---

# --description--

Tu devrais maintenant centrer la tete du chat.

Donne à l'élément `.cat-head` une propriété `position` définie sur `absolute`. Définis une valeur de `0` pour les propriétés `right`, `left`, `top` et `bottom`, puis définis sa propriété `margin` sur `auto` de tous les cotes. C'est une facon de centrer un élément verticalement et horizontalement avec le positionnement CSS.

# --hints--

Ton sélecteur `.cat-head` devrait avoir une propriété `position` définie sur `absolute`. Assure-toi d'ajouter un point-virgule.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'absolute');
```

Ton sélecteur `.cat-head` devrait avoir une propriété `top` définie sur `0`. Assure-toi d'ajouter un point-virgule.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '0px');
```

Ton sélecteur `.cat-head` devrait avoir une propriété `left` définie sur `0`. Assure-toi d'ajouter un point-virgule.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left, '0px');
```

Ton sélecteur `.cat-head` devrait avoir une propriété `right` définie sur `0`. Assure-toi d'ajouter un point-virgule.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.right, '0px');
```

Ton sélecteur `.cat-head` devrait avoir une propriété `bottom` définie sur `0`. Assure-toi d'ajouter un point-virgule.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.bottom, '0px');
```

Ton sélecteur `.cat-head` devrait avoir une propriété `margin` définie sur `auto`. Assure-toi d'ajouter un point-virgule.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.margin, 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: sticky;
  top: 0;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 400px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
