---
id: 647d855ac12cd436059acd39
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

Utilise le sélecteur universel pour ajouter `box-sizing: border-box;` à ton CSS. Cela garantit que les éléments incluent le remplissage et la bordure dans leur largeur et leur hauteur définies.

# --hints--

Tu devrais avoir le sélecteur universel (`*`).

```js
// assert.match(code, /\*\s*{\s*.*\s*}/)
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

Ton sélecteur universel devrait avoir une propriété `box-sizing` définie sur `border-box`.

```js
// assert.match(code, /\*\s*{\s*box-sizing:\s*border-box;?\s*}/)
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing, 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body></body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
