---
id: 6676a8a8d9861319ee901a90
title: Étape 22
challengeType: 0
dashedName: step-22
---

# --description--

Tu peux maintenant commencer a travailler sur les oreilles de ton chat.

Nettoie ton code de test en supprimant le sélecteur `.cat-right-ear` et toutes ses propriétés.

À l'aide d'un sélecteur de classe, donne à l'élément `.cat-left-ear` une bordure gauche et une bordure droite de `35px solid transparent` chacune. Définis aussi la bordure du bas sur `70px solid #5e5e5e`.

# --hints--


Ton sélecteur `.cat-right-ear` ne devrait pas avoir de propriété `height`.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height)
}
```

Ton sélecteur `.cat-right-ear` ne devrait pas avoir de propriété `width`.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width)
}
```


Ton sélecteur `.cat-right-ear` ne devrait pas avoir de propriété `border-left`.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft)
}
```

Ton sélecteur `.cat-right-ear` ne devrait pas avoir de propriété `border-right`.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight)
}
```

Ton sélecteur `.cat-right-ear` ne devrait pas avoir de propriété `border-top`.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop)
}
```

Ton sélecteur `.cat-right-ear` ne devrait pas avoir de propriété `border-bottom`.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom)
}
```

Tu ne devrais pas avoir de sélecteur `.cat-right-ear`.

```js
assert.notExists(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

Tu devrais avoir un sélecteur `.cat-left-ear`.

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.cat-left-ear'))
```

Ton sélecteur `.cat-left-ear` devrait avoir une propriété `border-left` définie sur `35px solid transparent`

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderLeft, '35px solid transparent')
```

Ton sélecteur `.cat-left-ear` devrait avoir une propriété `border-right` définie sur `35px solid transparent`

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderRight, '35px solid transparent')
```

Ton sélecteur `.cat-left-ear` devrait avoir une propriété `border-bottom` définie sur `70px solid #5e5e5e`

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderBottom, '70px solid rgb(94, 94, 94)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--
.cat-right-ear {
  height: 0;
  width: 0;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-top: 35px solid transparent;
  border-bottom: 35px solid red;
}
--fcc-editable-region--
```
