---
id: 6676a8b01e56ec1a1e07c202
title: Étape 21
challengeType: 0
dashedName: step-21
---

# --description--

Observe que tu as maintenant un carré blanc avec des bordures épaisses. Les bordures ont des bords diagonaux qui peuvent servir a créer des triangles. Tu en auras besoin pour créer les oreilles de ton chat.

Dans le même sélecteur de classe, ajuste `height` et `width` a `0`. Définis les bordures gauche, droite et du haut sur `transparent`.

Supprime la propriété `background-color`, et tu devrais voir un triangle.

# --hints--


Tu devrais avoir un sélecteur `.cat-right-ear`.

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

Ton sélecteur `.cat-right-ear` devrait avoir une propriété `height` définie sur `0`. N'oublie pas d'ajouter un point-virgule.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height, '0px')
```

Ton sélecteur `.cat-right-ear` devrait avoir une propriété `width` définie sur `0`. N'oublie pas d'ajouter un point-virgule.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width, '0px')
```

Ton sélecteur `.cat-right-ear` devrait avoir une propriété `border-left` définie sur `35px solid transparent`. N'oublie pas d'ajouter un point-virgule.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft, '35px solid transparent')
```

Ton sélecteur `.cat-right-ear` devrait avoir une propriété `border-right` définie sur `35px solid transparent`. N'oublie pas d'ajouter un point-virgule.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight, '35px solid transparent')
```

Ton sélecteur `.cat-right-ear` devrait avoir une propriété `border-top` définie sur `35px solid transparent`. N'oublie pas d'ajouter un point-virgule.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop, '35px solid transparent')
```


Ton sélecteur `.cat-right-ear` devrait avoir une propriété `border-bottom` définie sur `35px solid red`. N'oublie pas d'ajouter un point-virgule.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom, '35px solid red')
```

Ton sélecteur `.cat-right-ear` ne devrait pas avoir de propriété `background-color`.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--
.cat-right-ear {
  height: 100px;
  width: 100px;
  background-color: white;
  border-left: 35px solid blue;
  border-right: 35px solid blue;
  border-top: 35px solid red;
  border-bottom: 35px solid red;
}
--fcc-editable-region--
```
