---
id: 691f7bbf76229bb97c60827d
title: Étape 1
challengeType: 1
dashedName: step-1
---

# --description--

Dans ce workshop, tu vas t'entraîner à utiliser les boucles `for...of` en construisant une fonction qui compte combien de fois une chaîne apparaît dans un tableau de chaînes.

Avant d'écrire cette fonction, tu vas en créer une plus simple qui affiche chaque caractère d'une chaîne dans la console.

Commence par définir une fonction vide appelée `printCharacters` avec le paramètre `str`.

# --hints--

Tu dois créer une fonction `printCharacters`.

```js
assert.isFunction(printCharacters);
```

Ta fonction `printCharacters` doit avoir un paramètre appelé `str`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const functionRegex = __helpers.functionRegex('printCharacters', ['str']);

assert.match(codeWithoutJSComments, functionRegex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
