---
id: 6920d189e36b6b60a6330fee
title: Étape 3
challengeType: 1
dashedName: step-3
---

# --description--

Maintenant, à l'intérieur de la boucle, affiche la variable `char` dans la console.

# --hints--

Tu dois avoir une instruction `console.log` avec `char` comme argument à l'intérieur de la boucle.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const functionRegex = __helpers.functionRegex('printCharacters', ['str'], { capture: true });
const functionCode = codeWithoutJSComments.match(functionRegex)[1];
const normalizedFunctionCode = __helpers.removeWhiteSpace(functionCode);
const loopCodeBlock = normalizedFunctionCode.match(/for\(constcharofstr\)\{(\S*?)}/)[1]

assert.match(loopCodeBlock, /console\.log\(char\)/);
```

Lorsque la fonction `printCharacters` est appelée avec l'argument `'test'`, l'instruction `console.log()` à l'intérieur de la fonction doit être appelée quatre fois avec les arguments suivants, dans l'ordre : `'t'`, `'e'`, `'s'` et `'t'`. Au lieu de cela, elle a été appelée avec les arguments suivants : `--fcc-actual--`.

```js
const consoleLogSpy = __helpers.spyOn(console, 'log');
printCharacters('test');
assert.deepEqual(consoleLogSpy.calls, [['t'], ['e'], ['s'], ['t']]);
consoleLogSpy.restore();
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
