---
id: 6920c9cda1c7f056f75e4b1a
title: Étape 2
challengeType: 1
dashedName: step-2
---

# --description--

Maintenant, il est temps d'écrire ta première boucle `for...of`. Dans les leçons précédentes, tu as appris que tu peux utiliser une boucle `for...of` pour itérer sur les valeurs d'un itérable (par exemple, des chaînes ou des tableaux).

Voici un exemple de boucle `for...of` :

```js
for (const num of [1, 2, 3]) {
  // code block to be executed
}
```

Ajoute une boucle `for...of` avec un bloc de code vide à l'intérieur de ta fonction. Elle doit itérer sur chaque caractère de l'argument `str`.

# --hints--

Ta fonction `printCharacters` doit avoir une boucle `for...of` à l'intérieur de son bloc de code.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);

assert.match(normalizedCode, /functionprintCharacters\(str\)\{for\(\S+of\S+\)\{\S*?\}\}/);
```

Ta boucle `for...of` doit itérer sur l'argument `str` de ta fonction.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const forLoopRegex = /function printCharacters\(str\) \{\s*for\s*\(([\S\s]*?)\)\s*\{[\S\s]*?\}\s*\}/;
const forLoopCode = codeWithoutJSComments.match(forLoopRegex)[1];

assert.match(forLoopCode, /(const|let|var)\s+\w+\s+of\s+str/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
--fcc-editable-region--

--fcc-editable-region--
}
```
