---
id: 6921de8d2a111d0341cff702
title: Étape 8
challengeType: 1
dashedName: step-8
---

# --description--

Super, ta fonction retourne `0`, mais elle ne compte encore rien.

Pour corriger cela, à l'intérieur de `getMatchedWordCount`, crée une boucle `for...of` avec un bloc de code vide qui itère sur chaque mot de `sentence`.

# --hints--

Ta fonction `getMatchedWordCount` doit avoir une boucle `for...of` à l'intérieur de son bloc de code.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);

assert.match(normalizedCode, /functiongetMatchedWordCount\(sentence,match\)\{letcount=0\;for\(\S+of\S+\)\{\S*?\}returncount;\}/);
```

Ta boucle `for...of` doit itérer sur l'argument `sentence` de ta fonction.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const forLoopRegex = /function getMatchedWordCount\(sentence, match\) \{\s*let count = 0\;\s*for\s*\(([\S\s]*?)\)\s*\{[\S\s]*?\}\s*return count\;\s*\}/;
const forLoopCode = codeWithoutJSComments.match(forLoopRegex)[1];

assert.match(forLoopCode, /(const|let|var)\s+\w+\s+of\s+sentence/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
--fcc-editable-region--

--fcc-editable-region--
  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);
```
