---
id: 6921d90b5ce50185505d2004
title: Étape 6
challengeType: 1
dashedName: step-6
---

# --description--

Comme la fonction `getMatchedWordCount` doit retourner un décompte numérique, tu vas commencer par mettre en place une variable compteur.

Crée une variable nommée `count` en utilisant le mot-clé `let`, initialise-la avec la valeur `0`, et ajoute une instruction `return` qui retourne la variable `count`.

# --hints--

Tu dois déclarer la variable `count` à l'intérieur de la fonction `getMatchedWordCount`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const functionRegex = __helpers.functionRegex('getMatchedWordCount', ['sentence', 'match'], { capture: true });
const functionCode = codeWithoutJSComments.match(functionRegex)[1];
const normalizedFunctionCode = __helpers.removeWhiteSpace(functionCode);

assert.match(normalizedFunctionCode, /(const|let|var)count/);
```

Pour permettre à la variable `count` d'être incrémentée plus tard, déclare-la avec le mot-clé `let`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const functionRegex = __helpers.functionRegex('getMatchedWordCount', ['sentence', 'match'], { capture: true });
const functionCode = codeWithoutJSComments.match(functionRegex)[1];
const normalizedFunctionCode = __helpers.removeWhiteSpace(functionCode);

assert.match(normalizedFunctionCode, /letcount/);
```

Ta fonction doit contenir une affectation de la valeur `0` à la variable `count`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const functionRegex = __helpers.functionRegex('getMatchedWordCount', ['sentence', 'match'], { capture: true });
const functionCode = codeWithoutJSComments.match(functionRegex)[1];
const normalizedFunctionCode = __helpers.removeWhiteSpace(functionCode);

assert.match(normalizedFunctionCode, /count=0/);
```

Ta fonction doit contenir une instruction `return` qui retourne la variable `count`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const functionRegex = __helpers.functionRegex('getMatchedWordCount', ['sentence', 'match'], { capture: true });
const functionCode = codeWithoutJSComments.match(functionRegex)[1];
const normalizedFunctionCode = __helpers.removeWhiteSpace(functionCode);

assert.match(normalizedFunctionCode, /returncount/);
```

Ta fonction `getMatchedWordCount` doit retourner la valeur `0`.

```js
assert.strictEqual(getMatchedWordCount(["foo", "bar"], "foo"), 0);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
--fcc-editable-region--

--fcc-editable-region--
}
```
