---
id: 6921d688b8d58581d34459b9
title: Étape 5
challengeType: 1
dashedName: step-5
---

# --description--

Dans les prochaines étapes, tu vas construire une fonction qui compte combien de fois une chaîne apparaît dans un tableau de chaînes.

Pour commencer, définis une fonction vide nommée `getMatchedWordCount` avec les paramètres `sentence` et `match` dans cet ordre.

# --hints--

Tu dois créer une fonction `getMatchedWordCount`.

```js
assert.isFunction(getMatchedWordCount);
```

Ta fonction `getMatchedWordCount` doit avoir deux paramètres : `sentence` et `match`, dans cet ordre.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const functionRegex = __helpers.functionRegex('getMatchedWordCount', ['sentence', 'match']);

assert.match(codeWithoutJSComments, functionRegex);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

--fcc-editable-region--

--fcc-editable-region--
```
