---
id: 6921dae10700e300f26994e3
title: Étape 7
challengeType: 1
dashedName: step-7
---

# --description--

Pour voir la sortie de ta fonction `getMatchedWordCount` dans son état actuel, ajoute une instruction `console.log` en dessous. À l'intérieur du `console.log()`, appelle la fonction avec l'argument `["I", "really", "really", "really", "like", "to", "code"]` pour le paramètre `sentence` et `"really"` pour le paramètre `match`.

# --hints--

Tu dois appeler la fonction `getMatchedWordCount` à l'intérieur d'une instruction `console.log`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
assert.match(normalizedCode, /console\.log\(getMatchedWordCount\(\S*?\)\)/);
```

Tu dois appeler la fonction `getMatchedWordCount` avec l'argument `["I", "really", "really", "really", "like", "to", "code"]` pour le paramètre `sentence` et `"really"` pour le paramètre `match`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const normalizedConsoleLogCode = normalizedCode.match(/console\.log\(getMatchedWordCount\(\S*?\)\)/)[0];
assert.match(normalizedConsoleLogCode, /\[('|"|`)I\1,('|"|`)really\2,('|"|`)really\3,('|"|`)really\4,('|"|`)like\5,('|"|`)to\6,('|"|`)code\7\],('|"|`)really\8/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
