---
id: 69220a03729c43082ad3e9f6
title: Étape 11
challengeType: 1
dashedName: step-11
---

# --description--

Tu as fini de travailler sur ta fonction `getMatchedWordCount` !

Maintenant, tu vas la tester avec différentes données pour voir comment elle se comporte.

Ajoute une nouvelle instruction `console.log` en dessous de celle qui existe déjà, qui appelle `getMatchedWordCount` avec ces paramètres :

- sentence : `["Do", "not", "fear", "the", "dandy", "lion"]`
- match : `"dandy"`

Félicitations ! Tu as terminé ce workshop.

# --hints--

Tu dois avoir une deuxième instruction `console.log` qui appelle la fonction `getMatchedWordCount` à l'intérieur.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const normalizedConsoleLogCode = normalizedCode.match(/console\.log\(getMatchedWordCount\(\S*?\)\)/g)[1];
assert.exists(normalizedConsoleLogCode);
```

Dans ta deuxième instruction `console.log`, appelle la fonction `getMatchedWordCount` avec `["Do", "not", "fear", "the", "dandy", "lion"]` comme argument `sentence` et `"dandy"` comme argument `match`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const normalizedConsoleLogCode = normalizedCode.match(/console\.log\(getMatchedWordCount\(\S*?\)\)/g)[1];
assert.match(normalizedConsoleLogCode, /\[('|"|`)Do\1,('|"|`)not\2,('|"|`)fear\3,('|"|`)the\4,('|"|`)dandy\5,('|"|`)lion\6\],('|"|`)dandy\7/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;

  for (const word of sentence) {
    if (word === match) {
      count++;
    }
    console.log(`Checking "${word}" against "${match}" | Running count: ${count}`);
  }

  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;

  for (const word of sentence) {
    if (word === match) {
      count++;
    }
    console.log(`Checking "${word}" against "${match}" | Running count: ${count}`);
  }

  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);

console.log(
  getMatchedWordCount(
    ["Do", "not", "fear", "the", "dandy", "lion"],
    "dandy"
  )
);
```
