---
id: 6978bd465213832659f3338b
title: Étape 9
challengeType: 1
dashedName: step-9
---

# --description--

Maintenant, à l'intérieur de la boucle, affiche le littéral de gabarit suivant dans la console : `Checking "${word}" against "${match}" | Running count: ${count}`

# --hints--

Tu dois avoir exactement une instruction `console.log` à l'intérieur de la boucle dans ta fonction `getMatchedWordCount`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const functionRegex = /(functiongetMatchedWordCount\(sentence,match\)\{\S*?returncount;\})/;
const normalizedFunctionCode = normalizedCode.match(functionRegex)[1];
console.log(normalizedFunctionCode);

assert.match(normalizedFunctionCode, /for\(constwordofsentence\)\{\S*?console\.log\(\S*?\)\S*?\}/);
assert.notMatch(normalizedFunctionCode, /for\(constwordofsentence\)\{\S*?(console\.log\(\S*?\))\S*?(console\.log\(\S*?\))\S*?\}/);
```

Ton instruction `console.log` à l'intérieur de la boucle dans ta fonction `getMatchedWordCount` doit utiliser un littéral de gabarit avec les espaces réservés suivants : `${word}`, `${match}` et `${count}`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const functionRegex = /(functiongetMatchedWordCount\(sentence,match\)\{\S*?returncount;\})/;
const normalizedFunctionCode = normalizedCode.match(functionRegex)[1];
const consoleLogArguments = normalizedFunctionCode.match(/for\(constwordofsentence\)\{\S*?console\.log\((\S*?)\)\S*?\}/)[1]
const templateLiteral = consoleLogArguments.match(/(^\`\S*\`$)/)[1];

assert.include(templateLiteral, "${word}");
assert.include(templateLiteral, "${match}");
assert.include(templateLiteral, "${count}");
```

Lorsque la fonction `getMatchedWordCount` est appelée avec `(['foo'], 'bar')`, l'instruction `console.log()` à l'intérieur de la fonction doit être appelée une fois avec la chaîne `Checking "foo" against "bar" | Running count: 0`. Au lieu de cela, elle a été appelée avec chaque chaîne de la liste suivante séparée par des virgules : `--fcc-actual--`.

```js
const consoleLogSpy = __helpers.spyOn(console, 'log');
getMatchedWordCount(['foo'], 'bar');
assert.deepEqual(consoleLogSpy.calls, [['Checking "foo" against "bar" | Running count: 0']]);
consoleLogSpy.restore();
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;

  for (const word of sentence) {
--fcc-editable-region--

--fcc-editable-region--
  }

  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);
```
