---
id: 66be24cb4144f955b6bcc550
title: Construire une page de cartes à jouer
challengeType: 25
dashedName: build-a-page-of-playing-cards
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**Objectif :** Remplis les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu devrais construire une page web qui affiche au moins trois cartes à jouer.
1. Tu devrais avoir un élément `main` avec un ID de `playing-cards`.
1. Dans ton élément `#playing-cards`, tu devrais avoir au moins trois éléments `div`, chacun avec une classe `card`.
1. Dans chaque élément `.card`, tu devrais avoir trois éléments `div`, le premier avec une classe `left`, le deuxième avec une classe `middle` et le troisième avec une classe `right`.
1. Ton élément `#playing-cards` devrait utiliser Flexbox pour centrer horizontalement ses enfants, leur permettre de passer à la ligne et mettre un espace de `20px` entre eux.
1. Chacun de tes éléments `.card` devrait utiliser Flexbox pour répartir ses enfants avec `space-between`.
1. Chacun de tes éléments `.left` devrait utiliser des propriétés d'élément flex pour s'aligner au début de son conteneur parent.
1. Chacun de tes éléments `.middle` devrait utiliser des propriétés d'élément flex pour s'aligner au centre de son conteneur parent.
1. Chacun de tes éléments `.right` devrait utiliser des propriétés d'élément flex pour s'aligner à la fin de son conteneur parent.
1. Chacun de tes éléments `.middle` devrait utiliser Flexbox pour afficher ses enfants en colonne.

Voici quelques caractères que tu peux copier-coller pour les utiliser dans tes cartes si tu veux : `♠`, `♣`, `♥` et `♦`.

# --hints--

Tu devrais avoir un élément `main` avec un `id` de `playing-cards`.

```js
assert.exists(document.querySelector('main#playing-cards'));
```

Tu devrais avoir au moins trois éléments `div` avec une classe `card` dans ton élément `#playing-cards`.

```js
assert.isAtLeast(document.querySelectorAll('#playing-cards > div.card')?.length, 3);
```

Chacun de tes éléments `.card` devrait avoir une largeur et une hauteur.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => {
  assert.isAbove(el.clientWidth, 0);
  assert.isAbove(el.clientHeight, 0);
});
```

Chacun de tes éléments `.card` devrait avoir exactement trois éléments `div` comme enfants.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => {
  const children = el.children;
  assert.lengthOf(children, 3);
  Array.from(children).forEach(child => {
    assert.equal(child.tagName, 'DIV');
  });
});
```

Tu devrais avoir un `div` avec une classe `left` dans chaque élément `.card`.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.left')));
```

Tu devrais avoir un `div` avec une classe `middle` dans chaque élément `.card`.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.middle')));
```

Tu devrais avoir un `div` avec une classe `right` dans chaque élément `.card`.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.right')));
```

Tu devrais avoir un sélecteur `#playing-cards` qui définit le `display` de ses éléments à `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.display, 'flex');
```

Tu devrais avoir un sélecteur `#playing-cards` qui définit le `justify-content` de ses éléments à `center`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.justifyContent, 'center');
```

Tu devrais avoir un sélecteur `#playing-cards` qui définit le `flex-wrap` de ses éléments à `wrap`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.flexWrap, 'wrap');
```

Tu devrais avoir un sélecteur `#playing-cards` qui définit le `gap` de ses éléments à `20px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.gap, '20px');
```

Tu devrais avoir un sélecteur `.card` qui définit le `display` de ses éléments à `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.card')?.display, 'flex');
```

Tu devrais avoir un sélecteur `.card` qui définit le `justify-content` de ses éléments à `space-between`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.card')?.justifyContent, 'space-between');
```

Tu devrais avoir un sélecteur `.left` qui définit le `align-self` de ses éléments à `flex-start`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.left')?.alignSelf, 'flex-start');
```

Tu devrais avoir un sélecteur `.middle` qui définit le `align-self` de ses éléments à `center`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.alignSelf, 'center');
```

Tu devrais avoir un sélecteur `.right` qui définit le `align-self` de ses éléments à `flex-end`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.right')?.alignSelf, 'flex-end');
```

Tu devrais avoir un sélecteur `.middle` qui définit le `display` de ses éléments à `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.display, 'flex');
```

Tu devrais avoir une `flex-direction` définie à `column` dans ta classe `.middle`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.flexDirection, 'column');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Playing Cards</title>

</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Straight</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main id="playing-cards">
    <div class="card">
      <div class="left">
        <div>A</div>
        <div>♠</div>
      </div>
      <div class="middle">
        <div>♠</div>
      </div>
      <div class="right">
        <div class="flipped">♠</div>
        <div class="flipped">A</div>
      </div>
    </div>

    <div class="card">
      <div class="left">
        <div>2</div>
        <div>♣</div>
      </div>
      <div class="middle">
        <div>♣</div>
        <div class="flipped">♣</div>
      </div>
      <div class="right">
        <div class="flipped">♣</div>
        <div class="flipped">2</div>
      </div>
    </div>

    <div class="card red">
      <div class="left">
        <div>3</div>
        <div>♥</div>
      </div>
      <div class="middle">
        <div>♥</div>
        <div>♥</div>
        <div class="flipped">♥</div>
      </div>
      <div class="right">
        <div class="flipped">♥</div>
        <div class="flipped">3</div>
      </div>
    </div>

    <div class="card">
      <div class="left">
        <div>4</div>
        <div>♣</div>
      </div>
      <div class="middle">
        <div class="middle-section">
          <div>♣</div>
          <div>♣</div>
        </div>
        <div class="middle-section">
          <div class="flipped">♣</div>
          <div class="flipped">♣</div>
        </div>
      </div>
      <div class="right">
        <div class="flipped">♣</div>
        <div class="flipped">4</div>
      </div>
    </div>

    <div class="card red">
      <div class="left">
        <div>5</div>
        <div>♦</div>
      </div>
      <div class="middle">
        <div class="middle-section">
          <div>♦</div>
          <div>♦</div>
        </div>
        <div class="middle-section">
          <div>♦</div>
        </div>
        <div class="middle-section">
          <div>♦</div>
          <div>♦</div>
        </div>
      </div>
      <div class="right">
        <div class="flipped">♦</div>
        <div class="flipped">5</div>
      </div>
    </div>
  </main>
</body>

</html>
```

```css
body {
  background-color: Thistle;
  min-height: 100vh;
}

.red {
  color: red;
}

.flipped {
  transform: rotate(180deg);
}

#playing-cards {
  margin-top: 100px;
  display: flex;
  justify-content: center;
  gap: 20px;
  flex-wrap: wrap;
}

.card {
  background-color: white;
  border: 1px solid black;
  border-radius: 10px;
  box-shadow: 2px 2px 2px gray;
  min-width: 250px;
  max-width: 250px;
  height: 350px;
  font-size: 30px;
  font-weight: bold;
  padding: 5px;
  box-sizing: border-box;
  display: flex;
  justify-content: space-between;
  text-align: center;
}

.left {
  align-self: flex-start;
}

.middle {
  align-self: center;
  display: flex;
  font-size: 80px;
  flex-direction: column;
}

.right {
  align-self: flex-end;
}

.middle-section {
  display: flex;
  justify-content: center;
}
```
