---
id: 614385513d91ae5c251c2052
title: Étape 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Dans ce workshop, tu vas créer une page de magazine à partir d'un communiqué de presse issu d'une ancienne version du curriculum de freeCodeCamp.

Commence par ajouter un élément `link` pour la feuille de style de police `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`, un élément `link` pour la feuille de style FontAwesome `https://use.fontawesome.com/releases/v5.8.2/css/all.css`, et un élément `link` pour ta feuille de style `./styles.css`.

Ta feuille de style de police chargera trois polices distinctes : `Anton`, `Baskervville` et `Raleway`.

# --hints--

Ton code doit avoir trois éléments `link`.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 3);
```

Tes éléments `link` doivent être dans ton élément `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link')?.length, 3);
```

Tes trois éléments `link` doivent avoir un attribut `rel` défini sur `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

L'un de tes éléments `link` doit avoir le `href` défini sur `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap'));
```

L'un de tes éléments `link` doit avoir le `href` défini sur `https://use.fontawesome.com/releases/v5.8.2/css/all.css`.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css'));
```

L'un de tes éléments `link` doit avoir le `href` défini sur `styles.css`.

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
assert.strictEqual(styleElement?.getAttribute('data-href'), 'styles.css');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
