---
id: 6143869bb45bd85e3b1926aa
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

Dans ton élément `.heading`, crée un élément `header` avec la `class` définie sur `hero`.

Dans cet élément, crée un élément `img` avec le `src` défini sur `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png`, le `alt` défini sur `freecodecamp logo` et la `class` définie sur `hero-img`.

L'attribut `loading` d'un élément `img` peut être défini sur `lazy` pour indiquer au navigateur de ne pas récupérer la ressource de l'image avant qu'elle soit nécessaire, par exemple quand l'utilisateur fait défiler la page jusqu'à l'image. Comme avantage supplémentaire, les éléments chargés paresseusement ne se chargent qu'après les éléments non paresseux ; les utilisateurs avec une connexion lente peuvent donc voir le contenu de ta page sans attendre le chargement des images.

Donne à ton nouvel élément `img` un attribut `loading` défini sur `lazy`.

# --hints--

Tu dois créer un élément `header`.

```js
assert.exists(document.querySelector('header'));
```

Ton élément `header` doit être dans ton élément `.heading`.

```js
assert.equal(document.querySelector('header')?.parentElement?.className, 'heading');
```

Ton élément `header` doit avoir la `class` définie sur `hero`.

```js
assert.equal(document.querySelector('header')?.className, 'hero');
```

Tu dois créer un élément `img`.

```js
assert.exists(document.querySelector('img'));
```

Ton élément `img` doit être dans ton élément `header`.

```js
assert.equal(document.querySelector('img')?.parentElement?.localName, 'header');
```

Ton élément `img` doit avoir le `src` défini sur `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png`.

```js
assert.equal(document.querySelector('img')?.getAttribute('src'), 'https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png');
```

Ton élément `img` doit avoir le `alt` défini sur `freecodecamp logo`.

```js
assert.equal(document.querySelector('img')?.getAttribute('alt'), 'freecodecamp logo');
```

Ton élément `img` doit avoir un attribut `loading` défini sur `lazy`.

```js
assert.equal(document.querySelector('img')?.getAttribute('loading'), 'lazy');
```

Ton élément `img` doit avoir la `class` définie sur `hero-img`.

```js
assert.equal(document.querySelector('img')?.className, 'hero-img');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
