---
id: 6143862a5e54455d262c212e
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

Dans ton élément `body`, crée un élément `main`. Ensuite, dans cet élément, crée une `section` avec une `class` définie sur `heading`.

# --hints--

Tu dois avoir un élément `main`.

```js
assert.exists(document.querySelector('main'));
```

Ton élément `main` doit être dans ton élément `body`.

```js
assert.equal(document.querySelector('main')?.parentElement?.localName, 'body');
```

Tu dois avoir un élément `section`.

```js
assert.exists(document.querySelector('section'));
```

Ton élément `section` doit être dans ton élément `main`.

```js
assert.equal(document.querySelector('section')?.parentElement?.localName, 'main');
```

Ton élément `section` doit avoir la `class` définie sur `heading`.

```js
assert.isTrue(document.querySelector('section')?.classList?.contains('heading'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```

```css

```
