---
id: 614387cbefeeba5f3654a291
title: Étape 6
challengeType: 0
dashedName: step-6
---

# --description--

Après ton élément `header`, crée un élément `div` avec la `class` définie sur `author`.

Dans ce `div`, crée un élément `p` avec la `class` `author-name` et donne-lui le texte `By freeCodeCamp`. Enveloppe la partie `freeCodeCamp` dans un élément `a` avec le `href` défini sur `https://freecodecamp.org` et le `target` défini sur `_blank`.


Sous celui-ci, ajoute un second élément `p` avec la classe `publish-date` et le texte `March 7, 2019`.

# --hints--

Tu dois créer un nouvel élément `div`.

```js
assert.exists(document.querySelector('div'));
```

Ton élément `div` doit venir après ton élément `header`.

```js
assert.equal(document.querySelector('div')?.previousElementSibling?.localName, 'header');
```

Ton élément `div` doit avoir la `class` définie sur `author`.

```js
assert.equal(document.querySelector('div')?.className, 'author');
```

Tu dois créer deux nouveaux éléments `p`.

```js
assert.lengthOf(document.querySelectorAll('p'), 3);
```

Tes deux nouveaux éléments `p` doivent être dans ton élément `div`.

```js
assert.exists(document.querySelector('div')?.querySelectorAll('p')?.length === 2);
```

Ton premier nouvel élément `p` doit avoir la `class` définie sur `author-name`.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.className, 'author-name');
```

Ton premier nouvel élément `p` doit avoir le texte `By freeCodeCamp`.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.innerText, 'By freeCodeCamp');
```

Ton deuxième nouvel élément `p` doit avoir la `class` définie sur `publish-date`.

```js
assert.equal(document.querySelector('div')?.querySelectorAll('p')?.[1]?.className, 'publish-date');
```

Ton deuxième nouvel élément `p` doit avoir le texte `March 7, 2019`.

```js
assert.equal(document.querySelector('div')?.querySelectorAll('p')?.[1]?.innerText, 'March 7, 2019');
```

Tu dois créer un nouvel élément `a`.

```js
assert.exists(document.querySelector('a'));
```

Ton élément `a` doit être dans ton premier nouvel élément `p`.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.localName, 'a');
```

Ton élément `a` doit avoir le `href` défini sur `https://freecodecamp.org`.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('href'), 'https://freecodecamp.org');
```

Ton élément `a` doit avoir le `target` défini sur `_blank`.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('target'), '_blank');
```

Ton élément `a` doit entourer le texte `freeCodeCamp`.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.innerText, 'freeCodeCamp');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
--fcc-editable-region--

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
