---
id: 614389f601bb4f611db98563
title: Étape 8
challengeType: 0
dashedName: step-8
---

# --description--

Sous ton élément `.author`, crée un nouvel élément `div` avec la classe `social-icons`.

Ajoute cinq éléments `a` dans ce nouveau `div`, et donne-leur les attributs `href` suivants.

- Le premier élément `a` doit avoir un `href` défini sur `https://www.facebook.com/freecodecamp`.
- Le deuxième élément `a` doit avoir un `href` défini sur `https://twitter.com/freecodecamp`.
- Le troisième élément `a` doit avoir un `href` défini sur `https://instagram.com/freecodecamp`.
- Le quatrième élément `a` doit avoir un `href` défini sur `https://www.linkedin.com/school/free-code-camp`.
- Le cinquième élément `a` doit avoir un `href` défini sur `https://www.youtube.com/freecodecamp`.

# --hints--

Tu dois créer un nouvel élément `div`.

```js
assert.lengthOf(document.querySelectorAll('div'), 2);
```

Ton nouvel élément `div` doit venir après ton élément `.author`.

```js
assert.equal(document.querySelector('.author')?.nextElementSibling?.localName, 'div');
```

Ton nouvel élément `div` doit avoir la classe `social-icons`.

```js
assert.isTrue(document.querySelector('.author')?.nextElementSibling?.classList?.contains('social-icons'));
```

Ton élément `.social-icons` doit avoir cinq éléments `a`.

```js
assert.lengthOf(document.querySelector('.social-icons')?.querySelectorAll('a'), 5);
```

Ton premier élément `a` doit avoir le `href` défini sur `https://www.facebook.com/freecodecamp`.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[0]?.getAttribute('href'), 'https://www.facebook.com/freecodecamp');
```

Ton deuxième élément `a` doit avoir le `href` défini sur `https://twitter.com/freecodecamp`.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[1]?.getAttribute('href'), 'https://twitter.com/freecodecamp');
```

Ton troisième élément `a` doit avoir le `href` défini sur `https://instagram.com/freecodecamp`.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[2]?.getAttribute('href'), 'https://instagram.com/freecodecamp');
```

Ton quatrième élément `a` doit avoir le `href` défini sur `https://www.linkedin.com/school/free-code-camp`.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[3]?.getAttribute('href'), 'https://www.linkedin.com/school/free-code-camp');
```

Ton cinquième élément `a` doit avoir le `href` défini sur `https://www.youtube.com/freecodecamp`.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[4]?.getAttribute('href'), 'https://www.youtube.com/freecodecamp');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            By
            <a href="https://freecodecamp.org" target="_blank" rel="noreferrer"
              >freeCodeCamp</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
