---
id: 6165d3b702a5d92ad970b30c
title: Étape 4
challengeType: 0
dashedName: step-4
---

# --description--

Après ton élément `img`, ajoute un élément `h1` avec la `class` définie sur `hero-title` et le texte défini sur `OUR NEW CURRICULUM`, suivi d'un élément `p` avec la `class` définie sur `hero-subtitle` et le texte défini sur `Our efforts to restructure our curriculum with a more project-based focus`.

# --hints--

Tu dois créer un élément `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Ton élément `h1` doit venir après ton élément `img`.

```js
assert.equal(document.querySelector('h1')?.previousElementSibling?.localName, 'img');
```

Ton élément `h1` doit avoir la `class` définie sur `hero-title`.

```js
assert.equal(document.querySelector('h1')?.className, 'hero-title');
```

Ton élément `h1` doit avoir le texte `OUR NEW CURRICULUM`.

```js
assert.equal(document.querySelector('h1')?.textContent, 'OUR NEW CURRICULUM');
```

Tu dois créer un nouvel élément `p`.

```js
assert.exists(document.querySelector('p'));
```

Ton élément `p` doit venir après ton élément `h1`.

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.localName, 'h1');
```

Ton élément `p` doit avoir la `class` définie sur `hero-subtitle`.

```js
assert.equal(document.querySelector('p')?.className, 'hero-subtitle');
```

Ton élément `p` doit avoir le texte `Our efforts to restructure our curriculum with a more project-based focus`.

```js
assert.equal(document.querySelector('p')?.textContent?.trim()?.replace(/\s{2,}/, ' '), 'Our efforts to restructure our curriculum with a more project-based focus');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
          />
        </header>
      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
