---
id: 6169cd8a558aa8434e0ad7f6
title: Étape 7
challengeType: 0
dashedName: step-7
---

# --description--

L'en-tête HTTP `Referer` contient des informations sur l'adresse ou l'URL de la page d'où l'utilisateur peut venir. Ces informations peuvent être utilisées dans l'analytique pour suivre, par exemple, combien d'utilisateurs depuis ta page visitent freecodecamp.org. Définir l'attribut `rel` sur `noreferrer` omet ces informations dans la requête HTTP. Donne à ton élément `a` un attribut `rel` défini sur `noreferrer`.

# --hints--

Ton élément `a` doit avoir le `rel` défini sur `noreferrer`.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('rel'), 'noreferrer');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
--fcc-editable-region--
        <div class="author">
          <p class="author-name">
            By
            <a href="https://freecodecamp.org" target="_blank"
              >freeCodeCamp</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
