---
id: 669e2f60e83c011754f711f9
title: Construire une page d'agence de voyage
challengeType: 25
dashedName: build-a-travel-agency-page
demoType: onClick
---

# --description--

**Objectif :** respecte les histoires utilisateur ci-dessous et fais passer tous les tests pour terminer ce lab.

**Histoires utilisateur :**

1. Tu dois avoir une déclaration `DOCTYPE`.
1. Tu dois avoir un élément `html` avec `lang` défini sur `en`.
1. Tu dois avoir un élément `head` contenant un élément vide `meta` avec `charset` défini sur `utf-8` et un `title` avec le texte `Travel Agency Page`.
1. Tu dois avoir une balise `meta` dans ton élément `head` qui contient une courte description de ton site web pour le SEO.
1. Tu dois avoir un élément `h1` pour présenter tes destinations de voyage.
1. Tu dois avoir un paragraphe sous l'élément `h1` qui présente les possibilités de voyage.
1. Tu dois avoir un élément `h2` avec le texte `Packages`.
1. Tu dois avoir un élément `p` qui présente brièvement les différents forfaits.
1. Tu dois avoir un élément de liste non ordonnée avec deux éléments de liste. Les deux éléments de liste doivent avoir respectivement les textes `Group Travels` et `Private Tours`. Le texte de chaque élément de liste doit être entouré par un élément d'ancre.
1. Tu dois avoir un élément `h2` avec le texte `Top Itineraries`.
1. Tu dois avoir au moins trois éléments `figure`, chacun contenant un élément d'ancre et un élément `figcaption`.
1. Les trois éléments d'ancre doivent avoir comme contenu un élément `img` avec un attribut `alt` approprié et un attribut `src` défini sur une image valide. Tu peux utiliser `https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg`, `https://cdn.freecodecamp.org/curriculum/labs/alps.jpg` et `https://cdn.freecodecamp.org/curriculum/labs/sea.jpg` si tu veux.
1. Tes cinq éléments d'ancre doivent tous avoir un attribut `href` avec la valeur `https://www.freecodecamp.org/learn` et un attribut `target` avec la valeur `_blank`.

# --hints--

Ta page d'agence de voyage doit avoir une déclaration `!DOCTYPE html`.

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

Tu dois avoir un élément `html` avec `lang` défini sur `en`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[\s\S]*<\/\s*html\s*>/i);
```

Tu dois avoir un élément `head` dans l'élément `html`.

```js
assert.match(
  code,
  /<html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\/\s*html\s*>/i
);
```

Tu dois avoir deux éléments `meta` dans ton élément `head`.

```js
assert.match(code, /<\s*head\s*>[\s\S]*(<\s*meta[\s\S]*>[\s\S]*){2,}<\/\s*head\s*>/i);
```

Un élément `meta` doit avoir un attribut `name` avec la valeur `description` et un attribut `content` non vide.

```js
 assert.match(code,
  /<\s*meta\s+[^>]*name\s*=\s*["']description["'][^>]*content\s*=\s*["'][^"']+?["'][^>]*>/i,
);
```

Un élément `meta` doit avoir son attribut `charset` défini sur `UTF-8`.

```js
assert.match(code, /<\s*meta[\s\S]+?charset\s*=\s*('|")UTF-8\1/i);
```

Tu dois avoir un élément `title` dans ton élément `head`.

```js
assert.match(
  code,
  /<\s*head\s*>[\s\S]*<\s*title\s*>[\s\S]*<\/\s*title\s*>[\s\S]*<\/\s*head\s*>/i
);
```

Ton élément `title` doit contenir le nom de ton agence de voyage.

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

Tu dois avoir un élément `body` dans ton élément `html`.

```js
assert.match(
  code,
  /<\s*html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\s*body\s*>[\s\S]*<\/\s*body\s*>[\s\S]*<\/\s*html\s*>/i
);
```

Tu dois avoir un élément `h1` pour présenter tes destinations de voyage.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

Tu ne dois avoir qu'un seul élément `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Tu dois avoir un élément `p` sous l'élément `h1`.

```js
assert.strictEqual(
  document.querySelector('h1')?.nextElementSibling.tagName, 'P'
);
```

Ton premier élément `p` doit présenter les possibilités de voyage.

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

Ton premier élément `h2` doit avoir le texte `Packages`.

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText.trim(), 'Packages');
```

Tu dois avoir un élément `p` sous ton premier élément `h2`.

```js
assert.strictEqual(
    document.querySelector('h2')?.nextElementSibling.tagName, 'P'
);
```

Ton deuxième élément `p` doit présenter brièvement les différents forfaits.

```js
assert.isNotEmpty(document.querySelectorAll('p')?.[1]?.textContent?.trim());
```

Tu dois avoir un élément de liste non ordonnée sous ton deuxième élément `p`.

```js
assert.strictEqual(
  document.querySelectorAll('p')?.[1]?.nextElementSibling?.tagName,
  'UL'
);
```

Tu dois avoir deux éléments dans ta liste non ordonnée.

```js
assert.lengthOf(document.querySelectorAll('p + ul li'), 2);
```

Tes deux éléments de liste doivent contenir un élément d'ancre.

```js
const listItems = document.querySelectorAll('p + ul li');
assert.isNotEmpty(listItems);
for (let e of listItems) {
  assert.strictEqual(e.children[0].tagName, 'A');
  assert.lengthOf(e.children, 1);
}
```

L'élément d'ancre de ton premier élément de liste doit entourer le texte `Group Travels`.

```js
assert.equal(document.querySelectorAll('li > a')[0]?.innerText.trim(), 'Group Travels');
```

L'élément d'ancre de ton deuxième élément de liste doit entourer le texte `Private Tours`.

```js
assert.equal(document.querySelectorAll('li > a')[1]?.innerText.trim(), 'Private Tours');
```

Tu dois avoir un élément `h2` après ta liste non ordonnée.

```js
assert.strictEqual(
  document.querySelectorAll('h2')[1]?.previousElementSibling.tagName,
  'UL'
);
```

Ton deuxième élément `h2` doit avoir le texte `Top Itineraries`.

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText.trim(), 'Top Itineraries');
```

Tu dois avoir au moins trois éléments `figure`.

```js
assert.isAtLeast(document.querySelectorAll('figure').length, 3);
```

Chaque élément `figure` doit contenir un élément d'ancre comme premier enfant.

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[0].tagName, 'A');
}
```

Chaque élément `figure` doit contenir un élément `figcaption` comme deuxième enfant.

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[1]?.tagName, 'FIGCAPTION');
}
```

Chaque élément `figcaption` doit contenir du texte.

```js
const figcaptionEls = document.querySelectorAll('figcaption');
assert.isNotEmpty(figcaptionEls);
figcaptionEls.forEach((figcaption => assert.isNotEmpty(figcaption.innerText)));
```

Chacun des éléments `a` enfants de tes éléments `figure` doit contenir une image.

```js
const anchors = document.querySelectorAll('figure > a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.children[0]?.tagName, 'IMG');
}
```

Chaque élément `img` doit avoir un attribut `src` valide.

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.src.trim().length, 0);
}
```

Chaque élément `img` doit avoir un attribut `alt` avec une valeur appropriée.

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.alt.trim().length, 0);
}
```

Chaque élément `a` doit avoir un attribut `href` avec la valeur `https://www.freecodecamp.org/learn`. N'oublie pas les liens dans les éléments de liste.

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.href, 'https://www.freecodecamp.org/learn');
}
```

Chaque élément `a` doit avoir un attribut `target` avec la valeur `_blank`. N'oublie pas les liens dans les éléments de liste.

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.target, '_blank');
}
```
# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Travel Agency Page</title>
    <meta name="description" content="This is an example website."/>
  </head>
  <body>
    <h1>Discover Italy</h1>
    <p>
      Art, folklore, food, nature, and more. Choose among our wide selection of
      guided tours and excursions, and live an unforgettable experience
      exploring Italy.
    </p>
    <h2>Packages</h2>
    <p>
      We offer an extensive range of holiday solutions to accommodate the needs
      of all our clients. From daily excursions in the most beautiful cities, to
      thorough tours of hidden villages and medieval towns to discover Italy's
      lesser-known sides.
    </p>
    <ul>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Group Travels</a
        >
      </li>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Private Tours</a
        >
      </li>
    </ul>
    <h2>Top Itineraries</h2>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg"
          alt="colosseum view"
        />
      </a>
      <figcaption>Rome and Center Italy</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/alps.jpg"
          alt="dolomites mountain view"
        />
      </a>
      <figcaption>Nature and National Parks</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/sea.jpg"
          alt="sea view"
        />
      </a>
      <figcaption>South Italy and Islands</figcaption>
    </figure>
  </body>
</html>
```
