---
id: 690ddb5e6ee94eb73ed172b0
title: Déboguer un formulaire de don
challengeType: 25
dashedName: debug-donation-form
# Purposefully removed the demo type because this is a debugging project.
---

# --description--

Une association locale a créé un site avec un formulaire de don, mais plusieurs problèmes doivent être corrigés. Le formulaire n'est pas accessible et contient quelques erreurs de syntaxe HTML.

Ton objectif est de corriger toutes les erreurs pour que le formulaire fonctionne correctement et soit accessible à tous les utilisateurs. Termine les éléments des user stories ci-dessous, puis clique sur "Check Your Code" pour vérifier si tu as corrigé toutes les erreurs.

**User Stories :**

1. Les éléments `input` sont des éléments vides et ne doivent pas avoir de balises fermantes. Supprime toutes les balises fermantes `</input>` du formulaire.
2. Ajoute des éléments `label` pour chaque champ de saisie du formulaire afin que les utilisateurs sachent à quoi sert chaque champ. Le texte du libellé doit correspondre au texte actuellement placé à côté de chaque champ.
3. Le champ `Email Address:` doit avoir le type `email` au lieu de `text`.
4. Associe chaque élément `label` à son élément `input` correspondant avec l'attribut `for` sur le `label` et un attribut `id` correspondant sur l'`input`.
5. Ajoute l'attribut `required` aux champs de saisie de texte, d'e-mail et de nombre, mais pas à la case à cocher ni au bouton d'envoi, afin de t'assurer que les utilisateurs remplissent les informations obligatoires.

# --hints--

Tu ne devrais pas avoir de balises fermantes `</input>` dans ton code.

```js
assert.notMatch(code, /<\/input>/);
```

Tu devrais avoir exactement cinq éléments `input` dans ton formulaire.

```js
assert.lengthOf(document.querySelectorAll("input"), 5);
```

Tu devrais avoir exactement quatre éléments `label` dans ton formulaire.

```js
assert.lengthOf(document.querySelectorAll("label"), 4);
```

Ton premier élément `label` devrait avoir le texte `Full Name:`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[0]?.textContent, /Full\s*Name\s*:/i);
```

Ton premier élément `input` devrait avoir un attribut `required`.

```js
const firstInput = document.querySelector("input[name='name']");
assert.isTrue(firstInput?.hasAttribute("required"));
```

Ton deuxième élément `label` devrait avoir le texte `Email Address:`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[1]?.textContent, /Email\s*Address\s*:/i);
```

Ton deuxième élément `input` devrait avoir un attribut `required`.

```js
const secondInput = document.querySelector("input[name='email']");
assert.isTrue(secondInput?.hasAttribute("required"));
```

Ton troisième élément `label` devrait avoir le texte `Donation Amount ($):`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[2]?.textContent, /Donation\s*Amount\s*\(\$\)\s*:/i);
```

Ton troisième élément `input` devrait avoir un attribut `required`.

```js
const thirdInput = document.querySelector("input[name='amount']");
assert.isTrue(thirdInput?.hasAttribute("required"));
```

Ton quatrième élément `label` devrait avoir le texte `Subscribe`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[3]?.textContent, /Subscribe/i);
```

Tu devrais associer chaque élément `input` à un élément `label` avec les attributs `for` et `id`.

```js
const els = document.querySelectorAll('input:not([type="submit"])');
assert.isNotEmpty(els);
els.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`);
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
})
```

Ton champ `input` pour l'adresse e-mail devrait être de type `email`.

```js
const emailInput = document.querySelector("input[name='email']");
assert.strictEqual(emailInput?.getAttribute("type"), "email");
```

Ton élément `input` de case à cocher ne devrait pas avoir d'attribut `required`.

```js
const checkboxInput = document.querySelector("input[name='newsletter']");
assert.isFalse(checkboxInput?.hasAttribute("required"));
```

Ton bouton d'envoi ne devrait pas avoir d'attribut `required`.

```js
const submitButton = document.querySelector("input[type='submit']");
assert.isFalse(submitButton?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Donation Form</title>
</head>
<body>
  <h1>Donation Form</h1>
  <form>

    Full Name:
    <input type="text" name="name"></input>

    Email Address:
    <input type="text" name="email">

    Donation Amount ($):
    <input type="number" name="amount"></input>

    <input type="checkbox" name="newsletter"></input>
    Subscribe

    <input type="submit" value="Send"></input>
  </form>
</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Donation Form</title>
</head>
<body>
  <h1>Donation Form</h1>
  <form>
    <label for="full-name">Full Name:</label>
    <input id="full-name" type="text" name="name" required>

    <label for="email">Email Address:</label>
    <input id="email" type="email" name="email" required>

    <label for="amount">Donation Amount ($):</label>
    <input id="amount" type="number" name="amount" required>

    <label for="newsletter">Subscribe</label>
    <input id="newsletter" type="checkbox" name="newsletter">

    <input type="submit" value="Send">
  </form>
</body>
</html>
```
