---
id: 68f1f8c248306da204655117
title: Étape 6
challengeType: 0
dashedName: step-6
---

# --description--

L'attribut `muted` coupe le son lors de la lecture initiale. Si `controls` est activé, l'utilisateur pourra réactiver le son. Si tu omets l'attribut `muted`, le son sera lu dès la lecture initiale.

L'attribut `muted` est un attribut booléen et n'a pas besoin de valeur.

Ajoute l'attribut `muted` à l'élément `video`.

# --hints--

Ton élément `video` doit avoir l'attribut `muted`.

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('muted'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Travailler avec l'élément vidéo HTML</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640" loop controls>
--fcc-editable-region--
  </video>
</body>
</html>
```
