---
id: 68f1eb9ad0a7575e4dbc31d5
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

Ensuite, crée un élément `video` sous le `h1`. Au cours des prochaines étapes, tu ajouteras les attributs nécessaires pour rendre le lecteur vidéo fonctionnel.

# --hints--

Tu dois avoir un élément `video`.

```js
assert.exists(document.querySelector('video'));
```

Le nouvel élément `video` doit se trouver sous le `h1`.

```js
assert.exists(document.querySelector('h1 + video'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Travailler avec l'élément vidéo HTML</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--

--fcc-editable-region--
</body>
</html>
```
