---
id: 68f1f6f97af9749a42260a43
title: Étape 5
challengeType: 0
dashedName: step-5
---

# --description--

L'attribut `controls` fournit des contrôles de lecture, dont la lecture, le retour en arrière et le contrôle du volume, pour l'élément `video`.

L'attribut `controls` est un attribut booléen et n'a pas besoin de valeur.

Ajoute l'attribut `controls` à l'élément `video`.

Tu devrais maintenant voir l'élément `video` affiché sur la page.

# --hints--

Ton élément `video` doit avoir l'attribut `controls`.

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Travailler avec l'élément vidéo HTML</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640" loop>
--fcc-editable-region--
  </video>
</body>
</html>
```
