---
id: 68f1fac259a4afabb478b960
title: Étape 8
challengeType: 0
dashedName: step-8
---

# --description--

Tu as peut-être remarqué que tu n'as pas encore lié la vraie vidéo. Tu vas le faire dans la phase suivante. Pour les types de fichiers vidéo, la prise en charge varie selon les navigateurs. Pour gérer cela, tu peux utiliser des éléments `source` dans l'élément `video`, et le navigateur choisira le premier élément `source` compatible.

Voici un exemple d'élément `source` :

```html
<video controls width="250">
  <source src="src-url-goes-here" type="video-type-goes-here" />
</video>
```

L'élément `source` est un élément vide, il n'a donc pas de balise de fermeture.

Ajoute un élément `source` à l'intérieur de ton élément `video`.

# --hints--

Tu dois avoir un élément `source` dans ton élément `video`.

```js
assert.exists(document.querySelector('video > source'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Travailler avec l'élément vidéo HTML</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
--fcc-editable-region--

--fcc-editable-region--
  </video>
</body>
</html>
```
