---
id: 68f1f993460f06a64c033061
title: Étape 7
challengeType: 0
dashedName: step-7
---

# --description--

L'attribut `poster` est une image miniature de la vidéo. Pense aux vidéos que tu regardes sur YouTube. Elle s'affiche pendant le téléchargement de la vidéo. Si l'attribut est omis, la première image de la vidéo est affichée pendant la phase de téléchargement.

Maintenant, ajoute l'attribut `poster` avec la valeur `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg` à ton élément `video`.

# --hints--

Ton élément `video` doit avoir l'attribut `poster` avec la valeur `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg`.

```js
const video = document.querySelector('video');
assert.strictEqual(video?.getAttribute('poster'), 'https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Travailler avec l'élément vidéo HTML</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
--fcc-editable-region--

--fcc-editable-region--
  >
  </video>
</body>
</html>
```
