---
id: 68f1f57723debf9383c75947
title: Étape 4
challengeType: 0
dashedName: step-4
---

# --description--

L'attribut `loop` redémarre la vidéo une fois la lecture terminée. Pense à un mème sur Internet qui se répète. Si tu omets l'attribut `loop`, la vidéo ne sera lue qu'une seule fois.

L'attribut `loop` est un attribut booléen et n'a pas besoin de valeur.

Ajoute l'attribut `loop` à l'élément `video`.

# --hints--

Ton élément `video` doit avoir l'attribut `loop`.

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('loop'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Travailler avec l'élément vidéo HTML</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640">
--fcc-editable-region--
  </video>
</body>
</html>
```
