---
id: 68f18724d71db1dda8d6fec7
title: Étape 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Dans cet atelier, tu vas construire un lecteur vidéo HTML. Le code de base HTML t'est déjà fourni.

Crée un élément `h1` et donne-lui le texte `Working with the HTML Video Element`.

# --hints--

Tu dois avoir un élément `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Ton élément `h1` doit contenir le texte `Working with the HTML Video Element`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent.trim(), 'Working with the HTML Video Element');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Travailler avec l'élément vidéo HTML</title>
</head>
<body>
--fcc-editable-region--

--fcc-editable-region--
</body>
</html>
```
