---
id: 68f1f371fefd448b4f70b898
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

Dans une leçon précédente, tu as découvert différents attributs disponibles pour l'élément `video`. L'attribut `width` définit la largeur de la vidéo en pixels.

Ajoute l'attribut `width` à l'élément `video` avec la valeur `640`.

# --hints--

Ton élément `video` doit avoir l'attribut `width` avec la valeur `640`.

```js
const video = document.querySelector('video');
assert.strictEqual(video?.getAttribute('width'), '640');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Travailler avec l'élément vidéo HTML</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video>
--fcc-editable-region--
  </video>
</body>
</html>
```
