---
id: 68f2701b2160c765f23b60f8
title: Étape 12
challengeType: 0
dashedName: step-12
---

# --description--

Un autre type MIME courant est le type MIME `video/ogg`.

Ogg est un format de conteneur multimédia numérique conçu pour permettre une diffusion et une manipulation efficaces du multimédia numérique. Il est maintenu par la fondation Xiph.Org et il est libre, ouvert et non limité par des brevets logiciels. Son nom vient de "ogging", un terme du jargon du jeu informatique Netrek.

Sous ton deuxième élément `source`, ajoute un troisième élément `source` et donne-lui un attribut `src` avec la valeur `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg` et un attribut `type` avec la valeur `video/ogg`.

# --hints--

Tu dois avoir un troisième élément `source`.

```js
assert.exists(document.querySelector('video > source:nth-of-type(3)'));
```

Ton troisième élément `source` doit avoir un attribut `src`.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.isTrue(source?.hasAttribute('src'));
```

Ton troisième élément `source` doit avoir un attribut `src` avec la valeur `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg`.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg');
```

Ton troisième élément `source` doit avoir un attribut `type`.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.isTrue(source?.hasAttribute('type'));
```

Ton troisième élément `source` doit avoir un attribut `type` avec la valeur `video/ogg`.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.strictEqual(source?.getAttribute('type'), 'video/ogg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Travailler avec l'élément vidéo HTML</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
--fcc-editable-region--

--fcc-editable-region--
  </video>
</body>
</html>
```
