---
id: 68f27139bddf9f6bf2821acf
title: Étape 13
challengeType: 0
dashedName: step-13
---

# --description--

Le dernier élément `source` que tu vas ajouter sera pour le type MIME `video/quicktime`.

QuickTime est une architecture multimédia extensible créée par Apple. Elle prend en charge la lecture, la diffusion, l'encodage et le transcodage de nombreux formats de médias numériques. Même si ce format est moins populaire que le format MP4, tu peux en avoir besoin pour prendre en charge d'anciennes applications.

Sous ton troisième élément `source`, ajoute un quatrième élément `source` et donne-lui un attribut `src` avec la valeur `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov` et un attribut `type` avec la valeur `video/quicktime`.

Félicitations ! Tu as terminé l'atelier Lecteur vidéo HTML.

# --hints--

Tu dois avoir un quatrième élément `source`.

```js
assert.exists(document.querySelector('video > source:nth-of-type(4)'));
```

Ton quatrième élément `source` doit avoir un attribut `src`.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.isTrue(source?.hasAttribute('src'));
```

Ton quatrième élément `source` doit avoir un attribut `src` avec la valeur `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov`.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov');
```

Ton quatrième élément `source` doit avoir un attribut `type`.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.isTrue(source?.hasAttribute('type'));
```

Ton quatrième élément `source` doit avoir un attribut `type` avec la valeur `video/quicktime`.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.strictEqual(source?.getAttribute('type'), 'video/quicktime');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Travailler avec l'élément vidéo HTML</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg"
      type="video/ogg"
    >
--fcc-editable-region--

--fcc-editable-region--
  </video>
</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Travailler avec l'élément vidéo HTML</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    controls
    loop
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
       src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg"
      type="video/ogg"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov"
      type="video/quicktime"
    >
  </video>
</body>
</html>
```
