---
id: 68f268880219c2485bd71d2a
title: Étape 9
challengeType: 0
dashedName: step-9
---

# --description--

Pour préciser la ressource multimédia de la vidéo, tu dois ajouter l'attribut `src` à l'élément `source`.

Ajoute l'attribut `src` avec la valeur `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4`.

# --hints--

Ton élément `source` doit avoir un attribut `src`.

```js
const source = document.querySelector('source');
assert.isTrue(source?.hasAttribute('src'));
```

Ton élément `source` doit avoir un attribut `src` avec la valeur `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4`.

```js
const source = document.querySelector('source');
assert.strictEqual(source?.getAttribute('src'),
'https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Travailler avec l'élément vidéo HTML</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
--fcc-editable-region--
    <source>
--fcc-editable-region--
  </video>
</body>
</html>
```
