---
id: 68f26a5f79e97c5032c10642
title: Étape 10
challengeType: 0
dashedName: step-10
---

# --description--

Tu as utilisé un fichier vidéo avec l'extension `mp4`, et tu dois l'indiquer au navigateur pour qu'il sache comment lire le fichier.

Tu utiliseras l'attribut `type` pour préciser le type MIME `video/mp4`.

MIME (Multipurpose Internet Mail Extensions) est une norme qui sert à décrire des documents dans d'autres formes que le texte ASCII, par exemple l'audio, la vidéo et les images.

MP4, officiellement appelé MPEG-4 Part 14, est un format de conteneur multimédia numérique. Il est largement utilisé pour stocker de la vidéo et de l'audio, mais il peut aussi inclure d'autres types de données, comme des sous-titres et des images fixes. Les fichiers MP4 sont conçus pour être diffusés sur Internet et sont compatibles avec de nombreux appareils et plateformes.

Maintenant, ajoute l'attribut `type` et la valeur `video/mp4`.

# --hints--

Ton élément `source` doit avoir un attribut `type`.

```js
const source = document.querySelector('source');
assert.isTrue(source?.hasAttribute('type'));
```

Ton élément `source` doit avoir un attribut `type` avec la valeur `video/mp4`.

```js
const source = document.querySelector('source');
assert.strictEqual(source?.getAttribute('type'), 'video/mp4');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Travailler avec l'élément vidéo HTML</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
--fcc-editable-region--

--fcc-editable-region--
    >
  </video>
</body>
</html>
```
