---
id: 68f26d52b13e0f5b6e2d2e09
title: Étape 11
challengeType: 0
dashedName: step-11
---

# --description--

Un autre type MIME courant est le type MIME `video/webm`.

WebM est un format de fichier multimédia audiovisuel open source développé par Google, principalement conçu pour le contenu multimédia sur le Web. Il prend en charge des codecs vidéo comme VP8, VP9 et AV1, ainsi que des codecs audio comme Vorbis et Opus, ce qui en fait un choix courant pour les éléments vidéo et audio HTML5.

Sous ton premier élément `source`, ajoute un autre élément `source` et donne-lui un attribut `src` avec la valeur `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm` et un attribut `type` avec la valeur `video/webm`.

# --hints--

Tu dois avoir un deuxième élément `source`.

```js
assert.exists(document.querySelector('video > source:nth-of-type(2)'));
```

Ton deuxième élément `source` doit avoir un attribut `src`.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.isTrue(source?.hasAttribute('src'));
```

Ton deuxième élément `source` doit avoir un attribut `src` avec la valeur `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm`.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm');
```

Ton deuxième élément `source` doit avoir un attribut `type`.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.isTrue(source?.hasAttribute('type'));
```

Ton deuxième élément `source` doit avoir un attribut `type` avec la valeur `video/webm`.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.strictEqual(source?.getAttribute('type'), 'video/webm');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Travailler avec l'élément vidéo HTML</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
--fcc-editable-region--

--fcc-editable-region--
  </video>
</body>
</html>
```
