---
id: 612e96fc87fe8e44f69f7ec5
title: Étape 11
challengeType: 0
dashedName: step-11
---

# --description--

Il est temps de styliser les touches. Sous la règle `#piano`, cible l'élément `.keys` avec un sélecteur de `class`. Donne à la nouvelle règle une propriété `background-color` de `#040404`, une propriété `width` de `949px` et une propriété `height` de `180px`.

# --hints--

Tu devrais avoir un sélecteur `.keys`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.keys'));
```

Ton sélecteur `.keys` devrait avoir une propriété `background-color` définie sur `#040404`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.keys')?.backgroundColor, 'rgb(4, 4, 4)');
```

Ton sélecteur `.keys` devrait avoir la propriété `width` définie sur `949px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.keys')?.width, '949px');
```

Ton sélecteur `.keys` devrait avoir une propriété `height` définie sur `180px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.keys')?.height, '180px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
}

--fcc-editable-region--

--fcc-editable-region--
```
