---
id: 612e7d1c29fb872d6384379c
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

Imbrique un second `div` dans ton `div` existant, et définis la `class` sur `keys`.

# --hints--

Tu devrais créer un second élément `div`.

```js
const divDiv = document.querySelectorAll('div');
assert.lengthOf(divDiv, 2);
```

Ton nouvel élément `div` devrait être dans ton élément `div` existant.

```js
const div = document.querySelector('div');
assert.lengthOf(div?.children, 1);
assert.equal(div?.children?.[0]?.localName, 'div');
```

Ton nouvel élément `div` devrait avoir la `class` définie sur `keys`.

```js
const div = document.querySelector('div');
assert.equal(div?.children?.[0]?.className, 'keys');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    <div id="piano">
  --fcc-editable-region--

  --fcc-editable-region--
    </div>
  </body>
</html>
```

```css

```
