---
id: 612e89d254fe5d3df7d6693d
title: Étape 8
challengeType: 0
dashedName: step-8
---

# --description--

Maintenant que tu as réinitialisé le modèle de boîte de `html`, tu dois aussi l'appliquer aux éléments qu'il contient. Pour cela, tu devras définir la propriété `box-sizing` de tous les autres éléments pour qu'elle soit héritée, ce qui garantit que les éléments ciblés adoptent la même valeur que leur élément parent.

Tu devras aussi cibler les pseudo-éléments, qui sont des mots-clés spéciaux ajoutés après un sélecteur. Les deux pseudo-éléments que tu utiliseras sont les pseudo-éléments `::before` et `::after`.

Le sélecteur `::before` crée un pseudo-élément qui est le premier enfant de l'élément sélectionné, tandis que le sélecteur `::after` crée un pseudo-élément qui est le dernier enfant de l'élément sélectionné. Ces pseudo-éléments sont souvent utilisés pour créer du contenu décoratif, comme tu le verras plus tard dans ce projet.

Pour l'instant, crée un sélecteur CSS qui cible tous les éléments avec `*`, et inclus les pseudo-éléments `::before` et `::after`. Définis la propriété `box-sizing` sur `inherit`.

# --hints--

Tu devrais avoir un sélecteur `html`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('html'));
```

Ton sélecteur `html` devrait avoir la propriété `box-sizing` définie sur `border-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.boxSizing, 'border-box');
```

Tu devrais avoir un sélecteur `*, ::before, ::after`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after'));
```

Ton sélecteur `*, ::before, ::after` devrait avoir la propriété `box-sizing` définie sur `inherit`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after')?.boxSizing, 'inherit');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
html {
  box-sizing: border-box;
}
--fcc-editable-region--
```
