---
id: 612e89562043183c86df287c
title: Étape 7
challengeType: 0
dashedName: step-7
---

# --description--

Les navigateurs appliquent souvent des valeurs de marge et de remplissage par défaut à certains éléments. Pour t'assurer que la mise en page de ton piano s'affiche correctement, tu dois réinitialiser le modèle de boîte.

Ajoute un sélecteur de règle `html` à ton fichier CSS, et définis la propriété `box-sizing` sur `border-box`.

# --hints--

Tu devrais avoir un sélecteur `html`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('html'));
```

Ton sélecteur `html` devrait avoir la propriété `box-sizing` définie sur `border-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.boxSizing, 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
