---
id: 612e804c54d5e7308d7ebe56
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

Dans ton élément `.keys`, ajoute sept éléments `div`, chacun avec la classe `key`.

# --hints--

Tu devrais créer sept nouveaux éléments `div`.

```js
const divDivDiv = document.querySelectorAll('div');
assert.lengthOf(divDivDiv, 9);
```

Tes sept nouveaux éléments `div` devraient être dans ton élément `.keys`.

```js
const keys = document.querySelector('.keys');
assert.lengthOf([...keys?.children], 7);
[...keys?.children].forEach(
  child => assert.equal(child.tagName, 'DIV')
)
```

Tes sept nouveaux éléments `div` devraient tous avoir la `class` définie sur `key`.

```js
const keys = document.querySelector('.keys');
[...keys?.children].forEach(child => assert.isTrue(child?.classList.contains('key')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys"></div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
