---
id: 612e813b3ba67633222cbe54
title: Étape 4
challengeType: 0
dashedName: step-4
---

# --description--

Rappelle-toi qu'un attribut `class` peut contenir plusieurs valeurs. Pour différencier tes touches blanches et noires, ajoute une deuxième valeur de `class`, `black--key`, à tes deuxième, troisième, cinquième, sixième et septième éléments `.key`.

# --hints--

Ton deuxième élément `.key` devrait aussi avoir une `class` de `black--key`.

```js
const key = document.querySelectorAll('.key')[1];
assert.isTrue(key?.className.includes('black--key'));
```

Ton troisième `.key` devrait avoir `black--key` dans sa `class`.

```js
const third = document.querySelectorAll('.key')?.[2];
assert.isTrue(third?.classList?.contains('black--key'));
```

Ton cinquième `.key` devrait avoir `black--key` dans sa `class`.

```js
const fifth = document.querySelectorAll('.key')?.[4];
assert.isTrue(fifth?.classList?.contains('black--key'));
```

Ton sixième `.key` devrait avoir `black--key` dans sa `class`.

```js
const sixth = document.querySelectorAll('.key')?.[5];
assert.isTrue(sixth?.classList.contains('black--key'));
```

Ton septième `.key` devrait avoir `black--key` dans sa `class`.

```js
const seventh = document.querySelectorAll('.key')?.[6];
assert.isTrue(seventh?.classList?.contains('black--key'));
```

Tu devrais avoir cinq éléments `.black--key`.

```js
const blackKeys = document.querySelectorAll('.black--key');
assert.lengthOf(blackKeys, 5);
```

Tu devrais avoir sept éléments `.key`.

```js
const keys = document.querySelectorAll('.key');
assert.lengthOf(keys, 7);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
