---
id: 612e9d142affc44a453655db
title: Étape 14
challengeType: 0
dashedName: step-14
---

# --description--

Il est temps de styliser les touches elles-mêmes. Crée un sélecteur de `class` pour les éléments `.key`. Définis `background-color` sur la valeur `#ffffff`, la propriété `position` sur `relative`, la propriété `width` sur `41px` et la propriété `height` sur `175px`.

# --hints--

Tu devrais avoir un sélecteur `.key`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.key'));
```

Ton sélecteur `.key` devrait avoir une propriété `background-color` définie sur `#ffffff`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key')?.backgroundColor, 'rgb(255, 255, 255)');
```

Ton sélecteur `.key` devrait avoir la propriété `position` définie sur `relative`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key')?.position,'relative');
```

Ton sélecteur `.key` devrait avoir une propriété `width` définie sur `41px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key')?.width, '41px');
```

Ton sélecteur `.key` devrait avoir une propriété `height` définie sur `175px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key')?.height, '175px');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

--fcc-editable-region--

--fcc-editable-region--
```
