---
id: a97fd23d9b809dac9921074f
title: Crée une fonction de somme à arguments optionnels
challengeType: 26
dashedName: build-an-optional-arguments-sum-function
---

# --description--

Dans ce lab, tu vas créer une fonction qui accepte jusqu'à deux arguments et les additionne, mais si un seul argument est fourni, elle renvoie une fonction qui attend le second nombre pour faire la somme.

**Objectif :** Remplis les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User Stories :**

1. Crée une fonction `addTogether`.
2. Si la fonction reçoit deux arguments, elle doit renvoyer la somme des deux arguments.
3. Si un seul argument est fourni, `addTogether` doit renvoyer une fonction.
   - Lorsque la fonction renvoyée est appelée avec un seul argument, elle doit renvoyer la somme.

   ```js
   const sumTwoAnd = addTogether(2);
   sumTwoAnd(3); // 5
   ```

4. Si l'un des arguments n'est pas un nombre valide, renvoie `undefined`.

# --hints--

Tu dois avoir la fonction `addTogether`.

```js
assert.isFunction(addTogether);
```

`addTogether(2, 3)` devrait renvoyer 5.

```js
assert.deepEqual(addTogether(2, 3), 5);
```

`addTogether(23.4, 30)` devrait renvoyer 53.4.

```js
assert.deepEqual(addTogether(23.4, 30), 53.4);
```

`addTogether("2", 3)` devrait renvoyer `undefined`.

```js
assert.isUndefined(addTogether('2', 3));
```

`addTogether(5, undefined)` devrait renvoyer `undefined`.

```js
assert.isUndefined(addTogether(5, undefined));
```

`addTogether("https://www.youtube.com/watch?v=dQw4w9WgXcQ")` devrait renvoyer `undefined`.

```js
assert.isUndefined(addTogether('https://www.youtube.com/watch?v=dQw4w9WgXcQ'));
```

`addTogether(5)` devrait renvoyer une fonction.

```js
assert.isFunction(addTogether(5));
```

`addTogether(5)(7)` devrait renvoyer 12.

```js
assert.deepEqual(addTogether(5)(7), 12);
```

`addTogether(2)([3])` devrait renvoyer `undefined`.

```js
assert.isUndefined(addTogether(2)([3]));
```

`addTogether(2, "3")` devrait renvoyer `undefined`.

```js
assert.isUndefined(addTogether(2, '3'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function addTogether() {
  const first = arguments[0];
  if (typeof(first) !== 'number') {
    return undefined;
  }
  if (arguments.length === 1) {
    return function(second) {
      if (typeof(second) !== 'number') {
        return undefined;
      }
      return first + second;
    };
  }
  const second = arguments[1];
  if (typeof(second) !== 'number') {
    return undefined;
  }
  return first + second;
}
```
