---
id: 669e81368e52b3a5c35a2dc5
title: Construire une page de compilation de vidéos
challengeType: 25
dashedName: build-a-video-compilation-page
demoType: onClick
---

# --description--

**Objectif :** réponds aux user stories ci-dessous et fais passer tous les tests pour terminer le labo.

**User Stories :**

1. Tu dois avoir un élément `main` comme seul enfant de l'élément `body`.
1. Tu dois avoir un élément `h1` avec le sujet de ta page.
1. Tu dois avoir un paragraphe qui introduit le sujet de ta page sous ton élément `h1`.
1. Tu dois avoir trois éléments `section` sous ton premier paragraphe.
1. Chaque section doit contenir un élément `h2`, un paragraphe et un élément `iframe`, dans cet ordre.
1. Les trois éléments `iframe` doivent avoir un attribut `src` défini sur une vidéo valide.
1. Chaque élément `iframe` doit aussi avoir un attribut `title` pour décrire le contenu intégré, ainsi qu'un attribut `height` et un attribut `width` pour définir une taille correcte.

# --hints--

Tu dois avoir un élément `main` dans ton élément `body`.

```js
assert.equal(document.querySelector('body :first-child').tagName, 'MAIN');
```

Ton élément `main` doit être le seul enfant de l'élément `body`.

```js
assert.equal(document.querySelector('body').children.length, 1);
```

Tu dois avoir un élément `h1` avec le sujet de ta page dans l'élément `main`.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

Tu dois avoir un paragraphe qui introduit le sujet de ta page sous ton élément `h1`.

```js
const p = document.querySelectorAll('p')[0];
assert.equal(p?.previousElementSibling?.tagName, 'H1');
assert.isAbove(p?.innerText.length, 0);
```

Tu dois avoir trois éléments `section` sous ton premier élément `p`.

```js
assert.equal(document.querySelectorAll('section')?.length, 3);
assert.equal(document.querySelectorAll('section')[0]?.previousElementSibling?.tagName, 'P')
```

Chaque élément `section` doit commencer par un élément `h2` qui sert de titre à cette section.

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[0]?.tagName, 'H2');
    assert.isAbove(section.children[0]?.innerText.length, 0);
}
```

Chaque élément `section` doit contenir un élément `p` comme deuxième enfant pour présenter le contenu de la vidéo.

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[1]?.tagName, 'P');
    assert.isAbove(section.children[1]?.innerText.length, 0);
}
```

Chaque élément `section` doit contenir un élément `iframe` comme troisième enfant.

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[2]?.tagName, 'IFRAME');
}
```

Ton premier élément `iframe` doit avoir un attribut `src` défini sur une vidéo valide.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.src.length, 0)
```

Ton premier élément `iframe` doit avoir un attribut `title` pour décrire le contenu intégré.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.title.length, 0)
```

Ton premier élément `iframe` doit avoir un attribut `height`.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.height.length, 0)
```

Ton premier élément `iframe` doit avoir un attribut `width`.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.width.length, 0)
```

Ton deuxième élément `iframe` doit avoir un attribut `src` défini sur une vidéo valide.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.src.length, 0)
```

Ton deuxième élément `iframe` doit avoir un attribut `title` pour décrire le contenu intégré.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.title.length, 0)
```

Ton deuxième élément `iframe` doit avoir un attribut `height`.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.height.length, 0)
```

Ton deuxième élément `iframe` doit avoir un attribut `width`.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.width.length, 0)
```

Ton troisième élément `iframe` doit avoir un attribut `src` défini sur une vidéo valide.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.src.length, 0)
```

Ton troisième élément `iframe` doit avoir un attribut `title` pour décrire le contenu intégré.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.title.length, 0)
```

Ton troisième élément `iframe` doit avoir un attribut `height`.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.height.length, 0)
```

Ton troisième élément `iframe` doit avoir un attribut `width`.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.width.length, 0)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Page de compilation de vidéos</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Page de compilation de vidéos</title>
</head>

<body>
    <main>
        <h1>Développement web front-end</h1>
        <p>Le développement web front-end consiste à concevoir et construire les interfaces utilisateur des sites web avec HTML, CSS et
            JavaScript. Il se concentre sur la création de sites visuellement agréables, interactifs et faciles à utiliser. Les développeurs
            front-end assurent l'accessibilité, l'optimisation des performances et
            le respect des standards du web, en mélangeant créativité et expertise technique.</p>
        <section>
            <h2>HTML</h2>
            <p>HTML, ou HyperText Markup Language, forme la base des pages web en définissant leur structure et leur
                mise en page de base. Il utilise une série d'éléments pour afficher du texte, des images, des listes, des tableaux et d'autres contenus statiques sur
                une page web.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/GDGejH3SDNQ?si=KJYLgcz4kyyroYMB"
                title="vidéo d'introduction à html et au code" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
        <section>
            <h2>CSS</h2>
            <p>CSS, ou Cascading Style Sheets, sert à contrôler la présentation visuelle des éléments HTML sur une
                page web. Il permet de personnaliser les couleurs, les polices, les mises en page et les espacements, ce qui rend les sites
                visuellement agréables. CSS prend aussi en charge le design responsive, pour que les pages web s'affichent bien sur
                tous les appareils.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/OXGznpKZ_sA?si=s3KDSZvrhU_PU9XL"
                title="vidéo tutoriel css" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
        <section>
            <h2>JavaScript</h2>
            <p>JavaScript est un langage de programmation qui ajoute de l'interactivité aux pages web statiques, avec des mises à jour de contenu dynamiques,
                des validations de formulaires, des animations et des comportements responsives basés sur les interactions et
                événements de l'utilisateur.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/jS4aFq5-91M?si=zKQhHEYwU4tnMmVm"
                title="vidéo de programmation javascript" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
    </main>
</body>

</html>
```
