---
id: 6909ed80e1c195f49564edac
title: Étape 1
challengeType: 1
dashedName: step-1
---

# --description--

Dans ce workshop, tu apprendras à travailler avec les instructions conditionnelles et les opérateurs de comparaison.

Commence par déclarer une variable nommée `truthyOrFalsy` et affecte-lui la valeur booléenne `true`.

# --hints--

Tu devrais avoir une variable nommée `truthyOrFalsy`.

```js
assert.exists(truthyOrFalsy);
```

Tu devrais affecter `true` à ta variable `truthyOrFalsy`.

```js
assert.isTrue(truthyOrFalsy);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
