---
id: 690a58cc4f828a49eab1b53a
title: Étape 6
challengeType: 1
dashedName: step-6
---

# --description--

Déclare une variable nommée `hasDeveloperJob` et affecte-lui la valeur booléenne `true`.

# --hints--

Tu devrais avoir une variable nommée `hasDeveloperJob`.

```js
assert.exists(hasDeveloperJob);
```

Tu devrais affecter `true` à ta variable `hasDeveloperJob`.

```js
assert.isTrue(hasDeveloperJob);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
