---
id: 690a06833d3e83b5226ade78
title: Étape 7
challengeType: 1
dashedName: step-7
---

# --description--

Une instruction conditionnelle est utilisée pour exécuter du code en fonction d'une condition. Elle est composée du mot-clé `if`, suivi d'une condition entre parenthèses, et d'un corps entre accolades. Voici un exemple d'instruction `if` :

```js
if (condition) {
  console.log("condition is truthy");
}
```

Quand `condition` est une valeur truthy (chaîne non vide, nombre non nul, `true`, etc.), le code à l'intérieur du corps de l'instruction `if` est exécuté.

Sous ta variable, crée une instruction `if` qui utilise `hasDeveloperJob` comme condition. À l'intérieur du corps de ton instruction `if`, affiche `"Timmy is employed as a developer."` dans la console.

# --hints--

Tu devrais avoir une instruction `if`.

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(.+\)\s*\{?/);
```

Ton instruction `if` devrait utiliser `hasDeveloperJob` comme condition.

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(\s*hasDeveloperJob\s*\)\s*\{?/);
```

Tu devrais afficher `"Timmy is employed as a developer."` dans la console à l'intérieur de ton instruction `if`.

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(\s*hasDeveloperJob\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is employed as a developer.\1\s*\)\s*;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;
--fcc-editable-region--

--fcc-editable-region--
```
