---
id: 690a0df27ee9ee36d935d530
title: Étape 8
challengeType: 1
dashedName: step-8
---

# --description--

Déclare une autre variable nommée `isTimmyAGamer` et affecte-lui le booléen `false`.

# --hints--

Tu devrais déclarer une variable nommée `isTimmyAGamer`.

```js
assert.exists(isTimmyAGamer);
```

Tu devrais affecter la valeur booléenne `false` à ta variable `isTimmyAGamer`.

```js
assert.isFalse(isTimmyAGamer);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}
--fcc-editable-region--

--fcc-editable-region--
```
