---
id: 690a3ce03013be44d116136e
title: Étape 12
challengeType: 1
dashedName: step-12
---

# --description--

Une instruction conditionnelle peut avoir une clause `else`, qui exécute du code quand la condition `if` est falsy. Voici un exemple d'instruction `if...else` :

```js
if (condition) {
  console.log("condition is truthy");
} else {
  console.log("condition is falsy");
}
```

Ajoute une clause `else` à l'instruction `if` existante. À l'intérieur du corps de ta clause `else`, affiche `"Timmy is not old enough to drive."` dans la console.

# --hints--

Ta troisième instruction `if` devrait avoir une clause `else`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is old enough to drive.\1\s*\)\s*;?\s*\}\s+else\s+{/s);
```

Tu devrais afficher `"Timmy is not old enough to drive."` dans la console à l'intérieur du corps de ta clause `else`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is old enough to drive.\1\s*\)\s*;?\s*\}\s+else\s+{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is not old enough to drive.\2\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

const timmyAge = 18;

if (timmyAge >= 16) {
  console.log("Timmy is old enough to drive.");
--fcc-editable-region--
}
--fcc-editable-region--
```
