---
id: 690a562fd611e8155599ce98
title: Étape 3
challengeType: 1
dashedName: step-3
---

# --description--

Change l'affectation de `truthyOrFalsy` en définissant sa valeur à `"freeCodeCamp"`. Tu verras toujours `true` affiché dans la console car une chaîne non vide est une valeur truthy.

# --hints--

Ta variable `truthyOrFalsy` devrait avoir la valeur `"freeCodeCamp"`.

```js
assert.strictEqual(truthyOrFalsy, "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = true;
--fcc-editable-region--

console.log(Boolean(truthyOrFalsy));
```
