---
id: 690a57095c4afb27bbe751b3
title: Étape 4
challengeType: 1
dashedName: step-4
---

# --description--

Les valeurs falsy sont des valeurs qui s'évaluent à `false` dans un contexte booléen. Voici des exemples de valeurs falsy :

- `""` (chaîne vide)
- `0`
- `false`
- `null`
- `undefined`
- `NaN`

Tu rencontreras toutes ces valeurs au fil de ce cours.

Pour l'instant, change l'affectation de `truthyOrFalsy` en définissant sa valeur à une chaîne vide. Après cela, tu verras `false` affiché dans la console.

# --hints--

Tu devrais affecter une chaîne vide à `truthyOrFalsy`.

```js
assert.strictEqual(truthyOrFalsy, "");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = "freeCodeCamp";
--fcc-editable-region--

console.log(Boolean(truthyOrFalsy));
```
