---
id: 690a0f7a5c22c15222a6aa5c
title: Étape 9
challengeType: 1
dashedName: step-9
---

# --description--

Maintenant, crée une autre instruction `if` qui utilise `isTimmyAGamer` comme condition. À l'intérieur du corps de ton instruction `if`, affiche `"Timmy loves to play World of Warcraft."` dans la console.

Remarque que tu ne verras rien s'afficher dans la console car la condition de ta nouvelle instruction `if` est fausse.

# --hints--

Tu devrais avoir deux instructions `if` dans ton code.

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(.+\)\s*\{.+\}/s);
```

Ta deuxième instruction `if` devrait utiliser `isTimmyAGamer` comme condition.

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(\s*isTimmyAGamer\s*\)\s*\{.+\}/s);
```

Tu devrais afficher `"Timmy loves to play World of Warcraft."` dans la console à l'intérieur de ta deuxième instruction `if`.

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(\s*isTimmyAGamer\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy loves to play World of Warcraft.\1\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

--fcc-editable-region--

--fcc-editable-region--
```
