---
id: 690a51806fb743bc896c3af7
title: Étape 2
challengeType: 1
dashedName: step-2
---

# --description--

Comme tu l'as appris dans une leçon précédente, les booléens sont un type de données qui n'a que les valeurs `true` et `false`.

Une valeur truthy est une valeur qui s'évalue à `true` dans un contexte booléen. Voici des exemples de valeurs truthy :

- Les chaînes non vides
- Les nombres non nuls
- Le booléen `true`

Et d'autres valeurs que tu rencontreras plus tard.

Tu peux utiliser la fonction `Boolean()` pour vérifier le caractère truthy d'une valeur. Par exemple, `Boolean("Hello World!")` renverra `true` car `"Hello World!"` est truthy.

Maintenant, appelle `Boolean(truthyOrFalsy)` et affiche le résultat dans la console.

# --hints--

Tu devrais afficher `Boolean(truthyOrFalsy)` dans la console.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*Boolean\s*\(\s*truthyOrFalsy\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
const truthyOrFalsy = true;
--fcc-editable-region--

--fcc-editable-region--
```
