---
id: 690a110a76d0496d9e7c28eb
title: Étape 10
challengeType: 1
dashedName: step-10
---

# --description--

Déclare une autre variable nommée `timmyAge` et définis-la à `18`.

# --hints--

Tu devrais avoir une variable nommée `timmyAge`.

```js
assert.exists(timmyAge);
```

Tu devrais affecter `18` à ta variable `timmyAge`.

```js
assert.strictEqual(timmyAge, 18);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

--fcc-editable-region--

--fcc-editable-region--
```
