---
id: 690a57fef66aac39dffc61a2
title: Étape 5
challengeType: 1
dashedName: step-5
---

# --description--

Maintenant que tu sais ce que sont les valeurs truthy et falsy, supprime de ton code à la fois ta déclaration de variable et ton instruction `console.log`.

# --hints--

Tu ne devrais pas avoir de variable `truthyOrFalsy`.

```js
assert.notMatch(__helpers.removeJSComments(code), /const\s+truthyOrFalsy\s*=/);
```

Tu ne devrais pas avoir `console.log(Boolean(truthyOrFalsy))` dans ton code.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*Boolean\s*\(\s*truthyOrFalsy\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = "";

console.log(Boolean(truthyOrFalsy));
--fcc-editable-region--
```
