---
id: 690a3ad44874a928a2fe3077
title: Étape 11
challengeType: 1
dashedName: step-11
---

# --description--

Les opérateurs de comparaison te permettent de comparer des valeurs et d'écrire des instructions conditionnelles plus complexes.

```js
if (today == birthday) {
  console.log("Happy Birthday!");
}
```

Par exemple, le code ci-dessus affiche `"Happy Birthday!"` dans la console uniquement quand les valeurs des variables `today` et `birthday` sont égales.

Crée une troisième instruction `if`. Pour sa condition, utilise l'opérateur `>=` pour vérifier si `timmyAge` est supérieur ou égal à `16`.

À l'intérieur du corps de ta nouvelle instruction `if`, affiche `"Timmy is old enough to drive."` dans la console.

# --hints--

Tu devrais avoir une troisième instruction `if` dans ton code.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(.+\)\s*\{.+\}/s);
```

Ta troisième instruction `if` devrait utiliser `timmyAge >= 16` comme condition.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{.+\}/s);
```

Tu devrais afficher `"Timmy is old enough to drive."` dans la console à l'intérieur de ta troisième instruction `if`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is old enough to drive.\1\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

const timmyAge = 18;

--fcc-editable-region--

--fcc-editable-region--
```
