---
id: 61537c9eecea6a335db6da79
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

Sous ton élément `.header`, crée un nouvel élément `div` et attribue-lui une `class` de `gallery`. Ce `div` servira de conteneur pour les images de la galerie.

À l'intérieur de cet élément `.gallery`, crée neuf éléments `img`.

# --hints--

Tu devrais créer un élément `div` dans ton élément `body`.

```js
assert.isAtLeast(document.querySelector('body')?.querySelectorAll('div')?.length, 1);
```

Ton nouvel élément `div` devrait avoir une `class` avec `gallery` définie comme valeur.

```js
assert.exists(document.querySelector('body')?.querySelector('.gallery'));
```

Ton nouvel élément `div` devrait venir après ton élément `.header`.

```js
assert.exists(document.querySelector('header')?.nextElementSibling?.classList?.contains('gallery'));
```

Ton élément `.gallery` devrait avoir neuf éléments `img`.

```js
assert.lengthOf(document.querySelector('.gallery')?.querySelectorAll('img'), 9);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```

```css

```
