---
id: 61537c5f81f0cf325b4a854c
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

Ajoute un élément `header` dans l'élément `body` et attribue-lui une classe `header`.

À l'intérieur du `header`, crée un `h1` avec `css flexbox photo gallery` comme texte.

# --hints--

Tu devrais avoir un élément `header` dans ton élément `body`.

```js
assert.exists(document.querySelector('body')?.querySelector('header'));
```

Ton élément `header` devrait avoir un `class` avec `header` comme valeur.

```js
assert.isTrue(document?.querySelector('body')?.querySelector('header')?.classList?.contains('header'));
```

Ton élément `header` devrait avoir un élément `h1` à l'intérieur.

```js
assert.exists(document.querySelector('.header')?.querySelector('h1'));
```

Ton élément `h1` devrait avoir le texte `css flexbox photo gallery` à l'intérieur.

```js
assert.equal(document?.querySelector('.header')?.querySelector('h1')?.innerText, 'css flexbox photo gallery');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head lang="en">
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```

```css

```
