---
id: 6153893900438b4643590367
title: Étape 11
challengeType: 0
dashedName: step-11
---

# --description--

Aligne le texte de ton `.header` au centre. Mets le texte en majuscules avec la propriété `text-transform` et la valeur `uppercase`.

Donne-lui un padding de `32px` sur tous les côtés. Définis le fond à `#0a0a23` et le texte à `#fff` comme valeurs de couleur.

Ajoute un `border-bottom` avec `4px solid #fdb347` comme valeur.

# --hints--

Tu devrais avoir un sélecteur `.header`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.header'));
```

Ton sélecteur `.header` devrait avoir une propriété `text-align` définie à `center` comme valeur.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.textAlign, 'center');
```

Ton sélecteur `.header` devrait avoir une propriété `text-transform` définie à `uppercase` comme valeur.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.textTransform, 'uppercase');
```

Ton sélecteur `.header` devrait avoir une propriété `padding` définie à `32px` comme valeur.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.padding, '32px');
```

Ton sélecteur `.header` devrait avoir une propriété `background-color` définie à `#0a0a23` comme valeur.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.backgroundColor, 'rgb(10, 10, 35)');
```

Ton sélecteur `.header` devrait avoir une propriété `color` définie à `#fff` comme valeur.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.color, 'rgb(255, 255, 255)');
```

Ton sélecteur `.header` devrait avoir une propriété `border-bottom` définie à `4px solid #fdb347` comme valeur.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.borderBottom, '4px solid rgb(253, 179, 71)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}
body  {
  margin: 0;
  font-family: sans-serif;
  background-color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
