---
id: 61537d86bdc3dd343688fceb
title: Étape 4
challengeType: 0
dashedName: step-4
---

# --description--

Ensuite, donne à chaque `img` un attribut `src` selon son ordre dans le document. La première `img` devrait avoir un `src` de `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg`. Le reste devrait être identique, sauf en remplaçant le `1` par le numéro de l'`img` dans le document.

# --hints--

Tes neuf éléments `img` devraient tous avoir un attribut `src`.

```js
const images = [...document.querySelectorAll('img')];
assert.isTrue(images.every(image => image.getAttribute('src')));
```

Ton premier élément `img` devrait avoir `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg` comme valeur de l'attribut `src`.

```js
assert.equal(document.querySelectorAll('img')?.[0]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

Ton deuxième élément `img` devrait avoir `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg` comme valeur de l'attribut `src`.

```js
assert.equal(document.querySelectorAll('img')?.[1]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg');
```

Ton troisième élément `img` devrait avoir `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg` comme valeur de l'attribut `src`.

```js
assert.equal(document.querySelectorAll('img')?.[2]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg');
```

Ton quatrième élément `img` devrait avoir `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg` comme valeur de l'attribut `src`.

```js
assert.equal(document.querySelectorAll('img')?.[3]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg');
```

Ton cinquième élément `img` devrait avoir `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg` comme valeur de l'attribut `src`.

```js
assert.equal(document.querySelectorAll('img')?.[4]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg');
```

Ton sixième élément `img` devrait avoir `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg` comme valeur de l'attribut `src`.

```js
assert.equal(document.querySelectorAll('img')?.[5]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg');
```

Ton septième élément `img` devrait avoir `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg` comme valeur de l'attribut `src`.

```js
assert.equal(document.querySelectorAll('img')?.[6]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg');
```

Ton huitième élément `img` devrait avoir `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg` comme valeur de l'attribut `src`.

```js
assert.equal(document.querySelectorAll('img')?.[7]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg');
```

Ton neuvième élément `img` devrait avoir `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg` comme valeur de l'attribut `src`.

```js
assert.equal(document.querySelectorAll('img')?.[8]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
--fcc-editable-region--
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css

```
