---
id: 6153908a366afb4d57185c8d
title: Étape 13
challengeType: 0
dashedName: step-13
---

# --description--

Flexbox a un axe principal et un axe transversal. L'axe principal est défini par la propriété `flex-direction`, qui a quatre valeurs possibles :

- `row` (défaut) : axe horizontal avec les éléments flex de gauche à droite
- `row-reverse` : axe horizontal avec les éléments flex de droite à gauche
- `column` : axe vertical avec les éléments flex de haut en bas
- `column-reverse` : axe vertical avec les éléments flex de bas en haut

**Note** : Les axes et directions seront différents selon la direction du texte. Les valeurs affichées sont pour une direction de texte de gauche à droite.

Essaie les différentes valeurs pour voir comment elles affectent la mise en page.

Quand tu as terminé, définis explicitement un `flex-direction` de `row` sur l'élément `.gallery`.

# --hints--

Ton sélecteur `.gallery` devrait avoir une propriété `flex-direction` définie à `row` comme valeur.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.flexDirection, 'row');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}


.gallery {
  display: flex;
--fcc-editable-region--
  
--fcc-editable-region--
}

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
