---
id: 615380dff67172357fcf0425
title: Étape 9
challengeType: 0
dashedName: step-9
---

# --description--

Maintenant tes images sont trop grandes.

Crée un sélecteur `.gallery img` pour les cibler. Donne-leur toutes une `width` de `100%` et une `max-width` de `350px`.

Définis aussi la propriété `height` à `300px` pour garder tes images d'une taille uniforme.

# --hints--

Tu devrais avoir un sélecteur `.gallery img`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery img'));
```

Ton sélecteur `.gallery img` devrait avoir une propriété `width` définie à `100%` comme valeur.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery img')?.width, '100%');
```

Ton sélecteur `.gallery img` devrait avoir une propriété `max-width` définie à `350px` comme valeur.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery img')?.maxWidth, '350px');
```

Ton sélecteur `.gallery img` devrait avoir une propriété `height` définie à `300px` comme valeur.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery img')?.height, '300px');
```

# --seed--


## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

--fcc-editable-region--

--fcc-editable-region--
```
