---
id: 6493bc0d99879635209565aa
title: Étape 7
challengeType: 0
dashedName: step-7
---

# --description--

Le modèle de dimensionnement `border-box` fait le contraire de `content-box`. La largeur totale de l'élément, incluant le padding et la border, sera la largeur explicite définie. Le contenu de l'élément rétrécira pour faire de la place au padding et à la border.

Change la propriété `box-sizing` à `border-box`. Remarque comment tes bordures bleues d'image rentrent maintenant à l'intérieur de la bordure rouge de la galerie.

# --hints--


Ton sélecteur * devrait avoir une propriété `box-sizing` définie à `border-box` comme valeur.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing, 'border-box');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>


```

```css
 * {
--fcc-editable-region--
  box-sizing: content-box;
--fcc-editable-region--
}

.gallery {
  border: 5px solid red;
  width: 50%;
}

img {
  width: 100%;
  border: 5px solid blue;
  padding: 5px;
}
```
