---
id: 6494da0daf5df5197963671d
title: Étape 6
challengeType: 0
dashedName: step-6
---

# --description--

Remarque que la bordure bleue de l'image dépasse la bordure rouge de la galerie. C'est dû à la façon dont les navigateurs calculent la taille des éléments conteneurs.

La propriété `box-sizing` est utilisée pour définir ce comportement. Par défaut, le modèle `content-box` est utilisé. Avec ce modèle, quand un élément a une largeur spécifique, cette largeur est calculée en se basant uniquement sur le contenu de l'élément. Les valeurs de padding et de border sont ajoutées à la largeur totale, donc l'élément grandit pour les accueillir.

Essaie de définir explicitement `box-sizing` à `content-box`, avec le sélecteur global `*`. À ce stade, tu ne verras aucun changement, car tu utilises la valeur par défaut.

# --hints--

Tu devrais avoir un sélecteur `*`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

Ton sélecteur `*` devrait avoir une propriété `box-sizing`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing);
```

Ton sélecteur `*` devrait avoir sa propriété `box-sizing` définie à `content-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing, 'content-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
  ```

```css
--fcc-editable-region--

--fcc-editable-region--

.gallery {
  border: 5px solid red;
  width: 50%;
}

img {
  width: 100%;
  border: 5px solid blue;
  padding: 5px;
}
```
