---
id: 64dd076a1b33c86d84b4232f
title: Étape 5
challengeType: 0
dashedName: step-5
---

# --description--

Pour mieux visualiser comment tes éléments sont dimensionnés, ajouter une bordure peut être utile.

Donne à ton élément `.gallery` une `width` de `50%` et une `border` définie à `5px solid red`.

Ensuite, donne à tes éléments `img` une `width` de `100%`, un `padding` défini à `5px`, et une `border` définie à `5px solid blue`.

# --hints--

Tu devrais avoir un sélecteur `.gallery`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery'));
```

Ton sélecteur `.gallery` devrait avoir une propriété `width`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.width);
```

Ton sélecteur `.gallery` devrait avoir une propriété `width` définie à `50%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.width, '50%');
```

Ton sélecteur `.gallery` devrait avoir une propriété `border`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.border);
```

Ton sélecteur `.gallery` devrait avoir une propriété `border` définie à `5px solid red`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.border, '5px solid red');
```

Tu devrais avoir un sélecteur `img`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img'));
```

Ton sélecteur `img` devrait avoir une propriété `width`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.width);
```

Ton sélecteur `img` devrait avoir une propriété `width` définie à `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.width, '100%');
```

Ton sélecteur `img` devrait avoir une propriété `padding`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.padding);
```

Ton sélecteur `img` devrait avoir une propriété `padding` définie à `5px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.padding, '5px');
```

Ton sélecteur `img` devrait avoir une propriété `border`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.border);
```

Ton sélecteur `img` devrait avoir une propriété `border` définie à `5px solid blue`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.border, '5px solid blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
