---
id: 66bb3e20d3dc5b6d0a21f5dd
title: Construire un compteur en temps réel
challengeType: 25
dashedName: build-a-real-time-counter
demoType: onClick
---

# --description--

Dans ce lab, tu vas construire un compteur de caractères en temps réel avec JavaScript. Le compteur affichera le nombre de caractères saisis dans un élément `textarea`. Il se mettra à jour en temps réel à mesure que l'utilisateur saisit du texte dans le `textarea`.

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois avoir un élément `textarea` avec l'`id` `text-input`.
2. Tu dois avoir un élément `p` avec l'`id` `char-count`. Il doit initialement contenir le texte `Character Count: 0/50`. Ce texte indicatif doit être remplacé par le nombre actuel de caractères dans l'élément `textarea`.
3. Lorsque l'élément `#text-input` contient le texte `hello coder`, l'élément `#char-count` doit contenir le texte `"Character Count: 11/50"`.
4. Lorsque l'élément `#text-input` contient le texte `i am a programmer`, l'élément `#char-count` doit contenir le texte `"Character Count: 17/50"`.
5. Lorsque l'élément `#text-input` contient le texte `hello world`, l'élément `#char-count` doit contenir le texte `"Character Count: 11/50"`.
6. Lorsque l'élément `#text-input` contient le texte `I am learning a new language and it is called c++.`, l'élément `#char-count` doit contenir le texte `"Character Count: 50/50"`. **REMARQUE** : bien que l'attribut `maxlength` permette d'obtenir le bon comportement, ce lab exige une solution reposant uniquement sur JavaScript.
7. L'utilisateur ne doit pas pouvoir saisir plus de `50` caractères dans l'élément `textarea`. Lorsque le nombre de caractères atteint `50`, toute saisie supplémentaire doit être automatiquement tronquée et le texte `Character Count: 50/50` doit s'afficher en `red`.

**Remarque :** assure-toi de lier ta feuille de style et ton fichier JavaScript à ton fichier HTML.

# --hints--

Tu dois avoir un élément `textarea` avec l'`id` `text-input`.

```js
const textInput = document.querySelector('textarea');
assert.equal(textInput.id, 'text-input');
```

Tu dois avoir un élément `p` avec l'`id` `char-count`.

```js
assert.exists("p#char-count");
```

L'élément `#char-count` doit initialement contenir le texte `Character Count: 0/50`.

```js
const charCount = document.querySelector('p#char-count')?.innerText;
assert.equal(charCount, 'Character Count: 0/50');
```

Lorsque l'élément `#text-input` contient le texte `hello coder`, l'élément `#char-count` doit contenir le texte `"Character Count: 11/50"`.

```js
const textInput = document.getElementById('text-input');
const charCount = document.getElementById('char-count');
textInput.value = "hello coder";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
assert.strictEqual(charCount.innerText.trim(), 'Character Count: 11/50');
```

Lorsque le nombre de caractères est égal à `50`, le texte doit s'afficher en `red`.

```js
const textInput = document.getElementById('text-input');
const charCount = document.getElementById('char-count');
textInput.value = "I am learning a new language and it is called c++.";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
const computedStyle = window.getComputedStyle(charCount);
assert.oneOf(computedStyle.color, [
  "red",
  "rgb(255, 0, 0)",
  "rgb(100%, 0%, 0%)",
  "rgba(255, 0, 0, 1)",
  "rgba(100%, 0%, 0%, 1)",
  "hsl(0, 100%, 50%)",
  "hsla(0, 100%, 50%, 1)"
]);
```

Si le nombre de caractères est supérieur ou égal à `50`, la saisie supplémentaire doit être tronquée.

```js
const textInput = document.getElementById('text-input');
textInput.value = "I am learning a new programming language and it is called c++.";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
assert.strictEqual(textInput.value.length, 50);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Real Time Counter</title>
</head>

<body>

</body>


</html>
```

```css

```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Real-time Character Counter</title>
    <link rel="stylesheet" href="styles.css">

<body>
    <div class="container">
        <h1>Real-Time Character Counter</h1>
        <textarea id="text-input" placeholder="Type something..."></textarea>
        <p id="char-count">Character Count: 0/50</p>
    </div>

<script src="script.js"></script>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    font-family: Arial, sans-serif;
    background-color: #c0c1ec;
    margin: 0;
}

h1{
    background-color: #f2f48b;
    padding-left: 10px;
    padding-right: 10px;
}
.container {
    text-align: center;
}

textarea {
    width: 300px;
    height: 100px;
    font-size: 18px;
    padding: 10px;
    border: 1px solid #ccc;
    border-radius: 5px;
    resize: none;
    display: block;
    margin: 0 auto;
    border: 1px solid slategray;
    box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

textarea:focus {
    outline: none;
    border-color: #f2f48b;
}

#char-count {
    margin-top: 10px;
    font-size: 18px;
    color: darkslategray;
    font-weight: bold;
}
```

```js
document.getElementById("text-input").addEventListener("input", function () {
  let textInput = document.getElementById("text-input");
  let charCount = textInput.value.length;

  if (charCount > 50) {
    textInput.value = textInput.value.substring(0, 50);
    charCount = 50;
  }

  let charCountDisplay = document.getElementById("char-count");
  charCountDisplay.innerHTML = `Character Count: ${charCount}/50`;

  if (charCount === 50) {
    charCountDisplay.style.color = "red";
  }

});
```
