---
id: a0b5010f579e69b815e7c5d6
title: Crée une fonction de remplacement intelligent de mots
challengeType: 26
dashedName: build-a-smart-word-replacement-function
---

# --description--

Dans ce lab, tu vas créer une fonction qui effectue une recherche et un remplacement dans la phrase à l'aide des arguments fournis, puis renvoie la nouvelle phrase.

**Remarque :** conserve la casse du premier caractère du mot d'origine lors du remplacement. Par exemple, si tu veux remplacer le mot `Book` par le mot `dog`, il doit être remplacé par `Dog`

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois créer une fonction nommée `myReplace`.
1. La fonction `myReplace` doit prendre trois arguments : une chaîne, un mot à remplacer, et le mot de remplacement.
1. La fonction `myReplace` doit renvoyer la nouvelle chaîne avec le remplacement.
1. Tu dois conserver la casse du premier caractère du mot d'origine lors du remplacement.

# --hints--

Tu dois créer une fonction nommée `myReplace`.

```js
assert.isFunction(myReplace);
```

`myReplace` doit prendre trois arguments.

```js
assert.lengthOf(myReplace, 3);
```

`myReplace("Let us go to the store", "store", "mall")` doit renvoyer la chaîne `Let us go to the mall`.

```js
assert.deepEqual(
  myReplace('Let us go to the store', 'store', 'mall'),
  'Let us go to the mall'
);
```

`myReplace("He is Sleeping on the couch", "Sleeping", "sitting")` doit renvoyer la chaîne `He is Sitting on the couch`.

```js
assert.deepEqual(
  myReplace('He is Sleeping on the couch', 'Sleeping', 'sitting'),
  'He is Sitting on the couch'
);
```

`myReplace("I think we should look up there", "up", "Down")` doit renvoyer la chaîne `I think we should look down there`.

```js
assert.deepEqual(
  myReplace('I think we should look up there', 'up', 'Down'),
  'I think we should look down there'
);
```

`myReplace("This has a spellngi error", "spellngi", "spelling")` doit renvoyer la chaîne `This has a spelling error`.

```js
assert.deepEqual(
  myReplace('This has a spellngi error', 'spellngi', 'spelling'),
  'This has a spelling error'
);
```

`myReplace("His name is Tom", "Tom", "john")` doit renvoyer la chaîne `His name is John`.

```js
assert.deepEqual(
  myReplace('His name is Tom', 'Tom', 'john'),
  'His name is John'
);
```

`myReplace("Let us get back to more Coding", "Coding", "algorithms")` doit renvoyer la chaîne `Let us get back to more Algorithms`.

```js
assert.deepEqual(
  myReplace('Let us get back to more Coding', 'Coding', 'algorithms'),
  'Let us get back to more Algorithms'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function myReplace(str, before, after) {
  if (before.charAt(0) === before.charAt(0).toUpperCase()) {
    after = after.charAt(0).toUpperCase() + after.substring(1);
  } else {
    after = after.charAt(0).toLowerCase() + after.substring(1);
  }
  return str.replace(before, after);
}
```
