---
id: 67fccc2463253b213a000142
title: Construire un lecteur multimédia
challengeType: 25
dashedName: lab-multimedia-player
demoType: onClick
---

# --description--

Dans les leçons précédentes, tu as découvert le travail avec les éléments `audio` et `video`. Dans ce lab, tu vas construire un lecteur multimédia qui affichera une piste `audio` et une `video` avec une transcription.

Pour l'élément `audio`, tu devras inclure un élément `source` qui est utilisé pour préciser le média utilisé. 

Voici un exemple :

```html
<audio controls aria-label="descriptive label goes here">
  <source src="url-to-audio-goes-here" type="audio/mpeg">
</audio>
```

L'élément `source` peut aussi être utilisé dans l'élément `video` comme ceci :

```html
<video controls width="600" aria-label="descriptive label goes here">
  <source src="link-to-mp4-goes-here" type="video/mp4">
  <!-- Remaining code goes here -->  
</video>
```

**Objectif :** Réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu devrais avoir un élément `h1` pour le titre principal de la page.
2. Tu devrais avoir trois éléments `section`.
3. À l'intérieur du premier élément `section`, tu devrais avoir un élément `h2` pour le titre de la chanson qui joue.
4. Sous l'élément `h2`, tu devrais avoir un élément `audio` avec l'attribut `controls` et un attribut `aria-label`.
5. À l'intérieur de l'élément `audio`, tu devrais avoir un élément `source` avec un attribut `src` qui pointe vers un fichier audio et un attribut `type`. Tu peux utiliser cette URL audio si tu veux : `https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3`
6. À l'intérieur du deuxième élément `section`, tu devrais avoir un élément `h2` pour le titre de la vidéo qui joue.
7. Sous l'élément `h2`, tu devrais avoir un élément `video` avec les attributs `controls`, `width` et un attribut `aria-label`.
8. À l'intérieur de l'élément `video`, tu devrais avoir un élément `source` avec un attribut `src` qui pointe vers un fichier vidéo et un attribut `type`. Tu peux utiliser cette URL vidéo si tu veux : `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4`
9. Sous l'élément `source`, tu devrais avoir un élément `track` avec un attribut `src` qui pointe vers un fichier de sous-titres et un attribut `kind`, un attribut `srclang` et un attribut `label`. Tu peux utiliser cette URL de sous-titres si tu veux : `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.vtt`
10. À l'intérieur du troisième élément `section`, tu devrais avoir un élément `h2` pour le titre de la section, par exemple « Transcription ».
11. Sous l'élément `h2`, tu devrais avoir un élément `p` avec la transcription de la vidéo.

# --hints--

Tu devrais avoir un élément `h1`.

```js
assert.isNotNull(document.querySelector("h1"));
```

Ton `h1` doit avoir du texte qui représente le titre principal de la page.  

```js  
assert.isNotEmpty(document.querySelector("h1")?.textContent?.trim());  
```  

Tu devrais avoir trois éléments `section`.

```js
const sectionElements = document.querySelectorAll("section");
assert.lengthOf(sectionElements, 3);
```

La première `section` doit contenir un élément `h2`.

```js
const sectionElements = document.querySelectorAll("section");
const firstSection = sectionElements[0];
assert.isNotNull(firstSection?.querySelector("h2"));
```

Ton `h2` doit avoir du texte qui représente le titre de la chanson qui joue.

```js
const firstSection = document.querySelectorAll("section")[0];
assert.isNotEmpty(firstSection?.querySelector("h2")?.textContent?.trim());
```

La première `section` doit contenir un élément `audio`.  

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
assert.isNotNull(audio);
```

Ton élément `audio` doit avoir un attribut `controls`.  

```js  
const firstSection = document.querySelectorAll("section")[0];  
const audio = firstSection?.querySelector("audio");  
assert.isTrue(audio?.hasAttribute("controls"));
```

Ton élément `audio` doit avoir un attribut `aria-label` avec du texte qui décrit l'audio.

```js
const firstSection = document.querySelectorAll("section")[0];  
const audio = firstSection?.querySelector("audio");  
assert.isTrue(audio?.hasAttribute("aria-label"));  

const ariaLabel = audio?.getAttribute("aria-label");  
assert.isString(ariaLabel);  
assert.isNotEmpty(ariaLabel.trim());
```

Ton élément `audio` doit avoir un élément `source`.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");
assert.isNotNull(source);
```

Ton élément `source` doit avoir un attribut `src` qui pointe vers un fichier audio.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");
assert.isTrue(source?.hasAttribute("src"));
assert.isTrue(source?.getAttribute("src")?.trim().length > 0);
```

Ton élément `source` doit avoir un attribut `type` avec une valeur qui précise le type de l'audio.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");

assert.isTrue(source?.hasAttribute("type"));
assert.isTrue(source?.getAttribute("type")?.trim().length > 0);
```

La deuxième `section` doit contenir un élément `h2`.

```js
const sectionElements = document.querySelectorAll("section");
const h2 = sectionElements[1]?.querySelector("h2");
assert.isNotNull(h2);
```

Ton `h2` doit avoir du texte qui représente le titre de la vidéo qui joue.

```js
assert.isNotEmpty(document.querySelector("h2")?.textContent?.trim());
```

La deuxième `section` doit contenir un élément `video`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isNotNull(video);
```

Ton élément `video` doit avoir un attribut `controls`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isTrue(video?.hasAttribute("controls"));
```

Ton élément `video` doit avoir un attribut `width` avec une valeur qui précise la largeur de la vidéo.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isTrue(video?.hasAttribute("width"));

const width = video?.getAttribute("width");
assert.isTrue(width.trim().length > 0);
```

L'élément `video` doit avoir un élément `source`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isNotNull(source);
```

L'élément `source` doit avoir un attribut `src` qui pointe vers un fichier vidéo.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isTrue(source?.hasAttribute("src"));
assert.isTrue(source.getAttribute("src").trim().length > 0);
```

L'élément `source` doit avoir un attribut `type` avec une valeur qui précise le type de vidéo.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isTrue(source?.hasAttribute("type"));
assert.isTrue(source.getAttribute("type").trim().length > 0);
```

Ton élément `video` doit avoir un élément `track`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isNotNull(track);
```

Ton élément `track` doit avoir un attribut `src` qui pointe vers un fichier de sous-titres.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("src"));
assert.isTrue(track?.getAttribute("src")?.trim().length > 0);
```

Ton élément `track` doit avoir un attribut `kind` avec du texte qui décrit le type de fichier.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("kind"));
assert.isTrue(track?.getAttribute("kind")?.trim().length > 0);
```

Ton élément `track` doit avoir un attribut `srclang` avec du texte qui décrit la langue des sous-titres.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("srclang"));
assert.isTrue(track?.getAttribute("srclang")?.trim().length > 0);
```

Ton élément `track` doit avoir un attribut `label` avec du texte qui décrit la langue des sous-titres.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("label"));
assert.isTrue(track?.getAttribute("label")?.trim().length > 0);
```

La troisième `section` doit contenir un élément `h2`.

```js
const sectionElements = document.querySelectorAll("section");
const thirdSection = sectionElements[2];
assert.isNotNull(thirdSection.querySelector("h2"));
```

Ton `h2` doit avoir du texte qui représente le titre de la section.

```js
assert.isNotEmpty(document.querySelector("h2")?.textContent?.trim());
```

La troisième `section` doit contenir un élément `p` pour la transcription.

```js
const thirdSection = document.querySelectorAll("section")[2];
const transcript = thirdSection?.querySelector("p");
assert.isNotNull(transcript);
assert.isAbove(transcript.textContent.trim().length, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Multimedia Player</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Multimedia Player</title>
</head>
<body>
  <h1>Multimedia Player</h1>
  <section>
    <h2>Now Playing: Sailing Away</h2>
    <audio controls aria-label="Audio Player for Sailing Away">
        <source src="https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3" type="audio/mpeg">
    </audio>
  </section>

  <section>
    <h2>What is a map method and how does it work?</h2>
    <video controls width="600" aria-label="What is a map method and how does it work?">
      <source src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4" type="video/mp4">
      
      <track 
      src="what-is-a-map-method.vtt" 
      kind="captions" 
      srclang="en" 
      label="English" 
      default>

    </video>
  </section>

  <section>
    <h2>Transcript</h2>
    <p>What is a map method and how does it work? The map method is a powerful and widely used function in JavaScript that operates on arrays. It's designed to create a new array by applying a given function to each element of the original array. This method does not modify the original array, but instead returns the new array containing the results of the function applied to each element. 
    </p>
  </section>

</body>
</html> 
```
