---
id: aa7697ea2477d1316795783b
title: Implémente un traducteur Pig Latin
challengeType: 26
dashedName: lab-pig-latin
---

# --description--

Le Pig Latin est une façon de transformer des mots anglais en suivant des règles de transformation spécifiques.

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois créer une fonction `translatePigLatin` qui accepte une chaîne comme argument.
1. Si l'argument chaîne commence par une consonne, ta fonction doit prendre la première consonne ou le premier groupe de consonnes, le déplacer à la fin du mot, y ajouter `ay`, et renvoyer le résultat.
1. Si l'argument chaîne commence par une voyelle, ta fonction doit ajouter `way` à la fin et renvoyer le résultat.
1. Ta fonction doit gérer les arguments chaîne où la première voyelle se trouve au milieu du mot et renvoyer la chaîne transformée de façon appropriée.
1. Si l'argument chaîne n'a aucune voyelle, ta fonction doit ajouter `ay` à la fin et renvoyer le résultat.

Remarque : dans le contexte de ce lab, les voyelles sont `a`, `e`, `i`, `o` et `u`. La lettre `y` n'est pas considérée comme une voyelle.

# --hints--

Tu dois avoir une fonction `translatePigLatin`.

```js
assert.isFunction(translatePigLatin);
```

`translatePigLatin("california")` doit renvoyer la chaîne `aliforniacay`.

```js
assert.deepEqual(translatePigLatin('california'), 'aliforniacay');
```

`translatePigLatin("paragraphs")` doit renvoyer la chaîne `aragraphspay`.

```js
assert.deepEqual(translatePigLatin('paragraphs'), 'aragraphspay');
```

`translatePigLatin("glove")` doit renvoyer la chaîne `oveglay`.

```js
assert.deepEqual(translatePigLatin('glove'), 'oveglay');
```

`translatePigLatin("algorithm")` doit renvoyer la chaîne `algorithmway`.

```js
assert.deepEqual(translatePigLatin('algorithm'), 'algorithmway');
```

`translatePigLatin("eight")` doit renvoyer la chaîne `eightway`.

```js
assert.deepEqual(translatePigLatin('eight'), 'eightway');
```

Doit gérer les mots où la première voyelle se trouve au milieu du mot. `translatePigLatin("schwartz")` doit renvoyer la chaîne `artzschway`.

```js
assert.deepEqual(translatePigLatin('schwartz'), 'artzschway');
```

Doit gérer les mots sans voyelles. `translatePigLatin("rhythm")` doit renvoyer la chaîne `rhythmay`.

```js
assert.deepEqual(translatePigLatin('rhythm'), 'rhythmay');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function translatePigLatin(str) {
  if (isVowel(str.charAt(0))) return str + "way";
  let front = [];
  str = str.split('');
  while (str.length && !isVowel(str[0])) {
    front.push(str.shift());
  }
  return [].concat(str, front).join('') + 'ay';
}

function isVowel(c) {
  return ['a', 'e', 'i', 'o', 'u'].indexOf(c.toLowerCase()) !== -1;
}
```
