---
id: 66d75dd0aa65a71600dc669b
title: Crée un programme de gestion d'inventaire
challengeType: 26
dashedName: build-an-inventory-management-program
---

# --description--

Dans ce lab, tu vas créer un programme de gestion d'inventaire qui te permettra d'ajouter, de mettre à jour, de trouver et de retirer des produits de l'inventaire. Tu utiliseras un tableau d'objets pour représenter ton inventaire, où chaque objet aura `name` et `quantity` comme clés.

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois déclarer un tableau vide nommé `inventory` qui stockera des objets produit ayant une clé `name` avec la valeur d'une chaîne en minuscules, et une clé `quantity` avec la valeur d'un entier.
1. Tu dois créer une fonction nommée `findProductIndex` qui prend le nom du produit en argument et renvoie l'indice de l'objet produit correspondant dans le tableau `inventory`. La fonction doit toujours utiliser la forme en minuscules du nom du produit pour effectuer la recherche. Si le produit n'est pas trouvé, la fonction doit renvoyer `-1`.
1. Tu dois créer une fonction nommée `addProduct` qui prend un objet produit en argument.
1. Si le produit est déjà présent dans l'inventaire, la fonction `addProduct` doit mettre à jour sa valeur `quantity` en ajoutant la quantité passée à la fonction à la `quantity` actuelle et afficher dans la console le nom du produit suivi d'un espace et de `quantity updated`.
1. Si le produit n'est pas présent dans l'inventaire, la fonction `addProduct` doit ajouter le produit au tableau `inventory` et afficher dans la console le nom du produit suivi d'un espace et de `added to inventory`.
1. Tu dois créer une fonction nommée `removeProduct` qui prend le nom du produit et la quantité en arguments.
1. La fonction `removeProduct` doit soustraire la quantité passée de l'objet produit correspondant dans l'inventaire et afficher dans la console la chaîne `Remaining <product-name> pieces: <product-quantity>`, où `<product-name>` doit être remplacé par le nom du produit et `<product-quantity>` par la quantité du produit.
1. Si la quantité après la soustraction est zéro, `removeProduct` doit retirer l'objet produit de l'inventaire. Si la quantité dans l'inventaire n'est pas suffisante pour effectuer la soustraction, la fonction `removeProduct` doit afficher dans la console la chaîne `Not enough <product-name> available, remaining pieces: <product-quantity>`.
1. Si le produit à retirer n'est pas présent dans l'inventaire, la fonction `removeProduct` doit afficher `<product-name> not found` dans la console.

**Note :** pour éviter les conflits, conserve uniquement les affichages mentionnés dans les user stories lors de l'exécution des tests.

# --hints--

Tu dois déclarer un tableau vide nommé `inventory`.

```js
assert.match(__helpers.removeJSComments(code), /(let|const|var)\s+inventory\s*=\s*\[\s*\]\s*;?/)
```

Tu dois avoir une fonction nommée `findProductIndex`.

```js
assert.isFunction(findProductIndex)
```

`findProductIndex("flour")` doit renvoyer l'indice de l'objet ayant `name` égal à `"flour"` dans le tableau `inventory`.

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("flour"), 0);
assert.strictEqual(findProductIndex("rice"), 1);
assert.strictEqual(findProductIndex("sugar"), 2);
```

`findProductIndex("FloUr")` doit renvoyer l'indice de l'objet ayant `name` égal à `"flour"` dans le tableau `inventory`.

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("FLour"), 0);
assert.strictEqual(findProductIndex("RICE"), 1);
assert.strictEqual(findProductIndex("suGar"), 2);
```

`findProductIndex("Flour")` doit renvoyer `-1` lorsqu'aucun objet ayant `name` égal à `"flour"` n'est trouvé dans le tableau `inventory`.

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("salt"), -1)
```

Tu dois avoir une fonction nommée `addProduct`.

```js
assert.isFunction(addProduct)
```

`addProduct({name: "FLOUR", quantity: 5})` doit ajouter `5` à la valeur `quantity` de l'objet ayant `name` égal à `"flour"` dans le tableau `inventory`.

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'Flour', quantity: 5});
addProduct({name: 'RICE', quantity: 5});
addProduct({name: 'SuGar', quantity: 5});
const expected = JSON.stringify([{name: 'flour', quantity: 20}, {name: 'rice', quantity: 15}, {name: 'sugar', quantity: 10}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`addProduct({name: "FLOUR", quantity: 5})` doit afficher `flour quantity updated` lorsqu'un objet ayant `name` égal à `"flour"` est trouvé dans le tableau `inventory`.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'flOur', quantity: 5});
addProduct({name: 'RICE', quantity: 5});
addProduct({name: 'sugar', quantity: 5});
console.log = temp;
assert.lengthOf(logArr, 3);
assert.strictEqual(logArr[0], "flour quantity updated")
assert.strictEqual(logArr[1], "rice quantity updated")
assert.strictEqual(logArr[2], "sugar quantity updated")
```

`addProduct({name: "FLOUR", quantity: 5})` doit ajouter `{name: "flour", quantity: 5}` au tableau `inventory` lorsqu'aucun objet ayant `name` égal à `"flour"` n'est trouvé dans l'inventaire.

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'SALT', quantity: 8});
addProduct({name: 'Spam', quantity: 1});
const expected = JSON.stringify([{name: 'flour', quantity: 15}, {name: 'rice', quantity: 10}, {name: 'sugar', quantity: 5}, {name: 'salt', quantity: 8}, {name: 'spam', quantity: 1}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected);
```

`addProduct({name: "FLOUR", quantity: 5})` doit afficher `flour added to inventory` lorsqu'aucun objet ayant `name` égal à `"flour"` n'est trouvé dans l'inventaire.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'SAlt', quantity: 5});
addProduct({name: 'SPAM', quantity: 1});
console.log = temp;
assert.lengthOf(logArr, 2);
assert.equal(logArr[0], "salt added to inventory");
assert.equal(logArr[1], "spam added to inventory");
```

Tu dois créer une fonction nommée `removeProduct`.

```js
assert.isFunction(removeProduct);
```

`removeProduct("FLOUR", 5)` doit afficher `flour not found` lorsqu'aucun objet ayant `name` égal à `"flour"` n'est trouvé dans le tableau `inventory`.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("SALT", 2);
removeProduct("Spam", 1);
console.log = temp;
assert.lengthOf(logArr, 2);
assert.equal(logArr[0], "salt not found");
assert.equal(logArr[1], "spam not found");
```

`removeProduct("FLOUR", 5)` doit soustraire `5` de la valeur `quantity` de l'objet ayant `name` égal à `"flour"` dans le tableau `inventory`.

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 10);
removeProduct("RICE", 5);
removeProduct("Sugar", 2);
const expected = JSON.stringify([{name: 'flour', quantity: 5}, {name: 'rice', quantity: 5}, {name: 'sugar', quantity: 3}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`removeProduct("FLOUR", 5)` doit afficher `Remaining flour pieces: 15` dans la console, lorsque `inventory` est égal à `[{name: "flour", quantity: 20}, {name: "rice", quantity: 5}]`.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("Flour", 5);
removeProduct("RICE", 5);
removeProduct("Sugar", 2);
console.log = temp;
assert.lengthOf(logArr, 3);
assert.equal(logArr[0], "Remaining flour pieces: 10");
assert.equal(logArr[1], "Remaining rice pieces: 5");
assert.equal(logArr[2], "Remaining sugar pieces: 3");
```

Si la `quantity` après la soustraction est zéro, `removeProduct` doit retirer l'objet produit de l'inventaire.

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 15);
removeProduct("rice", 10);
const expected = JSON.stringify([{name: 'sugar', quantity: 5}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`removeProduct("FLOUR", 10)` doit afficher `Not enough flour available, remaining pieces: 5` lorsque `inventory` est égal à `[{name: "flour", quantity: 5}, {name: "rice", quantity: 5}]`.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 20);
removeProduct("RICE", 20);
removeProduct("Sugar", 20);
console.log = temp;
assert.lengthOf(logArr, 3);
assert.equal(logArr[0], "Not enough flour available, remaining pieces: 15")
assert.equal(logArr[1], "Not enough rice available, remaining pieces: 10")
assert.equal(logArr[2], "Not enough sugar available, remaining pieces: 5")
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const inventory = [];

const findProductIndex = (productName) => {
    let productIndex = -1;
    inventory.forEach((element, index) => {
        if (element.name === productName.toLowerCase()) {
            productIndex = index;
        }
    })
    return productIndex;
}

const addProduct = (product) => {
    product.name = product.name.toLowerCase();
    const productIndex = findProductIndex(product.name);
    if (productIndex < 0) {
        inventory.push(product);
        console.log(`${product.name} added to inventory`)
        return;
    }
    inventory[productIndex].quantity += product.quantity;
    console.log(`${product.name} quantity updated`)
}

const removeProduct = (productName, productQuantity) => {
    productName = productName.toLowerCase();
    const productIndex = findProductIndex(productName);
    if (productIndex < 0) {
        console.log(`${productName} not found`);
        return;
    }
    const foundProduct = inventory[productIndex];
    if (foundProduct.quantity >= productQuantity) {
        foundProduct.quantity -= productQuantity;
        console.log(`Remaining ${productName} pieces: ${foundProduct.quantity}`)
        if (!foundProduct.quantity) {
            inventory.splice(productIndex, 1);
        }
        return;
    }
    console.log(`Not enough ${productName} available, remaining pieces: ${foundProduct.quantity}`);
}

// Example usage

const product1 = {
    name: 'flour',
    quantity: 15
}
const product2 = {
    name: 'flour',
    quantity: 10
}
const product3 = {
    name: 'sugar',
    quantity: 20
}
const product4 = {
    name: 'salt',
    quantity: 12
}

addProduct(product1);
addProduct(product2);
addProduct(product3);
addProduct(product4);
removeProduct('flour', 30);
removeProduct('flour', 20);
removeProduct('flour', 5);
```
