---
id: 694bf6eb530e19a9c48c59f4
title: Étape 13
challengeType: 20
dashedName: step-13
---

# --description--

Concaténer beaucoup de chaînes avec `+` et convertir des nombres avec `str()` peut devenir confus et difficile à lire.

Python 3.6 a introduit les **f-strings** pour résoudre ce problème. En ajoutant la lettre `f` avant la quote d'ouverture, tu peux placer des variables et des expressions dans des champs de remplacement représentés par des accolades `{}`. Par exemple :

```python
name = 'John'
print(f'Hello {name}') # Output: Hello John
```

Crée une variable `employee_card` et assigne-lui une f-string qui affiche `Employee:` suivi d'un espace et de la valeur de la variable `full_name`.

# --hints--

Tu devrais avoir une variable nommée `employee_card`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_card")`)) })
```

Ta variable `employee_card` devrait avoir la valeur `f'Employee: {full_name}'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name}'")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--

--fcc-editable-region--
```
