---
id: 694bf6eb530e19a9c48c59f6
title: Étape 17
challengeType: 20
dashedName: step-17
---

# --description--

Tu peux découper n'importe quelle partie de la chaîne, pas seulement le début. C'est utile dans de nombreux cas.

Crée une variable `year_code` et assigne-lui le découpage de `employee_code` de l'indice `4` à `8`. Cela extraira `2026`.

Puis crée une variable `initials` et assigne-lui le découpage de `employee_code` de l'indice `9` à `11`. Cela extraira `JD`.

Enfin, affiche les deux variables dans le terminal.

# --hints--

Tu devrais avoir une variable nommée `year_code`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("year_code")`)) })
```

Tu devrais découper `employee_code` de l'indice `4` à `8` et assigner le résultat à ta variable `year_code`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("year_code").is_equivalent("year_code = employee_code[4:8]")`)) })
```

Tu devrais avoir une variable nommée `initials`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("initials")`)) })
```

Tu devrais découper `employee_code` de l'indice `9` à `11` et assigner le résultat à ta variable `initials`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("initials").is_equivalent("initials = employee_code[9:11]")`)) })
```

Tu devrais afficher `year_code` dans le terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(year_code)")`)) })
```

Tu devrais afficher `initials` dans le terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(initials)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department)
--fcc-editable-region--

--fcc-editable-region--
```
