---
id: 694bf6ea530e19a9c48c59f0
title: Étape 3
challengeType: 20
dashedName: step-3
---

# --description--

Si tu concatènes deux chaînes comme `'John' + 'Doe'`, le résultat est `'JohnDoe'` sans espace. Pour corriger cela, tu dois concaténer une chaîne contenant un espace (`' '`) entre elles.


Mets à jour ta variable `full_name` pour qu'elle concatène `first_name`, un espace et `last_name`.

# --hints--

Tu devrais ajouter un espace (`' '`) entre `first_name` et `last_name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("full_name").is_equivalent("full_name = first_name + ' ' + last_name")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
--fcc-editable-region--
full_name = first_name + last_name
--fcc-editable-region--
print(full_name)
```
