---
id: 694bf3af05d10fd89dab05cc
title: Étape 2
challengeType: 20
dashedName: step-2
---

# --description--

En Python, tu peux combiner des chaînes avec l'opérateur `+`. C'est ce qu'on appelle la concaténation. Voici un exemple :

```python
greeting = 'Hello' + 'World'
print(greeting) # Output: HelloWorld
```

Crée une variable `full_name` en concaténant `first_name` et `last_name`. Puis affiche `full_name`.

# --hints--

Tu devrais avoir une variable nommée `full_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("full_name")`)) })
```

Tu devrais assigner à ta variable `full_name` une chaîne formée en concaténant `first_name` et `last_name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("full_name").is_equivalent("full_name = first_name + last_name")`)) })
```

Tu devrais afficher la variable `full_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(full_name)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
--fcc-editable-region--

--fcc-editable-region--
```
