---
id: 694bf6eb530e19a9c48c59f7
title: Étape 18
challengeType: 20
dashedName: step-18
---

# --description--

Tu peux aussi utiliser des nombres négatifs pour découper depuis la fin d'une chaîne. Par exemple, `-1` fait référence au dernier caractère, `-2` à l'avant-dernier, et ainsi de suite.

Pour obtenir les trois derniers caractères d'une chaîne, tu peux utiliser `string[-3:]`. Note comment le paramètre stop est omis après les deux-points. Cela signifie que le découpage doit aller jusqu'à la limite de la chaîne.

Crée une variable nommée `last_three`. Utilise l'indexation négative pour extraire les trois derniers caractères de `employee_code` (qui représentent le numéro d'identification). Enfin, affiche `last_three` dans le terminal.

Avec cela, le workshop de la carte employé est terminé.

# --hints--

Tu devrais avoir une variable nommée `last_three`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("last_three")`)) })
```

Tu devrais découper les trois derniers caractères de `employee_code` avec l'indexation négative et les assigner à ta variable `last_three`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("last_three").is_equivalent("last_three = employee_code[-3:]")`)) })
```

Tu devrais afficher `last_three` dans le terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(last_three)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department)
year_code = employee_code[4:8]
print(year_code)
initials = employee_code[9:11]
print(initials)
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department)
year_code = employee_code[4:8]
initials = employee_code[9:11]
print(year_code)
print(initials)
last_three = employee_code[-3:]
print(last_three)
```
