---
id: 694bf6eb530e19a9c48c59f5
title: Étape 16
challengeType: 20
dashedName: step-16
---

# --description--

Quand tu travailles avec des chaînes, tu auras souvent besoin d'extraire une portion spécifique d'une chaîne. C'est ce qu'on appelle le **découpage** (slicing).

La syntaxe est `string[start:stop]`, où :

* `start` est l'indice où le découpage commence (**inclusif**).
* `stop` est l'indice où le découpage se termine (**exclusif**).

Par exemple, si `text = 'Python'`, alors `text[0:2]` renvoie `'Py'`.

Définis `employee_code` comme `'DEV-2026-JD-001'`. Ensuite, crée une variable `department` et assigne-lui le découpage de `employee_code` de l'indice `0` à `3`. Puis affiche `department` dans le terminal.

# --hints--
Tu devrais avoir une variable nommée `employee_code`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_code")`)) })
```

Tu devrais assigner la chaîne `DEV-2026-JD-001` à ta variable `employee_code`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_code").is_equivalent("employee_code = 'DEV-2026-JD-001'")`)) })
```

Tu devrais avoir une variable nommée `department`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("department")`)) })
```

Tu devrais découper les trois premiers caractères de `employee_code` et les assigner à ta variable `department`.

```js
({ test: () => runPython(`
dep = _Node(_code).find_variable("department")
assert dep.is_equivalent("department = employee_code[0:3]") or dep.is_equivalent("department = employee_code[:3]")
`) })
```

Tu devrais afficher `department` dans le terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(department)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)

--fcc-editable-region--

--fcc-editable-region--
```
