---
id: 694bf6eb530e19a9c48c59f2
title: Étape 11
challengeType: 20
dashedName: step-11
---

# --description--

Termine maintenant la phrase en mettant à jour l'affectation de `employee_info` pour concaténer aussi la chaîne ` years old` à la fin. N'oublie pas d'inclure un espace au début de ta chaîne.

Enfin, affiche `employee_info`.

# --hints--

Tu devrais concaténer `' years old'` à la fin de la variable `employee_info`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is ' + str(employee_age) + ' years old'")`)) })
```

Tu devrais afficher `employee_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(employee_info)")`)) })
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is ' + str(employee_age)

--fcc-editable-region--
```
