---
id: 694bf6eb530e19a9c48c59f1
title: Étape 5
challengeType: 20
dashedName: step-5
---

# --description--

Maintenant, ton adresse semble incomplète. Tu veux aussi ajouter le numéro d'appartement où habite l'employé, tu dois donc modifier la variable.

Quand tu veux ajouter du contenu à la fin d'une variable de chaîne existante, tu peux utiliser l'opérateur d'**affectation augmentée** `+=`. C'est plus court que d'écrire `var = var + 'new text'`. Par exemple :

```py
greeting = 'Hello'
greeting += ' World'
print(greeting) # Hello World
```

Souviens-toi que les chaînes sont immuables : cette opération ne modifie donc pas la chaîne d'origine. Elle crée plutôt une nouvelle chaîne et la réassigne à la variable.

Utilise l'opérateur `+=` pour ajouter la chaîne `, Apartment 4B` à ta variable `address`.

# --hints--

Tu devrais utiliser l'opérateur `+=` pour ajouter la chaîne `, Apartment 4B` à la variable `address`.

```js
({ test: () => assert(runPython(`_Node(_code).find_aug_variable("address").is_equivalent("address += ', Apartment 4B'")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
full_name = first_name + ' ' + last_name
print(full_name)
address = '123 Main Street'
--fcc-editable-region--

--fcc-editable-region--
print(address)
```
