---
id: 694bf6eb530e19a9c48c59f3
title: Étape 12
challengeType: 20
dashedName: step-12
---

# --description--

Tu vas maintenant utiliser la fonction `str()` une fois de plus. Comme pour l'âge, tu dois convertir toute variable numérique en chaîne avant de la concaténer avec d'autres textes.

Crée une variable nommée `experience_years` et assigne-lui l'entier `5`.

Puis, crée une variable `experience_info`. Assigne-lui une chaîne formée en concaténant `'Experience: '`, la variable `experience_years` (convertie en chaîne) et `' years'`. Affiche le résultat dans le terminal.

# --hints--

Tu devrais avoir une variable nommée `experience_years`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("experience_years")`)) })
```

Tu devrais assigner l'entier `5` à ta variable `experience_years`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("experience_years").is_equivalent("experience_years = 5")`)) })
```

Tu devrais avoir une variable nommée `experience_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("experience_info")`)) })
```

Tu devrais assigner à ta variable `experience_info` une chaîne formée en concaténant `'Experience: '`, `str(experience_years)` et `' years'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("experience_info").is_equivalent("experience_info = 'Experience: ' + str(experience_years) + ' years'")`)) })
```

Tu devrais afficher `experience_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(experience_info)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
--fcc-editable-region--

--fcc-editable-region--
```
