---
id: 695c9f21a8b1d53cf1b97841
title: Étape 7
challengeType: 20
dashedName: step-7
---

# --description--

Crée maintenant une variable nommée `employee_age` et assigne-lui l'entier `28`.


# --hints--

Tu devrais avoir une variable nommée `employee_age`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_age")`)) })
```

Tu devrais assigner l'entier `28` à ta variable `employee_age`. N'entoure pas le texte de quotes.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_age").is_equivalent("employee_age = 28")`)) })
```


# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
--fcc-editable-region--

--fcc-editable-region--
```
