---
id: 695c9c7c61d446f1451e0ae7
title: Étape 4
challengeType: 20
dashedName: step-4
---

# --description--

Ensuite, crée une variable `address` pour stocker l'adresse de l'employé. Assigne-lui la chaîne `123 Main Street`, puis affiche `address`.

# --hints--

Tu devrais avoir une variable nommée `address`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("address")`)) })
```

Tu devrais assigner la chaîne `123 Main Street` à ta variable `address`. N'oublie pas d'entourer le texte de quotes simples ou doubles.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("address").is_equivalent("address = '123 Main Street'")`)) })
```

Tu devrais afficher la variable `address`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(address)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
full_name = first_name + ' ' + last_name
print(full_name)
--fcc-editable-region--

--fcc-editable-region--
```
