---
id: 695c96f82af4fb78e59fa9ed
title: Étape 1
challengeType: 20
dashedName: step-1
---

# --description--

Les chaînes sont des séquences de caractères utilisées pour stocker du texte. Comme tu t'en souviens des leçons précédentes, tu peux créer une chaîne en entourant le texte de quotes simples (`'`) ou doubles (`"`). Par exemple :

```python
greeting = 'Hello'
print(greeting) # Output: Hello
```

Crée une variable `first_name` qui stocke la chaîne `John` et une variable `last_name` qui stocke la chaîne `Doe`. Puis affiche `first_name` et `last_name`.

# --hints--

Tu devrais avoir une variable nommée `first_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("first_name")`)) })
```

Tu devrais assigner la chaîne `John` à ta variable `first_name`. N'oublie pas d'entourer le texte de quotes simples ou doubles.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("first_name").is_equivalent("first_name = 'John'")`)) })
```

Tu devrais avoir une variable nommée `last_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("last_name")`)) })
```

Tu devrais assigner la chaîne `Doe` à ta variable `last_name`. N'oublie pas d'entourer le texte de quotes simples ou doubles.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("last_name").is_equivalent("last_name = 'Doe'")`)) })
```

Tu devrais afficher ta variable `first_name` dans le terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(first_name)")`)) })
```

Tu devrais afficher ta variable `last_name` dans le terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(last_name)")`)) })
```

# --seed--

## --seed-contents--

```python
--fcc-editable-region--

--fcc-editable-region--
```
