---
id: 695df6480de23758af5a29cd
title: Étape 8
challengeType: 20
dashedName: step-8
---

# --description--

Maintenant, tu veux créer une chaîne qui affiche l'âge de l'employé.

Commence par créer une variable `employee_info` et assigne-lui le résultat de la concaténation de :

- la variable `full_name`.
- une chaîne composée des caractères `is` précédés et suivis d'un espace.

# --hints--

Tu devrais avoir une variable nommée `employee_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_info")`)) })
```

Tu devrais assigner `full_name + ' is '` à ta variable `employee_info`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is '")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--

--fcc-editable-region--
```
