---
id: 695ea1e55c0e619e597fc5c9
title: Étape 10
challengeType: 20
dashedName: step-10
---

# --description--

Comme tu peux le voir, Python a levé une `TypeError: can only concatenate str (not "int") to str`. Cela arrive parce que Python ne te permet pas de concaténer directement du texte (des chaînes) et des nombres (des entiers).

Pour corriger cela, tu dois d'abord convertir le nombre en chaîne avec la fonction `str()`, qui renvoie la version chaîne d'un objet :

```py
my_num = str(42)
print(type(my_num)) # <class 'str'>
```

Mets à jour l'affectation de `employee_info` pour convertir `employee_age` en chaîne avec `str(employee_age)`.

# --hints--

Tu devrais modifier l'affectation de ta variable `employee_info` en convertissant `employee_age` en chaîne avec `str(employee_age)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is ' + str(employee_age)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is ' + employee_age
--fcc-editable-region--
```
