---
id: 695fcfbba648b3a96a9495a4
title: Étape 6
challengeType: 20
dashedName: step-6
---

# --description--

La sortie dans le terminal devient encombrée. Avant de continuer, il est temps de la nettoyer.

Supprime toutes les instructions `print()` de ton code.

# --hints--

Tu ne devrais pas avoir d'appels `print` dans ton code.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_calls("print")) == 0`)) })
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
--fcc-editable-region--
print(first_name)
print(last_name)
full_name = first_name + ' ' + last_name
print(full_name)
address = '123 Main Street'
address += ', Apartment 4B'
print(address)
--fcc-editable-region--
```
