---
id: 695fd3dde03273875e5d6709
title: Étape 9
challengeType: 20
dashedName: step-9
---

# --description--

Mets à jour l'affectation de `employee_info` pour concaténer aussi `employee_age` à la fin.

Une fois que c'est fait, tu verras une `TypeError` dans le terminal. À l'étape suivante, tu travailleras à la corriger.

# --hints--

Tu devrais mettre à jour l'affectation de `employee_info` pour concaténer `employee_age` à la fin.

```js
const regex = /^employee_info\s*=\s*full_name\s*\+\s*(["']) is \1\s*\+\s*employee_age\s*$/m;
assert.match(__helpers.python.removeComments(code), regex);
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is '
--fcc-editable-region--
```
