---
id: 69727ea5c9606bedfe36d3d2
title: Étape 15
challengeType: 20
dashedName: step-15
---

# --description--

Il est maintenant temps d'ajouter les derniers détails à la carte.

Crée une variable nommée `position` avec la valeur de la chaîne `Data Analyst` et une variable nommée `salary` avec la valeur de l'entier `75000`.

Puis, mets à jour ta f-string `employee_card` pour inclure le poste et le salaire. Elle doit suivre ce format exact : `Employee: [full_name] | Age: [employee_age] | Position: [position] | Salary: $[salary]`. Remplace les espaces réservés par les variables correspondantes.

Enfin, affiche `employee_card` pour voir le résultat.

# --hints--

Tu devrais avoir une variable nommée `position`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("position")`)) })
```

Tu devrais assigner la chaîne `Data Analyst` à ta variable `position`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("position").is_equivalent("position = 'Data Analyst'")`)) })
```

Tu devrais avoir une variable nommée `salary`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("salary")`)) })
```

Tu devrais assigner l'entier `75000` à ta variable `salary`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("salary").is_equivalent("salary = 75000")`)) })
```

Tu devrais mettre à jour `employee_card` pour utiliser la f-string `f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: \${salary}'")`)) })
```

Tu devrais afficher `employee_card` dans le terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(employee_card)")`))})
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--

employee_card = f'Employee: {full_name} | Age: {employee_age}'
--fcc-editable-region--
```
