---
id: 6980d3ecfe800d8a54130f30
title: Étape 14
challengeType: 20
dashedName: step-14
---

# --description--

Actuellement, `employee_card` n'affiche que le nom de l'employé. Tu vas maintenant y ajouter plus d'informations.

Mets à jour l'affectation de `employee_card` pour inclure l'âge de l'employé. La chaîne finale devrait ressembler à ceci : `Employee: [name] | Age: [age]`, avec `[name]` remplacé par le nom de l'employé et `[age]` par son âge.

# --hints--

Tu devrais mettre à jour `employee_card` pour utiliser la f-string `f'Employee: {full_name} | Age: {employee_age}'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name} | Age: {employee_age}'")`)) })
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--
employee_card = f'Employee: {full_name}'
--fcc-editable-region--
```
