---
id: 689f7996426c1534fa5bea66
title: Crée un traducteur de score de golf
challengeType: 26
dashedName: build-a-golf-score-translator
---

# --description--

Au golf, chaque trou possède un `par`, c'est-à-dire le nombre moyen de `strokes` (coups) qu'un golfeur est censé jouer pour faire entrer la balle dans le trou et terminer le jeu. Selon que tes `strokes` sont au-dessus ou en dessous du `par`, il existe un surnom différent.

Dans ce lab, tu vas écrire une fonction qui convertit le `par` et les `strokes` en leur surnom.

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois créer une fonction nommée `golfScore`.
1. `golfScore` doit prendre deux arguments numériques, qui sont le par du parcours et le nombre de coups joués.
1. `golfScore` doit retourner une chaîne.
1. `golfScore` doit retourner `"Hole-in-one!"` si `strokes` vaut `1`.
1. `golfScore` doit retourner `"Eagle"` si `strokes` est inférieur ou égal à `par` moins `2`.
1. `golfScore` doit retourner `"Birdie"` si `strokes` est égal à `par` moins `1`.
1. `golfScore` doit retourner `"Par"` si `strokes` est égal à `par`.
1. `golfScore` doit retourner `"Bogey"` si `strokes` est égal à `par` plus `1`.
1. `golfScore` doit retourner `"Double Bogey"` si `strokes` est égal à `par` plus `2`.
1. `golfScore` doit retourner `"Go Home!"` si `strokes` est supérieur ou égal à `par` plus `3`.

# --hints--

Tu dois créer une fonction nommée `golfScore`.

```js
assert.isFunction(golfScore);
```

`golfScore` doit prendre deux paramètres.

```js
assert.lengthOf(golfScore, 2);
```

`golfScore` doit retourner une chaîne

```js
assert.isString(golfScore(4,1));
```

`golfScore(1, 1)` doit retourner la chaîne `Hole-in-one!`

```js
assert.strictEqual(golfScore(1,1), 'Hole-in-one!');
```

`golfScore(3, 1)` doit retourner la chaîne `Hole-in-one!`

```js
assert.strictEqual(golfScore(3,1), 'Hole-in-one!');
```

`golfScore(4, 1)` doit retourner la chaîne `Hole-in-one!`

```js
assert.strictEqual(golfScore(4,1), 'Hole-in-one!');
```

`golfScore(5, 1)` doit retourner la chaîne `Hole-in-one!`

```js
assert.strictEqual(golfScore(5,1), 'Hole-in-one!');
```

`golfScore(4, 2)` doit retourner la chaîne `Eagle`

```js
assert.strictEqual(golfScore(4,2), 'Eagle');
```

`golfScore(5, 2)` doit retourner la chaîne `Eagle`

```js
assert.strictEqual(golfScore(5,2), 'Eagle');
```

`golfScore(3, 2)` doit retourner la chaîne `Birdie`

```js
assert.strictEqual(golfScore(3,2), 'Birdie');
```

`golfScore(4, 3)` doit retourner la chaîne `Birdie`

```js
assert.strictEqual(golfScore(4,3), 'Birdie');
```

`golfScore(5, 4)` doit retourner la chaîne `Birdie`

```js
assert.strictEqual(golfScore(5,4), 'Birdie');
```

`golfScore(3, 3)` doit retourner la chaîne `Par`

```js
assert.strictEqual(golfScore(3,3), 'Par');
```

`golfScore(4, 4)` doit retourner la chaîne `Par`

```js
assert.strictEqual(golfScore(4,4), 'Par');
```

`golfScore(5, 5)` doit retourner la chaîne `Par`

```js
assert.strictEqual(golfScore(5,5), 'Par');
```

`golfScore(3, 4)` doit retourner la chaîne `Bogey`

```js
assert.strictEqual(golfScore(3,4), 'Bogey');
```

`golfScore(4, 5)` doit retourner la chaîne `Bogey`

```js
assert.strictEqual(golfScore(4,5), 'Bogey');
```

`golfScore(5, 6)` doit retourner la chaîne `Bogey`

```js
assert.strictEqual(golfScore(5,6), 'Bogey');
```

`golfScore(3, 5)` doit retourner la chaîne `Double Bogey`

```js
assert.strictEqual(golfScore(3,5), 'Double Bogey');
```

`golfScore(4, 6)` doit retourner la chaîne `Double Bogey`

```js
assert.strictEqual(golfScore(4,6), 'Double Bogey');
```

`golfScore(5, 7)` doit retourner la chaîne `Double Bogey`

```js
assert.strictEqual(golfScore(5,7), 'Double Bogey');
```

`golfScore(3, 7)` doit retourner la chaîne `Go Home!`

```js
assert.strictEqual(golfScore(3,7), 'Go Home!');
```

`golfScore(4, 8)` doit retourner la chaîne `Go Home!`

```js
assert.strictEqual(golfScore(4,8), 'Go Home!');
```

`golfScore(5, 9)` doit retourner la chaîne `Go Home!`

```js
assert.strictEqual(golfScore(5,9), 'Go Home!');
```

# --seed--

## --seed-contents--

```js
const names = ["Hole-in-one!", "Eagle", "Birdie", "Par", "Bogey", "Double Bogey", "Go Home!"];

```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) {
    return "Hole-in-one!";
  }

  if (strokes <= par - 2) {
    return "Eagle";
  }

  if (strokes === par - 1) {
    return "Birdie";
  }

  if (strokes === par) {
    return "Par";
  }

  if (strokes === par + 1) {
    return "Bogey";
  }

  if(strokes === par + 2) {
    return "Double Bogey";
  }

  return "Go Home!";
}
```
