---
id: 66f17db06803d11a1bd19a20
title: Crée un jeu de quiz
challengeType: 26
dashedName: lab-quiz-game
---

# --description--

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois créer un tableau nommé `questions`.
1. Le tableau `questions` doit contenir au moins cinq objets, chacun ayant les clés `category`, `question`, `choices` et `answer`.
1. La clé `category` doit avoir pour valeur une chaîne représentant une catégorie de question.
1. La clé `question` doit avoir pour valeur une chaîne représentant une question.
1. La clé `choices` doit avoir pour valeur un tableau contenant trois chaînes, qui sont des réponses alternatives à la question.
1. La clé `answer` doit avoir pour valeur une chaîne, représentant la bonne réponse à la question. De plus, la valeur de `answer` doit être incluse dans le tableau `choices`.
1. Tu dois avoir une fonction nommée `getRandomQuestion` qui prend un tableau de questions en paramètre et retourne un objet question aléatoire du tableau.
1. Tu dois avoir une fonction nommée `getRandomComputerChoice` qui prend le tableau des choix disponibles en paramètre, et retourne une réponse aléatoire à la question sélectionnée.
1. Tu dois avoir une fonction nommée `getResults` qui prend l'objet question comme premier paramètre et le choix de l'ordinateur comme second paramètre. La fonction doit retourner `The computer's choice is correct!` si la réponse est correcte. Sinon, elle retourne `The computer's choice is wrong. The correct answer is: <correct-answer>`, où `<correct-answer>` est la valeur de la bonne réponse à la question choisie.

# --hints--

Tu dois créer un tableau nommé `questions`.

```js
assert.isArray(questions);
```

Le tableau `questions` doit contenir au moins cinq objets, chacun ayant les clés `category`, `question`, `choices` et `answer`.

```js
assert.isAtLeast(questions.length, 5);
questions.forEach(question => {
    assert.isObject(question);
    assert.hasAllKeys(question, ['category', 'choices', 'question', 'answer']);
});

```

La clé `category` doit avoir pour valeur une chaîne représentant une catégorie de question.

```js
assert.isNotEmpty(questions);
questions.forEach(({category}) => {
    assert.isString(category);
    assert.isNotEmpty(category);
});
```

La clé `question` doit avoir pour valeur une chaîne représentant une question.

```js
assert.isNotEmpty(questions);
questions.forEach(({question}) => {
    assert.isString(question);
    assert.include(question, '?');
    assert.isNotEmpty(question.replaceAll('?', ''));
})
```

La clé `choices` doit avoir pour valeur un tableau contenant trois chaînes différentes les unes des autres.

```js
assert.isNotEmpty(questions);
questions.forEach(({choices}) => {
    assert.isArray(choices);
    assert.lengthOf(choices, 3);
    assert.lengthOf(new Set(choices), 3);
    choices.forEach(choice => {
        assert.isString(choice);
        assert.isNotEmpty(choice);
    });
})
```

La clé `answer` doit avoir pour valeur une chaîne.

```js
assert.isNotEmpty(questions);
questions.forEach(({answer}) => {
    assert.isString(answer);
    assert.isNotEmpty(answer);
});
```

La valeur de `answer` doit être incluse dans le tableau `choices`.

```js
assert.isNotEmpty(questions);
questions.forEach(({answer, choices}) => assert.oneOf(answer, choices));
```

Tu dois avoir une fonction nommée `getRandomQuestion` qui prend un tableau de questions en paramètre et retourne un objet question aléatoire du tableau.

```js
const parameters = __helpers.getFunctionParams(getRandomQuestion.toString());

assert.isFunction(getRandomQuestion);
assert.lengthOf(parameters, 1);
assert.lengthOf(getRandomQuestion, 1);
assert.isObject(getRandomQuestion(questions));
assert.deepInclude(questions, getRandomQuestion(questions), "getRandomQuestion did not return one of the objects inside questions");
```

Tu dois avoir une fonction nommée `getRandomComputerChoice` qui prend le tableau des choix disponibles en paramètre, et retourne une réponse aléatoire à la question sélectionnée.

```js
assert.isFunction(getRandomComputerChoice);
const randomArray = ['a', 'b', 'c'];
assert.oneOf(getRandomComputerChoice(randomArray), randomArray)
```

Tu dois avoir une fonction nommée `getResults`.

```js
assert.isFunction(getResults);
```

Ta fonction `getResults` doit prendre l'objet question comme premier paramètre et le choix de l'ordinateur comme second paramètre.

```js
const testQuestion = questions[0];
const wrongChoice = questions[0].choices.find(choice => choice !== questions[0].answer);
assert.equal(getResults(testQuestion, testQuestion.answer), `The computer's choice is correct!`);
assert.equal(getResults(testQuestion, wrongChoice), `The computer's choice is wrong. The correct answer is: ${testQuestion.answer}`);
```

Si le choix de l'ordinateur correspond à la réponse, `getResults` doit retourner `The computer's choice is correct!`

```js
assert.equal(getResults(questions[0], questions[0].answer), `The computer's choice is correct!`);
```

Si le choix de l'ordinateur ne correspond pas à la réponse, `getResults` doit retourner `The computer's choice is wrong. The correct answer is: <correct-answer>`, où `<correct-answer>` est la valeur de la bonne réponse à la question choisie.

```js
assert.equal(getResults({category: 'misc', choices: ['a', 'b', 'c'], question: "question?", answer: "b"}, "a"), `The computer's choice is wrong. The correct answer is: b`)
```

Ta fonction `getResults` doit utiliser une comparaison d'égalité exacte, pas une correspondance de sous-chaîne.

```js
assert.equal(getResults({category: 'food', choices: ['Ham', 'Hamburger', 'Hot Dog'], question: "What food?", answer: "Hamburger"}, "Ham"), `The computer's choice is wrong. The correct answer is: Hamburger`);
assert.equal(getResults({category: 'food', choices: ['Ham', 'Hamburger', 'Hot Dog'], question: "What food?", answer: "Ham"}, "Hamburger"), `The computer's choice is wrong. The correct answer is: Ham`);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const questions = [
    {
        category: "science",
        question: "What is the chemical symbol for potassium?",
        choices: ["P", "K", "Pt"],
        answer: "K"
    },
    {
        category: "science",
        question: "What is the unit of electrical resistance?",
        choices: ["Ohm", "Coulomb", "Sievert"],
        answer: "Ohm"
    },
    {
        category: "geography",
        question: "What is the capital city of Australia?",
        choices: ["Sidney", "Canberra", "Wellington"],
        answer: "Canberra"
    },
    {
        category: "literature",
        question: 'Who wrote "1984"?',
        choices: ["Ray Bradbury", "Aldous Huxley", "George Orwell"],
        answer: "George Orwell"
    },
    {
        category: "sport",
        question: "How many players are on a standard volleyball team?",
        choices: ["6", "7", "12"],
        answer: "6"
    }
]

function getRandomQuestion(questionList) {
    const randomIndex = Math.floor(Math.random() * questionList.length);
    return questionList[randomIndex];
}

function getRandomComputerChoice(choices) {
    const randomIndex = Math.floor(Math.random() * choices.length);
    return choices[randomIndex];
}

function getResults(obj, computerChoice) {
    return computerChoice === obj.answer
        ? "The computer's choice is correct!"
        : `The computer's choice is wrong. The correct answer is: ${obj.answer}`;
}

const questionObj = getRandomQuestion(questions);
const { question, choices } = questionObj;

console.log(question);
console.log(`Choices: ${choices}`);

const computerChoice = getRandomComputerChoice(choices);
console.log(`Computer chooses: ${computerChoice}`);

const results = getResults(questionObj, computerChoice);
console.log(results);

```
