---
id: 66f2836c459cfb16ae76f24f
title: Crée un générateur de pyramide
challengeType: 26
dashedName: lab-pyramid-generator
---

# --description--

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois avoir une fonction nommée `pyramid` qui prend trois arguments.
1. Le premier argument doit être une chaîne représentant le caractère de motif à répéter dans ta pyramide.
1. Le second argument doit être un entier représentant le nombre de lignes de la pyramide.
1. Le troisième argument doit être une valeur booléenne.
1. La fonction `pyramid` doit renvoyer une chaîne dans laquelle le caractère de motif est répété et disposé pour former une pyramide dont le sommet est orienté vers le haut lorsque le troisième argument est `false`.
1. Lorsque le troisième argument est `true`, la pyramide doit avoir le sommet orienté vers le bas.
1. La ligne du sommet doit avoir un seul caractère de motif, et chaque autre ligne doit avoir deux caractères de motif de plus que la précédente.
1. Chaque ligne doit commencer par un nombre d'espaces suffisant pour placer le caractère central de chaque ligne dans la même colonne, et il ne doit pas y avoir d'espaces à la fin de chaque ligne.
1. La pyramide doit commencer et se terminer par un caractère de nouvelle ligne.

Par exemple, appeler `pyramid("o", 4, false)` doit donner cette sortie :

```js

   o
  ooo
 ooooo
ooooooo

```

# --hints--

Tu dois avoir une fonction nommée `pyramid`.

```js
assert.isFunction(pyramid);
```

Ta fonction `pyramid` doit avoir trois paramètres.

```js
assert.lengthOf(pyramid, 3);
```

`pyramid("o", 4, false)` doit renvoyer `"\n   o\n  ooo\n ooooo\nooooooo\n"`.

```js
assert.equal(pyramid("o", 4, false), "\n   o\n  ooo\n ooooo\nooooooo\n")
```

`pyramid("p", 5, true)` doit renvoyer `"\nppppppppp\n ppppppp\n  ppppp\n   ppp\n    p\n"`.

```js
assert.equal(pyramid("p", 5, true), "\nppppppppp\n ppppppp\n  ppppp\n   ppp\n    p\n")
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function pyramid(char, count, isInverted) {
    const rows = []
    for (let i = 1; i <= count; i++) {
        if (isInverted) {
            rows.unshift(" ".repeat(count - i) + char.repeat(2 * i - 1))
        } else {
            rows.push(" ".repeat(count - i) + char.repeat(2 * i - 1))
        }
    }
    return "\n" + rows.join("\n") + "\n";
}

console.log(pyramid("#", 10, false))

```
