---
id: 6823d9ac8bdc3853df65a1ff
title: Étape 7
challengeType: 0
dashedName: step-7
---

# --description--

Il existe six éléments de titre en HTML : de `h1` à `h6`. Ils servent à montrer l'importance des sections de ta page web, `h1` étant le plus important et `h6` le moins important.

Sous ton élément `p`, ajoute un titre `h3` avec le texte :

```md
Introduction to HTML
```

# --hints--

Ton élément `h3` doit avoir une balise ouvrante `<h3>`.

```js
assert.exists(document.querySelector("h3"));
```

Ton élément `h3` doit avoir une balise fermante `</h3>`.

```js
assert.match(code, /<\/h3\s*\>/);
```

Le texte de ton élément `h3` doit être `Introduction to HTML`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*HTML\s*\<\/h3\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

--fcc-editable-region--

--fcc-editable-region--
```
