---
id: 6823ac607bfdbc46331b2559
title: Étape 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

HTML signifie HyperText Markup Language. C'est le code qui définit la structure et le contenu d'une page web. Voici ton éditeur de code : c'est ici que tu vas écrire du HTML.

Trouve la ligne 1 dans l'éditeur et tape ce texte :

```md
Welcome to freeCodeCamp
```

Quand tu as terminé, clique sur le bouton « Vérifier votre code » pour voir si c'est correct.

# --hints--

Tu dois avoir le texte `Welcome to freeCodeCamp` dans ton éditeur. Vérifie bien l'orthographe.

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /welcome\s*to\s*freecodecamp/i)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
