---
id: 6823e637c1c0ed56f781b4fc
title: Étape 11
challengeType: 0
dashedName: step-11
---

# --description--

Enfin, JavaScript rend ta page web interactive : il te permet de dire à la page quoi faire quand quelqu'un clique sur un bouton, envoie un formulaire, et bien d'autres choses.

Pour la dernière étape de l'atelier, ajoute un autre élément `h3` et un autre élément `p` à la page pour décrire JavaScript.

D'abord, ajoute un élément `h3` avec le texte :

```md
Introduction to JavaScript
```

Puis, sous cet élément `h3`, ajoute un élément `p` avec le texte suivant :

```md
JavaScript adds interactivity to a webpage
```

# --hints--

Tu dois ajouter un troisième élément `h3` à la page.

```js
assert.lengthOf(document.querySelectorAll("h3"), 3);
```

Le texte de ton élément `h3` doit être `Introduction to JavaScript`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*JavaScript\s*\<\/h3\s*\>/i);
```

Tu dois ajouter un quatrième élément `p` à la page.

```js
assert.lengthOf(document.querySelectorAll("p"), 4);
```

Ton élément `p` doit avoir le texte `JavaScript adds interactivity to a webpage`

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*JavaScript\s*adds\s*interactivity\s*to\s*a\s*web\s*page\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

<h3>Introduction to JavaScript</h3>
<p>JavaScript adds interactivity to a webpage</p>
```
