---
id: 6823e036ea4b71553558c01b
title: Étape 8
challengeType: 0
dashedName: step-8
---

# --description--

Remarque que chaque titre a une apparence légèrement différente dans l'aperçu. C'est pareil pour le paragraphe.

Ensuite, ajoute un autre élément `p` qui affiche le texte suivant :

```md
HTML represents the content and structure of a webpage
```

# --hints--

Tu dois avoir un deuxième élément `p`.

```js
assert.lengthOf(document.querySelectorAll("p"), 2);
```

Ton élément `p` doit avoir le texte `HTML represents the content and structure of a webpage`. Tu peux copier le texte depuis les instructions.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*HTML\s*represents\s*the\s*content\s*and\s*structure\s*of\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
--fcc-editable-region--

--fcc-editable-region--
```
