---
id: 682ba2318000b62f179bdf04
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

Le HTML est composé d'éléments. Le premier que tu vas utiliser est l'élément `h1` :

```html
<h1>Welcome to freeCodeCamp</h1>
```

Il commence par une balise ouvrante (`<h1>`), se termine par une balise fermante (`</h1>`) et contient entre les deux le texte qui sera affiché.

Transforme ton texte `Welcome to freeCodeCamp` en élément `h1` en ajoutant une balise ouvrante avant le texte et une balise fermante après.

# --hints--

Ton élément `h1` doit avoir une balise ouvrante `<h1>`.

```js
assert.exists(document.querySelector("h1"));
```

Ton élément `h1` doit avoir une balise fermante `</h1>`.

```js
assert.match(code, /<\/h1\s*\>/);
```

Ton élément `h1` doit ressembler à ceci : `<h1>Welcome to freeCodeCamp</h1>`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h1\s*\>\s*Welcome\s*to\s*freeCodeCamp\s*\<\/h1\s*\>/i);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
Welcome to freeCodeCamp
--fcc-editable-region--
```
