---
id: 682cd20b883fc7b25eb539c6
title: Étape 5
challengeType: 0
dashedName: step-5
---

# --description--

Sous les deux autres lignes de texte, ajoute :

```md
Learn the skills to become a full-stack developer
```

# --hints--

Tu dois avoir le texte `Learn the skills to become a full-stack developer` dans ton éditeur. Vérifie bien l'orthographe.

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /learn\s*the\s*skills\s*to\s*become\s*a\s*full-stack\s*developer/i)
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
--fcc-editable-region--

--fcc-editable-region--
```
