---
id: 682cd206883fc7b25eb539c5
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

Remarque que le HTML que tu écris dans l'éditeur apparaît dans l'aperçu. Dans cet atelier, tu vas écrire le HTML d'une page partielle de curriculum.

Sous ton élément `h1`, tape le texte suivant sur la ligne vide :

```md
Full-Stack Curriculum
```

# --hints--

Tu dois avoir le texte `Full-Stack Curriculum` dans ton éditeur. Vérifie bien l'orthographe.

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /full-stack\s*curriculum/i)
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
--fcc-editable-region--

--fcc-editable-region--
```
