---
id: 683921c4769fd23dbadec2fe
title: Étape 10
challengeType: 0
dashedName: step-10
---

# --description--

Alors que HTML définit la structure et le contenu d'une page web, CSS sert à ajouter du style : couleurs, polices, espacements, mise en page, etc.

Sous le `h3` que tu viens d'ajouter, ajoute un autre élément paragraphe avec le texte :

```md
CSS is used to style a webpage
```

# --hints--

Tu dois ajouter un deuxième élément `h3` à la page.

```js
assert.lengthOf(document.querySelectorAll("h3"), 2);
```

Le texte de ton élément `h3` doit être `Introduction to CSS`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*CSS\s*\<\/h3\s*\>/i);
```

Tu dois ajouter un troisième élément `p` à la page.

```js
assert.lengthOf(document.querySelectorAll("p"), 3);
```

Ton élément `p` doit avoir le texte `CSS is used to style a webpage`. Tu peux copier le texte depuis les instructions.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*CSS\s*is\s*used\s*to\s*style\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

--fcc-editable-region--
<h3>Introduction to CSS</h3>

--fcc-editable-region--
```
