---
id: 641da5462576784453146ec2
title: Étape 7
challengeType: 0
dashedName: step-7
---

# --description--

Maintenant que tu as les données que tu veux, tu peux les utiliser pour peupler l'interface.

Mais les données que tu récupères auront un total de 26 auteurs. Au lieu d'afficher les 26 auteurs d'un coup, tu dois en afficher 8 à la fois.

Commence par utiliser le mot-clé `let` pour créer deux variables nommées `startingIndex` et `endingIndex`, et assigne-leur les nombres `0` et `8`, respectivement. Tu auras besoin d'utiliser `let` ici car tu réassigneras ces valeurs plus tard.

Puis, crée une variable `authorDataArr` et assigne-lui un tableau vide.

Déclare et initialise chaque variable sur une ligne séparée.

# --before-all--

```js
window.fetch = () => Promise.resolve({json: () => Promise.resolve([{ author: 'Whoever', image: 'data:image/gif;base64,R0lGODlhAQABAIAAAAAAAP///ywAAAAAAQABAAACAUwAOw==', url: "http://not-a-real-url.nowhere/", bio: 'words go here' }])});
```

# --hints--

Tu dois utiliser `let` pour déclarer une variable nommée `startingIndex`.

```js
assert.match(code, /let\s+startingIndex/);
```

Ta variable `startingIndex` doit être un nombre.

```js
assert.isNumber(startingIndex);
```

Ta variable `startingIndex` doit être définie à `0`.

```js
assert.strictEqual(startingIndex, 0);
```

Tu dois utiliser `let` pour déclarer une variable nommée `endingIndex`.

```js
assert.match(code, /let\s+endingIndex/);
```

Ta variable `endingIndex` doit être un nombre.
  
  ```js
assert.isNumber(endingIndex);
```

Tu dois définir ta variable `endingIndex` à `8`.

```js
assert.strictEqual(endingIndex, 8)
```

Tu dois avoir une variable nommée `authorDataArr`.

```js
assert.isDefined(authorDataArr);
```

Ta variable `authorDataArr` doit être un tableau vide.

```js
assert.isArray(authorDataArr);
assert.isEmpty(authorDataArr);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">freeCodeCamp News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
const authorContainer = document.getElementById('author-container');
const loadMoreBtn = document.getElementById('load-more-btn');

--fcc-editable-region--

--fcc-editable-region--

fetch('https://cdn.freecodecamp.org/curriculum/news-author-page/authors.json')
  .then((res) => res.json())
  .then((data) => {
    console.log(data);
  })
  .catch((err) => {
    console.error(`There was an error: ${err}`);
  });
```
