---
id: 69baafd6d683dd9921b75db3
title: Étape 1
challengeType: 1
dashedName: step-1
---

# --description--

Dans ce workshop, tu vas utiliser JavaScript pour simuler le flux des participants à un festival de musique.

Tu travailleras avec deux jeux de données qui décrivent les portes du festival pendant les créneaux du matin et de la nuit :

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];
```

Chaque objet porte contient les propriétés suivantes :

- `id` : une chaîne qui identifie la porte.
- `capacity` : le nombre de participants que la porte peut traiter par tick.
- `queue` : un tableau de nombres représentant combien de participants arrivent à la porte pendant un tick précis.

Pour commencer, crée une fonction vide nommée `initializeThroughput` qui accepte un paramètre nommé `gates`.

# --hints--

Tu dois créer une fonction nommée `initializeThroughput`.

```js
assert.isFunction(initializeThroughput);
```

Ta fonction `initializeThroughput` doit avoir un paramètre `gates`.

```js
const regex = __helpers.functionRegex('initializeThroughput', ['gates']);
assert.match(__helpers.removeJSComments(initializeThroughput.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

--fcc-editable-region--

--fcc-editable-region--
```
