---
id: 69babdf3f896711c11516373
title: Étape 4
challengeType: 1
dashedName: step-4
---

# --description--

Pendant chaque tick de simulation :

- Un certain nombre de participants arrivent à la porte (depuis sa `queue`).
- La porte traite les participants en fonction de sa `capacity`.
- Si plus de participants arrivent que la porte ne peut en gérer, certains resteront (débordement, overflow).

Tu vas maintenant construire une fonction qui gère cette logique pour une porte à un seul tick.

Crée une fonction vide nommée `processGateFlow` qui accepte deux paramètres :

- `gate` : l'objet porte en cours de traitement.
- `tickIndex` : la position du tick de simulation courant dans le tableau `queue`.

# --hints--

Tu dois créer une fonction nommée `processGateFlow`.

```js
assert.isFunction(processGateFlow);
```

Ta fonction `processGateFlow` doit avoir les paramètres `gate` et `tickIndex`.

```js
const regex = __helpers.functionRegex('processGateFlow', ['gate', 'tickIndex']);
assert.match(__helpers.removeJSComments(processGateFlow.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

--fcc-editable-region--

--fcc-editable-region--
```
