---
id: 69bac027df64d4433c708d44
title: Étape 12
challengeType: 1
dashedName: step-12
---

# --description--

Lors de la redirection du débordement, tu dois envoyer les participants vers la porte suivante du tableau `gates`. Pour ce faire, tu dois trouver l'index de la porte suivante. Normalement, tu pourrais ajouter 1 à l'index courant. Cependant, si la porte courante est la dernière du tableau, tu dois revenir à la première porte.

Tu peux utiliser l'opérateur modulo (`%`) pour gérer cela :

```js
const nextIndex = (currentIndex + 1) % array.length;
```

Cela fonctionne parce que lorsque `currentIndex + 1` est égal à `array.length`, le résultat devient 0.

En utilisant cette approche, crée une variable nommée `nextGateIndex` qui stocke l'index de la porte suivante dans le tableau `gates`.

# --hints--

Tu dois utiliser la technique du modulo pour accéder à l'index suivant à partir de `currentIndex` dans le tableau `gates` et l'affecter à une variable nommée `nextGateIndex`. Suis le même schéma en une ligne que dans l'exemple.

```js
const rawFuncStr = __helpers.removeJSComments(rerouteOverflow.toString());
assert.match(
  rawFuncStr,
  /(var|const|let)\s+nextGateIndex/,
  "You must have at least one space between the declaration keyword (let/const) and 'nextGateIndex'."
);

const cleaned = __helpers.removeWhiteSpace(rawFuncStr);
const solution = /(var|const|let)nextGateIndex=\((currentIndex\+1|1\+currentIndex)\)%gates\.length;?}$/;

assert(
  solution.test(cleaned),
  "Inside your `rerouteOverflow` function, use the modulo technique to create a variable named `nextGateIndex` that stores the index of the next gate."
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
--fcc-editable-region--

--fcc-editable-region--
}
```
