---
id: 69babea99725b43a01adb6d1
title: Étape 5
challengeType: 1
dashedName: step-5
---

# --description--

D'abord, tu dois obtenir le nombre de participants arrivant pendant le tick courant. Accède à la valeur du tableau `queue` de la porte fournie à l'index `tickIndex`, et affecte-la à une variable nommée `currentTickQueue`.

# --hints--

Tu dois affecter l'élément situé à `gate.queue[tickIndex]` à une variable nommée `currentTickQueue`.

```js
const rawFuncStr = __helpers.removeJSComments(processGateFlow.toString());
assert.match(
  rawFuncStr,
  /(var|const|let)\s+currentTickQueue/,
  "You must have at least one space between the declaration keyword (let/const) and 'currentTickQueue'."
);

const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(processGateFlow.toString()));
const solution = /(var|const|let)currentTickQueue=gate\.queue\[tickIndex\];?}$/;

assert(
  solution.test(cleaned),
  "Inside your `processGateFlow` function, assign the element at `gate.queue[tickIndex]` to a variable named `currentTickQueue` as the only statement inside the function."
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
--fcc-editable-region--

--fcc-editable-region--
}
```
