---
id: 69babfd65c4c4632a0c76f63
title: Étape 10
challengeType: 1
dashedName: step-10
---

# --description--

Lorsqu'une porte ne peut pas traiter tous les participants pendant un tick, certains participants restent dans la file. Pour gérer ce débordement, tu vas construire une fonction qui redirige les participants en débordement vers une autre porte.

Crée une fonction vide nommée `rerouteOverflow` avec les paramètres suivants :

- `gates` : le tableau complet des objets porte.
- `currentGate` : la porte où le débordement s'est produit.
- `tickIndex` : la position du tick courant dans la file.
- `overflowAmount` : le nombre de participants qui n'ont pas pu être traités.

# --hints--

Tu dois créer une fonction nommée `rerouteOverflow`.

```js
assert.isFunction(rerouteOverflow);
```

Ta fonction `rerouteOverflow` doit avoir les paramètres `gates`, `currentGate`, `tickIndex` et `overflowAmount`.

```js
const regex = __helpers.functionRegex('rerouteOverflow', ['gates', 'currentGate', 'tickIndex', 'overflowAmount']);
assert.match(__helpers.removeJSComments(rerouteOverflow.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

--fcc-editable-region--

--fcc-editable-region--
```
