---
id: 66c07238b01053abaf812065
title: Crée un calculateur de factorielle
challengeType: 26
dashedName: build-a-factorial-calculator
---

# --description--

Une factorielle est le produit d'un entier et de tous les entiers inférieurs à lui. Par exemple, la factorielle de `5` est `5 * 4 * 3 * 2 * 1 = 120`. Dans ce lab, tu vas créer un calculateur de factorielle qui prend un nombre de l'utilisateur et calcule la factorielle de ce nombre.

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois déclarer une variable `num` et lui affecter un nombre de ton choix. Le nombre affecté doit être compris entre `1` et `20` inclus.
1. Crée une fonction nommée `factorialCalculator` qui prend un nombre en argument et retourne la factorielle de ce nombre.
1. À l'intérieur de la fonction, déclare une variable `result` et affecte-lui la valeur `1`. À l'aide d'une boucle, parcours tous les nombres de `1` au nombre fourni (inclus) et, pour chaque nombre, multiplie la variable `result` par le nombre courant et affecte le résultat à la variable `result`. Tu peux choisir d'utiliser ici une boucle `for`, une boucle `while` ou une boucle `do...while`.
1. Tu dois appeler la fonction `factorialCalculator` avec `num` comme argument et affecter le résultat à la variable `factorial`.
1. Tu dois stocker la sortie finale au format `Factorial of [num] is [factorial]` et l'affecter à la variable `resultMsg`.
1. Tu dois afficher la valeur de `resultMsg` dans la console.

# --hints--

Tu dois avoir une variable `num`.

```js
assert.isDefined(num);
```

Tu dois affecter une valeur à la variable `num`

```js
assert.isNotNull(num);
```

La valeur de `num` doit être comprise entre `1` et `20`.

```js
assert.isAtLeast(num, 1);
assert.isAtMost(num, 20);
```

La valeur de `num` doit être un nombre.

```js
assert.isNumber(num);
```

Tu dois avoir une fonction nommée `factorialCalculator`.

```js
assert.isFunction(factorialCalculator);
```

La fonction `factorialCalculator` doit prendre un nombre en argument.

```js
assert.match(factorialCalculator.toString(), /\s*factorialCalculator\(\s*\w+\s*\)/);
```

La factorielle de `5` doit retourner `120`.

```js
assert.strictEqual(factorialCalculator(5), 120);
```

La factorielle de `7` doit retourner `5040`.

```js
assert.strictEqual(factorialCalculator(7), 5040);
```

Tu dois appeler ta fonction `factorialCalculator` avec la variable `num` comme argument.

```js
assert.match(__helpers.removeJSComments(code), /factorialCalculator\(\s*num\s*\)\s*;?\s*$/m);
```

Tu dois définir une variable `factorial` et lui affecter le résultat de la fonction `factorialCalculator`.

```js
assert.isDefined(factorial);
assert.strictEqual(factorial, factorialCalculator(num));
```

Ta fonction `factorialCalculator` doit produire le résultat correct.

```js
assert.strictEqual(factorialCalculator(6), 720);
assert.strictEqual(factorialCalculator(9), 362880);
assert.strictEqual(factorialCalculator(11), 39916800);
```

Ta variable `resultMsg` doit contenir une chaîne au format `Factorial of [num] is [factorial]`.

```js
assert.strictEqual(resultMsg, `Factorial of ${num} is ${factorialCalculator(num)}`);
assert.notMatch(__helpers.removeJSComments(code), /Factorial of \d{1,2} is \d+/);
```

Tu dois afficher la valeur de `resultMsg` dans la console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(resultMsg\)\s*;?\s?$/m);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const num = 5;

function factorialCalculator(n) {
  let result = 1;

  for (let i = 2; i <= n; i++) {
    result *= i;
  }

  return result;
}

const factorial = factorialCalculator(num);
const resultMsg = `Factorial of ${num} is ${factorial}`;
console.log(resultMsg);
```
