---
id: 5f332b23c2045fb843337579
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

Pour indiquer aux visiteurs que le café a été fondé en 2020, ajoute un `p` en dessous du `h1` avec le texte `Est. 2020`.

(« Est. » est l'abréviation anglaise de **established** — « établi/fondé en ».)

# --hints--

Tu dois avoir une balise ouvrante `<p>`.

```js
assert(code.match(/<p>/i));
```

Tu dois avoir une balise fermante `</p>`.

```js
assert(code.match(/<\/p>/i));
```

Ne touche pas au `h1` — vérifie que la balise fermante est toujours là.

```js
assert.lengthOf(document.querySelectorAll('h1'),1);
```

Le `p` doit être **en dessous** du `h1`.

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.tagName, 'H1');
```

Le texte du `p` doit être `Est. 2020`.

```js
assert(document.querySelector("p").innerText === "Est. 2020");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
--fcc-editable-region--
      <h1>CAMPER CAFE</h1>
--fcc-editable-region--
    </main>
  </body>
</html>
```
