---
id: 5f344f9c805cd193c33d829c
title: Étape 11
challengeType: 0
dashedName: step-7
---

# --description--

Vous pouvez ajouter du style à un élément en le spécifiant dans le `style` élément et la configuration d'une propriété pour elle comme ceci:

```css
element {
 property: value;
}
```

Centrer le contenu du `h1` en fixant `text-align` propriété à la valeur `center`.

# --hints--

Vous devriez avoir un `h1` sélecteur dans votre `style` élément.

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1');
assert(hasSelector);
```

Votre `text-align` la propriété doit avoir une valeur de `center`.

```js
const hasTextAlign = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['text-align'] === 'center');
assert(hasTextAlign);
```

Votre `h1` sélecteur doit définir le `text-align` propriété `center`.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert(h1TextAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

