---
id: 5f344fbc22624a2976425065
title: Étape 5
challengeType: 0
dashedName: step-5
---

# --description--

Créer un `h2` élément de `section` élément et lui donner le texte `Coffee`.

# --hints--

Tu devrais avoir une ouverture. `<h2>` marque.

```js
assert(code.match(/<h2\s*>/i));
```

Tu devrais fermer. `</h2>` marque.

```js
assert(code.match(/<\/h2\s*>/i));
```

Vous ne devriez pas changer votre `section` élément. Assurez-vous de ne pas supprimer l'étiquette de fermeture.

```js
assert.lengthOf(document.querySelectorAll('section'),1);
```

Votre `h2` élément devrait être dans votre `section` élément.

```js
const h2 = document.querySelector('h2');
assert(h2.parentElement.tagName === 'SECTION');
```

Votre `h2` l'élément doit avoir le texte `Coffee`.

```js
const h2 = document.querySelector('h2');
assert(h2.innerText === 'Coffee');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
--fcc-editable-region--
      <section>
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```
