---
id: 5f3477ae34c1239cafe128be
title: Étape 9
challengeType: 0
dashedName: step-9
---

# --description--

Vous avez maintenant trois sélecteurs de type avec le même style. Vous pouvez ajouter le même groupe de styles à de nombreux éléments en créant une liste de sélecteurs. Chaque sélecteur est séparé par des virgules comme ceci :

```css
selector1, selector2 {
  property: value;
}
```

Supprimer les trois sélecteurs de type existants et les remplacer par une liste de sélecteur qui centre le texte pour le `h1`, `h2`et `p` éléments.

# --hints--

Vous devriez utiliser un sélecteur de type unique pour les trois éléments, `h1, h2, p`. Assurez-vous d'utiliser cet ordre.

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1, h2, p');
assert(hasSelector);
```

Vous ne devriez avoir qu'un sélecteur dans votre `style` élément.

```js
const selectors = new __helpers.CSSHelp(document).getCSSRules().filter(x => x.selectorText)
assert(selectors.length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
      h1 {
        text-align: center;
      }
      h2 {
        text-align: center;
      }
      p {
        text-align: center;
      }
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

