---
id: 5f34a1fd611d003edeafd681
title: Étape 19
challengeType: 0
dashedName: step-15
---

# --description--

Ce fond brun rend difficile la lecture du texte. Changer `body` couleur de fond de l'élément à `burlywood` Il a donc une certaine couleur, mais vous êtes toujours capable de lire le texte.

# --hints--

Vous devriez mettre le `background-color` propriété `burlywood`.

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-color'] === 'burlywood');
assert(hasBackground);
```

Votre `body` élément devrait avoir un `burlywood` Contexte.

```js
const bodyBackground = new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('background-color');
assert(bodyBackground === 'burlywood');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  background-color: brown;
}
--fcc-editable-region--
h1, h2, p {
  text-align: center;
}
```
