---
id: 5f344fad8bf01691e71a30eb
title: Étape 6
challengeType: 0
dashedName: step-6
---

# --description--

Jusqu'à maintenant, vous aviez un contrôle limité sur la présentation et l'apparence de votre contenu. Pour changer cela, ajouter un `style` élément à l'intérieur de `head` élément.

# --hints--

Votre code devrait avoir une ouverture `<style>` marque.

```js
assert(code.match(/<style\s*>/i));
```

Votre code devrait être fermé. `</style>` marque.

```js
assert(code.match(/<\/style\s*>/));
```

Votre `style` élément doit être imbriqué dans votre `head` élément.

```js
assert(code.match(/<head\s*>[\w\W\s]*<style\s*>[\w\W\s]*<\/style\s*>[\w\W\s]*<\/head\s*>/i))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
--fcc-editable-region--
</html>
```
