---
id: 5f3477cb2e27333b1ab2b955
title: Étape 12
challengeType: 0
dashedName: step-12
---

# --description--

Maintenant vous devez lier le `styles.css` fichier, donc les styles seront appliqués à nouveau. Dans le `head` élément, ajouter un `link` élément. Donnez-lui une `rel` attribut avec la valeur de `"stylesheet"` et une `href` attribut avec la valeur de `"styles.css"`.

Noter que les `link` élément est un élément vide.

# --hints--

Votre code devrait avoir un `link` élément.

```js
const link = document.querySelector('link');
assert.isNotNull(link);
```

Les `link` élément est un élément vide, il ne devrait pas avoir une balise de fin `</link>`.

```js
assert.notMatch(code, /<\/link>/i);
```

Vous ne devriez pas changer votre `head` élément. Assurez-vous de ne pas supprimer l'étiquette de fermeture.

```js
const headElementCount = document.querySelectorAll('head')?.length;
assert.strictEqual(headElementCount, 1);
```

Tu devrais en avoir une. `link` élément.

```js
const linkElementCount = document.querySelectorAll('link')?.length;
assert.strictEqual(linkElementCount, 1);
```

Votre `link` élément devrait être dans votre `head` élément.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

Votre `link` élément devrait avoir un `rel` attribut avec la valeur `stylesheet`.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

Votre `link` élément devrait avoir un `href` attribut avec la valeur `styles.css`.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
