---
id: 5f3477ae8466a9a3d2cc953c
title: Étape 15
challengeType: 0
dashedName: step-11
---

# --description--

Maintenant que vous avez le CSS dans le `styles.css` fichier, aller de l'avant et supprimer le `style` élément et tout son contenu. Une fois retiré, le texte qui a été centré retournera à gauche.

# --hints--

Vous ne devriez pas en avoir `style` tags dans votre code.

```js
assert(!code.match(/style/i));
```

Vous ne devriez pas avoir de sélecteurs CSS dans votre fichier HTML.

```js
assert(!editableContents.includes('style'));
assert(!editableContents.includes('text-align'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```

