---
id: 5f356ed63c7807a4f1e6d054
title: Étape 21
challengeType: 0
dashedName: step-17
---

# --description--

L'objectif est maintenant de `div` ne pas prendre toute la largeur de la page. Le CSS `width` propriété est parfait pour cela.

Vous pouvez utiliser le `id` sélecteur pour cibler un élément spécifique avec un `id` attribut. Une `id` Le sélecteur est défini en plaçant le symbole de hachage `#` directement devant l'élément `id` valeur. Par exemple, si un élément a `id` des `cat` alors vous viseriez cet élément comme ceci:

```css 
#cat {
  width: 250px;
}
```

Utilisez la `#menu` sélecteur pour donner à votre élément une largeur `300px`.

# --hints--

Vous devriez avoir un `#menu` sélecteur.

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle("#menu");
assert(hasDiv);
```

Vous devriez mettre le `width` propriété `300px`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '300px');
assert(hasWidth);
```

Votre `div` devrait avoir une largeur de 300px.

```js
const divWidth = new __helpers.CSSHelp(document).getStyle("#menu")?.getPropertyValue('width');
assert(divWidth === '300px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  background-color: burlywood;
}

h1, h2, p {
  text-align: center;
}
--fcc-editable-region--

```

