---
id: 5f356ed69db0a491745e2bb6
title: Étape 23
challengeType: 0
dashedName: step-23
---

# --description--

Pour appliquer le style de la classe à la `div` élément, supprimer le `id` ajouter un attribut `class` attribut à la `div` l'élément ouvre la balise. Assurez-vous de définir la valeur de classe à `menu`.

# --hints--

Votre `div` devrait encore rendre. Assurez-vous que vous n'avez pas malformé la `<div>` marque.

```js
assert.lengthOf(document.querySelectorAll('div'),1);
```

Votre `div` l'élément devrait avoir `menu` En cours.

```js
assert.isTrue(document.querySelector('div').classList.contains('menu'));
```

Votre `div` l'élément ne devrait plus avoir `menu` Oui.

```js
assert.lengthOf(document.querySelectorAll('div#menu'),0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
--fcc-editable-region--
    <div id="menu">
--fcc-editable-region--
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

