---
id: 5f356ed60785e1f3e9850b6e
title: Étape 20
challengeType: 0
dashedName: step-20
---

# --description--

Maintenant, il est facile de voir que le texte est centré dans le `#menu` élément. Actuellement, la largeur `#menu` l'élément est spécifié en pixels (`px`).

Changer `width` valeur de propriété à être `80%`, pour en faire 80% de la largeur de son élément parent (`body`).

# --hints--

Vous devriez mettre le `width` propriété `80%`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '80%');
assert(hasWidth);
```

Vous ne devriez pas avoir `width` biens de `300px`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '300px');
assert(!hasWidth);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}
--fcc-editable-region--
#menu {
  width: 300px;
  background-color: burlywood;
}
--fcc-editable-region--
```

