---
id: 5f3c866daec9a49519871816
title: Étape 30
challengeType: 0
dashedName: step-26
---

# --description--

`article` les éléments contiennent généralement plusieurs éléments qui ont des informations connexes. Dans ce cas, il contiendra une saveur de café et un prix pour cette saveur.

Nid deux `p` éléments dans votre `article` élément. Le premier texte devrait être `French Vanilla`, et le texte de la seconde `3.00`.

# --hints--

Vous ne devriez pas changer l'existant `article` élément.

```js
assert.lengthOf(document.querySelectorAll('article'),1);
```

Votre `article` élément devrait avoir deux `p` éléments.

```js
const article = document.querySelector('article');
const paragraphChildren = article?.querySelectorAll(`:scope ${'p'}`);
assert.lengthOf(paragraphChildren,2);
```

Votre premier `p` l'élément doit avoir le texte `French Vanilla`.

```js
const article = document.querySelector('article');
const paragraphChildren = article?.querySelectorAll(`:scope ${'p'}`);
const firstP = paragraphChildren?.[0];
assert.match(firstP?.innerText.trim(), /French Vanilla/i);
```

Votre deuxième `p` l'élément doit avoir le texte `3.00`.

```js
const article = document.querySelector('article');
const paragraphChildren = article?.querySelectorAll(`:scope ${'p'}`);
const secondP = paragraphChildren?.[1];
assert.match(secondP?.innerText.trim(), /3\.00/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
--fcc-editable-region--
          <article>
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
