---
id: 5f3c866de7a5b784048f94b1
title: Étape 32
challengeType: 0
dashedName: step-32
---

# --description--

C'est un peu ce que vous voulez, mais maintenant ce serait bien si la saveur et le prix étaient sur la même ligne. `p` les éléments sont: <dfn>niveau bloc</dfn> les éléments, donc ils prennent toute la largeur de leur élément parent.

Pour les obtenir sur la même ligne, vous devez appliquer un style à la `p` éléments de sorte qu'ils se comportent plus comme <dfn>en ligne</dfn> éléments.

Pour cela, commencez par ajouter un `class` attribut avec la valeur `item` à la première `article` élément de `Coffee` En route.

# --hints--

Vous devriez appliquer la `item` cours à votre `article` élément.

```js
const el = document.querySelector('article.item');
assert.exists(el);
```

Tu ne devrais en avoir qu'un. `item` élément de classe.

```js
const elements = document.querySelectorAll('.item');
assert.lengthOf(elements, 1);
```

Votre premier `article` l'élément devrait avoir `item` En cours.

```js
assert.equal(document.querySelector('article')?.className, 'item');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
--fcc-editable-region--
          <h2>Coffee</h2>
          <article>
            <p class="flavor">French Vanilla</p>
            <p class="price">3.00</p>
          </article>
--fcc-editable-region--
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.flavor {
  text-align: left;
}

.price {
  text-align: right;
}
```
