---
id: 66f53dc2c5bd6a11d6c3282f
title: Crée une application génératrice de mots de passe
challengeType: 26
dashedName: lab-password-generator
---

# --description--

Dans ce lab, tu vas t'entraîner à utiliser les fonctions en créant un générateur de mots de passe aléatoires.

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois créer une fonction appelée `generatePassword` qui prend un paramètre indiquant la longueur du mot de passe généré. Tu peux nommer ce paramètre comme tu veux.
2. Ta fonction doit renvoyer une chaîne qui représente un mot de passe généré aléatoirement. Tu dois utiliser la chaîne suivante et différentes méthodes `Math` pour t'aider à renvoyer une nouvelle chaîne avec des caractères aléatoires : `ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()`.
3. Tu dois définir une variable appelée `password` et lui assigner le résultat de l'appel à la fonction `generatePassword` avec un argument numérique qui représente la longueur souhaitée du mot de passe.
4. Tu dois avoir un `console.log` qui affiche une seule chaîne formée en concaténant le message `Generated password:` et la variable `password`, séparés par un espace.

# --hints--

Tu dois avoir une fonction `generatePassword` avec un paramètre. Tu peux nommer ce paramètre comme tu veux.

```js
assert.isFunction(generatePassword);
assert.lengthOf(generatePassword, 1);
```

Ta fonction `generatePassword` doit renvoyer une chaîne.

```js
const result = generatePassword(5);
assert.isString(result);
```

Ta fonction `generatePassword` doit renvoyer une nouvelle chaîne de la bonne longueur.

```js
const length = 8;
const password = generatePassword(length);
assert.lengthOf(password, length);
```

Ta fonction doit renvoyer un mot de passe généré aléatoirement avec des caractères valides.

```js
const validChars = "ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()";
const password = generatePassword(10);
for (let char of password) {
    assert.include(validChars, char);
}
```

Ta fonction doit renvoyer un mot de passe généré aléatoirement qui contient plus d'un caractère unique.

```js
const passwordSetSize1 = new Set(generatePassword(12)).size;
const passwordSetSize2 = new Set(generatePassword(12)).size;
const passwordSetSize3 = new Set(generatePassword(12)).size;
assert.isAbove(passwordSetSize1 + passwordSetSize2 + passwordSetSize3, 3);
```

Ta fonction doit renvoyer une nouvelle chaîne aléatoire à chaque appel.

```js
const password1 = generatePassword(10);
const password2 = generatePassword(10);
assert.notStrictEqual(password1, password2);
```

Tu dois avoir une variable `password`.

```js
assert.isDefined(password);
```

Ta variable `password` doit être une chaîne.

```js
assert.isString(password);
```

Tu dois appeler la fonction `generatePassword` avec un argument numérique et stocker le mot de passe renvoyé dans la variable `password`.

```js
assert.isDefined(password);
assert.isString(password);
assert.match(password, /^[A-Za-z0-9!@#$%^&*()]+$/);
const length2 = 10;
const password3 = generatePassword(length2);
assert.lengthOf(password3, length2);
assert.equal(password3.length, generatePassword(length2).length);
assert.match(__helpers.removeJSComments(code), /(let|const|var)\s+password\s*=\s*generatePassword\(\d+\)\;?/);
```

Tu dois afficher une seule chaîne combinant `Generated password:` et `password`, séparés par un espace, en utilisant `+` ou un littéral de modèle.

```js
const cleanCode = __helpers.removeJSComments(code);
const condition1 = /console\.log\(\s*["']Generated\s+password:\s*["']\s*\+\s*password\s*\);?/gm.test(cleanCode);
const condition2 = /console\.log\(\s*`Generated\s+password:\s*\$\{password\}`\s*\);?/gm.test(cleanCode);
assert.isTrue(condition1 || condition2);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const generatePassword = (length) => {
    const chars = "ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()";
    let password = "";
    for (let i = 0; i < length; i++) {
        const randomIndex = Math.floor(Math.random() * chars.length);
        password += chars[randomIndex];
    }
    return password;
};

const password = generatePassword(12);
console.log(`Generated password: ${password}`);
```
