---
id: 68ed661e9f463c155bb677a1
title: Étape 1
challengeType: 0
dashedName: step-1
---

# --description--

Dans cet atelier, tu vas t’entraîner à utiliser CSS Flexbox en concevant un ensemble de boîtes colorées.

Pour commencer, dans l’élément `body`, ajoute un élément `header`. Imbrique un élément `h1` dans l’élément `header` avec le texte `Colored Boxes Layout`.

# --hints--

Tu devrais avoir un élément `header` dans l’élément `body`.

```js
const headerEl = document.querySelector('body > header');
assert.exists(headerEl);
```

Ton élément `header` devrait contenir un élément `h1`.

```js
const h1El = document.querySelector('body > header > h1');
assert.exists(h1El);
```

Ton élément `h1` devrait contenir le texte `Colored Boxes Layout`.

```js
const h1El = document.querySelector('body > header > h1');
assert.equal(h1El?.textContent.trim(), 'Colored Boxes Layout');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
</head>
<body>
--fcc-editable-region--

--fcc-editable-region--
</body>
</html>
```
