---
id: 68ef0b3237fef991c92d9c19
title: Étape 4
challengeType: 0
dashedName: step-4
---

# --description--

Ajoute maintenant un élément `div` avec la classe `flex-container` sous ton `header`. Ensuite, crée six éléments `div` avec la classe `box` dans le `div` que tu viens de créer.

Rappelle-toi qu’un attribut de classe sert souvent à pointer vers un nom de classe dans une feuille de style. Ici, tous les éléments `div` seront stylisés de la même façon selon les définitions de style `.flex-container` et `.box`.

# --hints--

Tu devrais avoir un élément `div` sous l’élément `header`.

```js
const divELBelowHeader = document.querySelector('header + div');
assert.exists(divELBelowHeader);
```

Ton élément `div` sous l’élément `header` devrait avoir la classe `flex-container`.

```js
const divELBelowHeader = document.querySelector('header + div');
assert.equal(divELBelowHeader.className, 'flex-container');
```

Tu devrais avoir six éléments `div` imbriqués dans l’élément `.flex-container`.

```js
const divElsInsideContainer = document.querySelectorAll('div.flex-container > div');
assert.lengthOf(divElsInsideContainer, 6);
```

Tes six éléments `div` dans l’élément `.flex-container` devraient avoir la classe `box`.

```js
const divElsInsideContainerWithBox = document.querySelectorAll('div.flex-container > div.box');
assert.lengthOf(divElsInsideContainerWithBox, 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
--fcc-editable-region--

--fcc-editable-region--
</body>
</html>
```

```css
h1 {
	text-align: center;
	margin-bottom: 10px;
}
```
