---
id: 68f91df5776e693049fb8942
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

Il est maintenant temps d’ajouter quelques règles CSS au fichier `styles.css`. Commence par créer un sélecteur pour l’élément `h1`.

Centre le contenu de l’élément `h1` en définissant sa propriété `text-align` sur `center`. Crée aussi une propriété `margin-bottom` avec la valeur `10px` pour ajouter une marge entre le bas de l’élément `h1` et tout élément HTML placé en dessous.

# --hints--

Tu devrais avoir un sélecteur `h1`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('h1'));
```

Ton sélecteur `h1` devrait définir la propriété `text-align` sur `center`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.getPropVal('text-align'), 'center');
```

Ton sélecteur `h1` devrait définir la propriété `margin-bottom` sur `10px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.getPropVal('margin-bottom'), '10px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
