---
id: 692f22933ba1cb86cd75a70c
title: Étape 6
challengeType: 0
dashedName: step-6
---

# --description--

Il est maintenant temps de créer le nom de couleur des boîtes. Pour cela, ajoute un élément `p` sous chaque élément `h2` avec ces noms de couleur dans l’ordre : `Red`, `Orange`, `Yellow`, `Green`, `Blue` et `Indigo`. Une couleur par boîte.

# --hints--

Tu devrais ajouter un élément `p` avec le texte `Red` sous l’élément `h2` dans le premier élément `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[0];
assert.equal(heading.textContent, "Red");
```

Tu devrais ajouter un élément `p` avec le texte `Orange` sous l’élément `h2` dans le deuxième élément `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[1];
assert.equal(heading.textContent, "Orange");
```

Tu devrais ajouter un élément `p` avec le texte `Yellow` sous l’élément `h2` dans le troisième élément `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[2];
assert.equal(heading.textContent, "Yellow");
```

Tu devrais ajouter un élément `p` avec le texte `Green` sous l’élément `h2` dans le quatrième élément `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[3];
assert.equal(heading.textContent, "Green");
```

Tu devrais ajouter un élément `p` avec le texte `Blue` sous l’élément `h2` dans le cinquième élément `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[4];
assert.equal(heading.textContent, "Blue");
```

Tu devrais ajouter un élément `p` avec le texte `Indigo` sous l’élément `h2` dans le sixième élément `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[5];
assert.equal(heading.textContent, "Indigo");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
--fcc-editable-region--
    <div class="box">
      <h2>Box 1</h2>

    </div>
    <div class="box">
      <h2>Box 2</h2>

    </div>
    <div class="box">
      <h2>Box 3</h2>

    </div>
    <div class="box">
      <h2>Box 4</h2>

    </div>
    <div class="box">
      <h2>Box 5</h2>

    </div>
    <div class="box">
      <h2>Box 6</h2>

    </div>
--fcc-editable-region--
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}
```
