---
id: 692f30e86027a4b253d0a42d
title: Étape 16
challengeType: 0
dashedName: step-16
---

# --description--

Ajoute la propriété `font-size` avec la valeur `1.125rem` au sélecteur `.box`.

La propriété `font-size` définit la taille du texte. Elle peut être définie avec différentes unités comme les pixels, `em`, `rem`, les pourcentages, etc., ce qui la rend flexible pour les designs responsives.

Dans ce cas, la valeur utilise l’unité `rem`, qui est relative à l’élément racine `html` et non à l’élément parent. Cela signifie que la taille de police reste cohérente et n’est pas affectée par la taille des éléments parents, ce qui rend les mises en page plus prévisibles.

# --hints--

Ton sélecteur `.box` devrait avoir la propriété `font-size` avec la valeur `1.125rem`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('font-size'), '1.125rem');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
--fcc-editable-region--

--fcc-editable-region--
}
```
