---
id: 692f22a2ad44ae887a24c8f9
title: Étape 5
challengeType: 0
dashedName: step-5
---

# --description--

Dans chaque élément `.box`, ajoute un élément `h2` avec le texte `Box`, suivi d’un espace et d’un numéro séquentiel à partir de `1`.

# --hints--

Tu devrais avoir un élément `h2` avec le texte `Box 1` dans le premier `div` ayant la classe `box`.

```js
const firstH2El = document.querySelector("div.box:nth-of-type(1) > h2");
assert.equal(firstH2El.textContent, "Box 1");
```

Tu devrais avoir un élément `h2` avec le texte `Box 2` dans le deuxième `div` ayant la classe `box`.

```js
const secondH2El = document.querySelector("div.box:nth-of-type(2) > h2");
assert.equal(secondH2El.textContent, "Box 2");
```

Tu devrais avoir un élément `h2` avec le texte `Box 3` dans le troisième `div` ayant la classe `box`.

```js
const thirdH2El = document.querySelector("div.box:nth-of-type(3) > h2");
assert.equal(thirdH2El.textContent, "Box 3");
```

Tu devrais avoir un élément `h2` avec le texte `Box 4` dans le quatrième `div` ayant la classe `box`.

```js
const fourthH2El = document.querySelector("div.box:nth-of-type(4) > h2");
assert.equal(fourthH2El.textContent, "Box 4");
```

Tu devrais avoir un élément `h2` avec le texte `Box 5` dans le cinquième `div` ayant la classe `box`.

```js
const fifthH2El = document.querySelector("div.box:nth-of-type(5) > h2");
assert.equal(fifthH2El.textContent, "Box 5");
```

Tu devrais avoir un élément `h2` avec le texte `Box 6` dans le dernier `div` ayant la classe `box`.

```js
const sixthH2El = document.querySelector("div.box:nth-of-type(6) > h2");
assert.equal(sixthH2El.textContent, "Box 6");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
--fcc-editable-region--
		<div class="box">

		</div>
    <div class="box">

    </div>
    <div class="box">

    </div>
    <div class="box">

    </div>
    <div class="box">

    </div>
    <div class="box">

    </div>
--fcc-editable-region--
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}
```
