---
id: 692f2c20cb7f9e5050b935a7
title: Étape 9
challengeType: 0
dashedName: step-9
---

# --description--

Ajoute la propriété `border` avec la valeur `1px solid #000` au sélecteur `.box`. La propriété `border` sert à créer un contour visible autour d’un élément. Elle peut être personnalisée avec `border-width`, `border-style` et `border-color`.

Dans ce cas, `1px` correspond à `border-width`, qui définit l’épaisseur de la bordure. `solid` correspond à `border-style` et définit l’apparence de la bordure (pleine, en tirets, en pointillés, etc.). Et `#000` correspond à `border-color`, qui définit la couleur de la bordure.

# --hints--

Ton sélecteur `.box` devrait avoir la propriété `border` avec la valeur `1px solid #000`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('border'), '1px solid rgb(0, 0, 0)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
--fcc-editable-region--

--fcc-editable-region--
}
```
