---
id: 68a23004e96a105df3a21fc4
title: Crée un assistant de comptage de cartes
challengeType: 26
dashedName: lab-counting-cards
---

# --description--

Dans le jeu de casino Blackjack, un joueur peut déterminer s'il a un avantage sur la maison pour la prochaine main en gardant la trace du nombre relatif de cartes hautes et basses restant dans le paquet. C'est ce qu'on appelle le comptage de cartes.

Avoir plus de cartes hautes restantes dans le paquet favorise le joueur. Lorsque le compte est positif, le joueur devrait miser haut. Lorsque le compte est nul ou négatif, le joueur devrait miser bas.

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois utiliser `let` pour déclarer une variable globale nommée `count` et la définir à `0`.
1. Tu dois avoir une fonction appelée `cardCounter`.
1. La fonction `cardCounter` doit recevoir un paramètre `card` qui peut être soit un nombre, soit une chaîne.
    - Pour les valeurs entre `2` et `10`, le paramètre `card` sera un nombre.
    - Pour toutes les autres valeurs, le paramètre `card` sera une chaîne.
1. La fonction `cardCounter` doit modifier la variable globale `count` selon certains critères.
1. La variable globale `count` doit être augmentée de `1` pour les cartes `2`, `3`,`4`, `5` ou `6`
1. La variable globale `count` doit rester inchangée pour les cartes `7`, `8`, `9`.
1. La variable globale `count` doit être diminuée de `1` pour les cartes `10`, `"J"`, `"Q"`, `"K"`, `"A"`
1. La fonction `cardCounter` doit retourner une chaîne avec le compte actuel et la chaîne `Bet` si le compte est positif.
1. La fonction `cardCounter` doit retourner une chaîne avec le compte actuel et la chaîne `Hold` si le compte est inférieur ou égal à `0`.
1. Dans la sortie de la fonction, le compte actuel et la décision du joueur (`Bet` ou `Hold`) doivent être séparés par un espace. Par exemple, `-3 Hold`.

# --hints--

Tu dois utiliser `let` pour déclarer une variable globale nommée `count` et la définir à `0`.

```js
assert.match(__helpers.removeJSComments(code), /^\s*let\s+count\s*=\s*0\s*;?\s*$/m);
```

Tu dois avoir une fonction nommée `cardCounter`.

```js
assert.isFunction(cardCounter);
```

Ta fonction doit retourner la valeur de `count` et le texte (`Bet` ou `Hold`) avec un seul caractère espace entre les deux.

```js
count = 0;
let out = cardCounter(10);
const pattern = /^-?\d+ (Bet|Hold)$/;
assert.match(out, pattern);
```

Après les cartes `2`, `3`, `4`, `5`, l'appel de `cardCounter(6)` doit retourner la chaîne `5 Bet`.

```js
count = 0;
cardCounter(2);
cardCounter(3);
cardCounter(4);
cardCounter(5);
const out = cardCounter(6);
assert.strictEqual(out, '5 Bet')
```

Après les cartes `7`, `8`, l'appel de `cardCounter(9)` doit retourner la chaîne `0 Hold`.

```js
count = 0;
cardCounter(7);
cardCounter(8);
const out = cardCounter(9);
assert.strictEqual(out, '0 Hold');
```

Après les cartes `10`, `"J"`, `"Q"`, `"K"`, l'appel de `cardCounter("A")` doit retourner la chaîne `-5 Hold`.

```js
count = 0;
cardCounter(10);
cardCounter('J');
cardCounter('Q');
cardCounter('K');
const out = cardCounter('A');
assert.strictEqual(out, '-5 Hold');
```

Après les cartes `3`, `7`, `"Q"`, `8`, l'appel de `cardCounter("A")` doit retourner la chaîne `-1 Hold`.

```js
count = 0;
cardCounter(3);
cardCounter(7);
cardCounter('Q');
cardCounter(8);
const out = cardCounter('A');
assert.strictEqual(out, '-1 Hold');
```

Après les cartes `2`, `"J"`, `9`, `2`, l'appel de `cardCounter(7)` doit retourner la chaîne `1 Bet`.

```js
count = 0;
cardCounter(2);
cardCounter('J');
cardCounter(9);
cardCounter(2);
const out = cardCounter(7);
assert.strictEqual(out, '1 Bet');
```

Après les cartes `2`, `2`, l'appel de `cardCounter(10)` doit retourner la chaîne `1 Bet`.

```js
count = 0;
cardCounter(2);
cardCounter(2);
const out = cardCounter(10);
assert.strictEqual(out, '1 Bet')
```

Après les cartes `3`, `2`, `"A"`, `10`, l'appel de `cardCounter("K")` doit retourner la chaîne `-1 Hold`.

```js
count = 0;
cardCounter(3);
cardCounter(2);
cardCounter('A');
cardCounter(10);
const out = cardCounter('K');
assert.strictEqual(out, '-1 Hold');
```

# --seed--

## --seed-contents--


```js



```

# --solutions--

```js
let count = 0;
function cardCounter(card) {
  switch(card) {
    case 2:
    case 3:
    case 4:
    case 5:
    case 6:
      count++;
      break;
    case 10:
    case 'J':
    case 'Q':
    case 'K':
    case 'A':
      count--;
  }
  if(count > 0) {
    return count + " Bet";
  } else {
    return count + " Hold";
  }
}
```
