---
id: 6998b99a8e195cd0c1f211df
title: Étape 1
challengeType: 1
dashedName: step-1
---

# --description--

Dans ce workshop, tu vas construire un simple traqueur de faune en utilisant les objets JavaScript.

Voici un exemple d'objet :

```js
const animal = {}; // empty object
```

Cela crée un objet vide.

Maintenant, crée une variable nommée `tiger` et affecte-lui un objet vide.

# --hints--

Tu dois créer une variable nommée `tiger`.

```js
assert.exists(tiger);
```

La variable `tiger` doit se voir affecter un objet vide.

```js
assert.deepEqual(tiger, {});
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
