---
id: 6999faf886634b72be7a819b
title: Étape 2
challengeType: 1
dashedName: step-2
---

# --description--

Pour ce traqueur de faune, tu vas mettre à jour un objet avec de nouvelles propriétés.

Voici un exemple de création d'un objet avec une propriété :

```js
const animal = {
  name: "Lion" // property: value
};
```

Maintenant, mets à jour l'objet `tiger` pour qu'il inclue une propriété appelée `species`.

Définis sa valeur à `"Tiger"`.

# --hints--

`tiger` doit être un objet.

```js
assert.isObject(tiger);
```

`tiger` doit avoir une propriété `species`.

```js
assert.property(tiger, "species");
```

`tiger.species` doit être égal à `"Tiger"`.

```js
assert.propertyVal(tiger, "species", "Tiger");
```

# --seed--

## --seed-contents--

```js
const tiger = {
--fcc-editable-region--

--fcc-editable-region--
};
```
