---
id: 6999adc72c02812a28c01e31
title: Étape 9
challengeType: 1
dashedName: step-9
---

# --description--

Dans cette étape, tu vas créer une fonction qui met à jour une propriété d'un objet.

Voici un exemple de mise à jour d'une propriété à l'intérieur d'une fonction :

```js
const dog = {
  age: 4
};

const changeAge = (pet, updatedAge) => {
  pet.age = updatedAge; // update existing property using dot notation
  return pet; // return the updated object
}

console.log(changeAge(dog, 6)); // { age: 6 }
```

Dans cet exemple, la propriété `age` est mise à jour avec une nouvelle valeur.

Maintenant, crée une fonction appelée `updateAge`. La fonction doit prendre deux paramètres : `animal` et `newAge`.

À l'intérieur de la fonction, mets à jour la propriété `age` de l'objet `animal` avec `newAge`. Retourne l'objet `animal` mis à jour.

Après avoir créé la fonction, utilise `console.log` pour appeler `updateAge(elephant, 12)` afin de voir l'objet `elephant` mis à jour dans la console.

# --hints--

Tu dois créer une fonction nommée `updateAge`.

```js
assert.isFunction(updateAge);
```

La fonction `updateAge` doit avoir deux paramètres : `animal` et `newAge`.

```js
const regex = __helpers.functionRegex('updateAge', ['animal', 'newAge']);
assert.match(__helpers.removeJSComments(code), regex);
```

`updateAge` doit utiliser la notation par point pour mettre à jour la propriété `age`.

```js
assert.match(code, /animal\.age\s*=\s*newAge/);
```

La fonction `updateAge` doit retourner l'objet `animal` mis à jour.

```js
const testAnimal = { species: "Dog", age: 4 };
const result = updateAge(testAnimal, 6);
assert.strictEqual(result, testAnimal);
```

Tu dois afficher `updateAge(elephant, 12)` dans la console.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*updateAge\s*\(\s*elephant\s*,\s*12\s*\)\s*\)/
);
```

Appeler `updateAge(elephant, 12)` doit mettre à jour la propriété age.

```js
const updatedElephant = updateAge(elephant, 12);

assert.propertyVal(updatedElephant, "age", 12);
```

`updateAge` doit utiliser les paramètres de la fonction et fonctionner avec n'importe quel objet.

```js
const lion = { species: "Lion", age: 4 };
const updatedLion = updateAge(lion, 9);

assert.propertyVal(updatedLion, "age", 9);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

--fcc-editable-region--

--fcc-editable-region--
```
