---
id: 6999ac171928136168654774
title: Étape 6
challengeType: 1
dashedName: step-6
---

# --description--

Dans cette étape, tu vas créer une fonction pour accéder à la propriété `species` d'un objet.

Voici un petit exemple d'accès à une propriété d'un objet :

```js
const dog = {
  species: "Dog"
};
const getAnimalSpecies = (pet) => {
  return pet.species; // access species using dot notation
};
console.log(getAnimalSpecies(dog)); // Dog
```

Dans cet exemple, la fonction prend un objet en paramètre et retourne sa propriété `species`.

Maintenant, crée une fonction appelée `getSpecies`.

La fonction doit prendre un paramètre nommé `animal`.

Retourne la propriété `species` de l'`animal` en utilisant la notation par point.

Après avoir créé la fonction, utilise `console.log` pour appeler `getSpecies(tiger)` afin de voir le résultat dans la console.

# --hints--

Tu dois créer une fonction nommée `getSpecies`.

```js
assert.isFunction(getSpecies);
```

La fonction `getSpecies` doit avoir un seul paramètre, `animal`.

```js
const regex = __helpers.functionRegex('getSpecies', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

Tu dois afficher `getSpecies(tiger)` dans la console.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*getSpecies\s*\(\s*tiger\s*\)\s*\)/);
```

`getSpecies` doit accéder à la propriété species en utilisant la notation par point.

```js
assert.match(code, /animal\.species/);
```

Appeler `getSpecies(tiger)` doit retourner `"Tiger"`.

```js
assert.strictEqual(getSpecies(tiger), "Tiger");
```

`getSpecies` doit utiliser le paramètre de la fonction et fonctionner avec n'importe quel objet.

```js
const lion = { species: "Lion" };
assert.strictEqual(getSpecies(lion), "Lion");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

--fcc-editable-region--

--fcc-editable-region--
```
