---
id: 69999ff7f7862fdfb885ea5d
title: Étape 3
challengeType: 1
dashedName: step-3
---

# --description--

Maintenant, mets à jour ton objet `tiger` en ajoutant une nouvelle propriété appelée `age`.

Définis la propriété `age` à la valeur `5`.

# --hints--

`tiger` doit avoir une propriété `age`.

```js
assert.property(tiger, "age");
```

`tiger.age` doit être égal à `5`.

```js
assert.propertyVal(tiger, "age", 5);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
--fcc-editable-region--

--fcc-editable-region--
};
```
