---
id: 6999b07c34c2cde21ee6e1d8
title: Étape 11
challengeType: 1
dashedName: step-11
---

# --description--

Dans cette étape, tu vas créer une fonction qui vérifie si un objet possède une propriété précise.

Voici un exemple d'utilisation de `hasOwnProperty` :

```js
const cat = {
  species: "Cat",
  color: "White"
};

const dog = {
  species: "Dog",
  weight: 50
};

const hasColor = (pet) => {
  return pet.hasOwnProperty("color"); // check if "color" property exists
};

console.log(hasColor(cat)); // true
console.log(hasColor(dog)); // false
```

Dans cet exemple, la fonction `hasColor` utilise `hasOwnProperty` pour vérifier si la propriété `color` existe sur un objet.

Maintenant, crée une fonction appelée `hasHabitat`. La fonction doit prendre un paramètre nommé `animal`.

Utilise la méthode `hasOwnProperty` pour retourner `true` si l'objet `animal` a une propriété appelée `"habitat"` et `false` sinon.

Après avoir créé la fonction, utilise `console.log` pour appeler `hasHabitat(tiger)` et `hasHabitat(elephant)` afin de voir les deux résultats dans la console.

# --hints--

Tu dois créer une fonction nommée `hasHabitat`.

```js
assert.isFunction(hasHabitat);
```

La fonction `hasHabitat` doit avoir un seul paramètre, `animal`.

```js
const regex = __helpers.functionRegex('hasHabitat', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

`hasHabitat` doit utiliser la méthode `hasOwnProperty` pour vérifier la propriété `habitat`.

```js
assert.match(code, /animal\.hasOwnProperty\s*\(\s*["']habitat["']\s*\)/);
```

Tu dois afficher `hasHabitat(tiger)` dans la console.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*hasHabitat\s*\(\s*tiger\s*\)\s*\)/);
```

Tu dois afficher `hasHabitat(elephant)` dans la console.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*hasHabitat\s*\(\s*elephant\s*\)\s*\)/);
```

Appeler `hasHabitat(tiger)` doit retourner `true` après l'ajout d'un habitat.

```js
addHabitat(tiger, "Rainforest");
assert.strictEqual(hasHabitat(tiger), true);
```

Appeler `hasHabitat(elephant)` doit retourner `false` si aucun habitat n'existe.

```js
const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

assert.strictEqual(hasHabitat(elephant), false);
```

`hasHabitat` doit utiliser les paramètres de la fonction et fonctionner avec n'importe quel objet.

```js
const lion = { species: "Lion", habitat: "Savanna" };
assert.strictEqual(hasHabitat(lion), true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

--fcc-editable-region--

--fcc-editable-region--
```
