---
id: 6999b0198d18d402c36cc5fe
title: Étape 10
challengeType: 1
dashedName: step-10
---

# --description--

Dans cette étape, tu vas créer une fonction qui supprime une propriété d'un objet.

Voici un exemple de suppression d'une propriété en utilisant `delete` :

```js
const bird = {
  species: "Parrot",
  canFly: true
};

const removeFlight = (pet) => {
  delete pet.canFly; // remove property using delete keyword
  return pet; // return the updated object
};

console.log(removeFlight(bird));
// { species: "Parrot" }
```

Dans cet exemple, la propriété `canFly` est supprimée de l'objet `bird`.

Maintenant, crée une fonction appelée `removeEndangeredStatus`. La fonction doit prendre un paramètre nommé `animal`.

À l'intérieur de la fonction, supprime la propriété `isEndangered` de l'objet `animal` en utilisant le mot-clé `delete`. Retourne l'objet `animal` mis à jour.

Après avoir créé la fonction, utilise `console.log` pour appeler `removeEndangeredStatus(tiger)` afin de voir l'objet mis à jour dans la console.

# --hints--

Tu dois créer une fonction nommée `removeEndangeredStatus`.

```js
assert.isFunction(removeEndangeredStatus);
```

La fonction `removeEndangeredStatus` doit avoir un seul paramètre, `animal`.

```js
const regex = __helpers.functionRegex('removeEndangeredStatus', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

La fonction `removeEndangeredStatus` doit retourner l'objet `animal` mis à jour.

```js
const testAnimal = { species: "Cat", isEndangered: true };
const result = removeEndangeredStatus(testAnimal);
assert.strictEqual(result, testAnimal);
```

`removeEndangeredStatus` doit utiliser le mot-clé `delete` pour supprimer la propriété `isEndangered`.

```js
assert.match(code, /delete\s+animal\.isEndangered/);
```

Tu dois afficher `removeEndangeredStatus(tiger)` dans la console.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*removeEndangeredStatus\s*\(\s*tiger\s*\)\s*\)/);
```

Appeler `removeEndangeredStatus(tiger)` doit supprimer la propriété `isEndangered`.

```js
const updatedTiger = removeEndangeredStatus(tiger);

assert.notProperty(updatedTiger, "isEndangered");
```

`removeEndangeredStatus` doit utiliser les paramètres de la fonction et fonctionner avec n'importe quel objet.

```js
const lion = {
  species: "Lion",
  age: 4,
  isEndangered: true
};

const updatedLion = removeEndangeredStatus(lion);

assert.notProperty(updatedLion, "isEndangered");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

--fcc-editable-region--

--fcc-editable-region--
```
