---
id: 6999b0d79ea938edcb032237
title: Étape 12
challengeType: 1
dashedName: step-12
---

# --description--

Dans cette dernière étape du workshop Wildlife Tracker, tu vas créer une fonction qui accède à une propriété d'un objet en utilisant la notation par crochets.

Voici un exemple d'accès à une propriété en utilisant la notation par crochets :

```js
const cat = {
  species: "Cat",
  age: 3
};

const property = "species";

console.log(cat[property]); // "Cat"
```

Dans cet exemple, la notation par crochets te permet d'accéder à une propriété en utilisant une variable.

Maintenant, crée une fonction appelée `getProperty`.

La fonction doit prendre deux paramètres : `animal` et `propertyName`.

Retourne la valeur de la propriété en utilisant la notation par crochets.

Après avoir créé la fonction, utilise `console.log` pour appeler `getProperty(tiger, "species")` et `getProperty(elephant, "age")` afin de voir les valeurs retournées dans la console.

Avec ça, le workshop Wildlife Tracker est terminé !

# --hints--

Tu dois créer une fonction nommée `getProperty`.

```js
assert.isFunction(getProperty);
```

La fonction `getProperty` doit avoir deux paramètres : `animal` et `propertyName`.

```js
const regex = __helpers.functionRegex('getProperty', ['animal', 'propertyName']);
assert.match(__helpers.removeJSComments(code), regex);
```

`getProperty` doit utiliser la notation par crochets pour accéder à la propriété.

```js
assert.match(code, /animal\s*\[\s*propertyName\s*\]/);
```

Tu dois afficher `getProperty(tiger, "species")` dans la console.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*getProperty\s*\(\s*tiger\s*,\s*["']species["']\s*\)\s*\)/
);
```

Tu dois afficher `getProperty(elephant, "age")` dans la console.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*getProperty\s*\(\s*elephant\s*,\s*["']age["']\s*\)\s*\)/
);
```

Appeler `getProperty(tiger, "species")` doit retourner `"Tiger"`.

```js
assert.strictEqual(getProperty(tiger, "species"), "Tiger");
```

Appeler `getProperty(elephant, "age")` doit retourner `12`.

```js
assert.strictEqual(getProperty(elephant, "age"), 12);
```

`getProperty` doit utiliser les paramètres de la fonction et fonctionner avec n'importe quel objet.

```js
const lion = {
  species: "Lion",
  weight: 200
};

assert.strictEqual(getProperty(lion, "weight"), 200);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

const hasHabitat = (animal) => {
  return animal.hasOwnProperty("habitat");
};

console.log(hasHabitat(tiger));
console.log(hasHabitat(elephant));

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

const hasHabitat = (animal) => {
  return animal.hasOwnProperty("habitat");
};

console.log(hasHabitat(tiger));
console.log(hasHabitat(elephant));

const getProperty = (animal, propertyName) => {
  return animal[propertyName];
};

console.log(getProperty(tiger, "species"));
console.log(getProperty(elephant, "age"));
```
