---
id: 6999a0c1698b93d5d41156ed
title: Étape 4
challengeType: 1
dashedName: step-4
---

# --description--

Maintenant, ajoute une autre propriété à l'objet `tiger` appelée `isEndangered`.

Définis la propriété `isEndangered` à `true`.

# --hints--

`tiger` doit avoir une propriété `isEndangered`.

```js
assert.property(tiger, "isEndangered");
```

`tiger.isEndangered` doit être égal à `true`.

```js
assert.propertyVal(tiger, "isEndangered", true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
--fcc-editable-region--

--fcc-editable-region--
};
```
