---
id: 699a00de9f564bb0effaa14d
title: Étape 7
challengeType: 1
dashedName: step-7
---

# --description--

Dans cette étape, tu vas créer une fonction pour accéder à la propriété `age` d'un objet.

Voici un exemple :

```js
const cat = {
  age: 3
};

const getAnimalAge = (pet) => {
  return pet.age; // access age using dot notation
}

console.log(getAnimalAge(cat)); // 3
```

Dans cet exemple, la fonction prend un objet en paramètre et retourne sa propriété `age`.

Maintenant, crée une fonction appelée `getAge`.

La fonction doit prendre un paramètre nommé `animal`.

Retourne la propriété `age` de l'`animal` en utilisant la notation par point.

Après avoir créé la fonction, utilise `console.log` pour appeler `getAge(tiger)` afin de voir le résultat dans la console.

# --hints--

Tu dois créer une fonction nommée `getAge`.

```js
assert.isFunction(getAge);
```

La fonction `getAge` doit avoir un seul paramètre, `animal`.

```js
const regex = __helpers.functionRegex('getAge', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

`getAge` doit accéder à la propriété `age` en utilisant la notation par point.

```js
assert.match(code, /animal\.age/);
```

Tu dois afficher `getAge(tiger)` dans la console.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*getAge\s*\(\s*tiger\s*\)\s*\)/);
```

Appeler `getAge(tiger)` doit retourner `5`.

```js
assert.strictEqual(getAge(tiger), 5);
```

`getAge` doit utiliser le paramètre de la fonction et fonctionner avec n'importe quel objet.

```js
const lion = { age: 8 };
assert.strictEqual(getAge(lion), 8);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

--fcc-editable-region--

--fcc-editable-region--
```
