---
id: 69ba8da6b296391d913096a0
title: Étape 5
challengeType: 1
dashedName: step-5
---

# --description--

Maintenant, crée un deuxième objet appelé `elephant`.

Ajoute les propriétés suivantes :

- `species` avec la valeur `"Elephant"`
- `age` avec la valeur `10`
- `isEndangered` avec la valeur `true`

# --hints--

Tu dois créer une variable nommée `elephant`.

```js
assert.exists(elephant);
```

`elephant` doit avoir une propriété `species` égale à `"Elephant"`.

```js
assert.propertyVal(elephant, "species", "Elephant");
```

`elephant` doit avoir une propriété `age` égale à `10`.

```js
assert.propertyVal(elephant, "age", 10);
```

`elephant` doit avoir une propriété `isEndangered` égale à `true`.

```js
assert.propertyVal(elephant, "isEndangered", true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

--fcc-editable-region--

--fcc-editable-region--
```
