---
id: a103376db3ba46b2d50db289
title: Implémente un convertisseur en spinal case
challengeType: 26
dashedName: implement-a-spinal-case-converter
---

# --description--

Le spinal case est un format de chaîne où tous les mots sont en minuscules et séparés par des tirets. `"this-is-spinal-tap"` est un exemple de chaîne en spinal case.

Dans ce lab, tu vas créer une fonction qui convertit une chaîne donnée en spinal case.

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois créer une fonction nommée `spinalCase`.
2. La fonction `spinalCase` doit prendre un seul argument, une chaîne.
3. La fonction `spinalCase` doit renvoyer la chaîne au format spinal case. Par exemple, si l'argument est `ProductLanding page`, la fonction doit renvoyer `product-landing-page`.

# --hints--

Tu dois créer une fonction nommée `spinalCase`.

```js
assert.isFunction(spinalCase);
```

`spinalCase` doit prendre un seul argument.

```js
assert.lengthOf(spinalCase, 1);
```

`spinalCase("This Is Spinal Tap")` doit renvoyer la chaîne `this-is-spinal-tap`.

```js
assert.deepEqual(spinalCase('This Is Spinal Tap'), 'this-is-spinal-tap');
```

`spinalCase("thisIsSpinalTap")` doit renvoyer la chaîne `this-is-spinal-tap`.

```js
assert.strictEqual(spinalCase('thisIsSpinalTap'), 'this-is-spinal-tap');
```

`spinalCase("The_Andy_Griffith_Show")` doit renvoyer la chaîne `the-andy-griffith-show`.

```js
assert.strictEqual(
  spinalCase('The_Andy_Griffith_Show'),
  'the-andy-griffith-show'
);
```

`spinalCase("Teletubbies say Eh-oh")` doit renvoyer la chaîne `teletubbies-say-eh-oh`.

```js
assert.strictEqual(
  spinalCase('Teletubbies say Eh-oh'),
  'teletubbies-say-eh-oh'
);
```

`spinalCase("AllThe-small Things")` doit renvoyer la chaîne `all-the-small-things`.

```js
assert.strictEqual(spinalCase('AllThe-small Things'), 'all-the-small-things');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function spinalCase(str) {
  str = str.replace(/([a-z](?=[A-Z]))/g, '$1 ');
  return str.toLowerCase().replace(/\ |\_/g, '-');
}
```
