---
id: 66adc82170acb1464bd348eb
title: Étape 13
challengeType: 1
dashedName: step-13
---

# --description--

Les derniers messages du bot porteront sur la matière préférée du bot.

Commence par créer une nouvelle variable appelée `favoriteSubject` et affecte-lui la chaîne `"Computer Science"`.

# --hints--

Tu devrais avoir une variable appelée `favoriteSubject`.

```js
assert.isNotNull(favoriteSubject);
```

Ta variable `favoriteSubject` devrait être une chaîne.

```js
assert.isString(favoriteSubject);
```

Tu devrais affecter la chaîne `"Computer Science"` à la variable `favoriteSubject`.

```js
assert.strictEqual(favoriteSubject, "Computer Science");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

const newNicknameGreeting = "I love my nickname but I wish people would call me " + bot + ".";
console.log(newNicknameGreeting);

--fcc-editable-region--

--fcc-editable-region--
```
