---
id: 66ad984e92ce213a66f7e51c
title: Étape 4
challengeType: 1
dashedName: step-4
---

# --description--

Maintenant, il est temps d'affecter une valeur à ta variable `bot`.

Dans les leçons précédentes, tu as appris à affecter des valeurs aux variables comme ceci :

```js
variableName = "Here is the value";
```

Souviens-toi que ce qui se trouve à droite de l'opérateur d'affectation `=` est la valeur que tu affectes à la variable située à gauche.

Affecte la chaîne `"teacherBot"` à la variable `bot`.

# --hints--

Tu devrais affecter une chaîne à ta variable `bot`.

```js
assert.isString(bot);
```

Ta variable `bot` devrait recevoir la chaîne `"teacherBot"`.

```js
assert.equal(bot, "teacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
--fcc-editable-region--

--fcc-editable-region--
```
