---
id: 66adc6046acc18453decf577
title: Étape 12
challengeType: 1
dashedName: step-12
---

# --description--

Pour voir le nouveau surnom du bot, tu devras afficher un nouveau message dans la console.

Crée une nouvelle variable appelée `newNicknameGreeting`.

Ensuite, utilise la concaténation de chaînes pour assembler la chaîne `"I love my nickname but I wish people would call me "` avec la variable `bot` suivie d'un point.

Affecte le résultat à la variable `newNicknameGreeting`.

Ensuite, affiche la valeur de `newNicknameGreeting` dans la console.

# --hints--

Tu devrais avoir une variable appelée `newNicknameGreeting`.

```js
assert.isNotNull(newNicknameGreeting);
```

Ta variable `newNicknameGreeting` devrait être une chaîne.

```js
assert.isString(newNicknameGreeting);
```

Tu devrais utiliser la concaténation de chaînes pour assembler la chaîne `"I love my nickname but I wish people would call me "` avec la variable `bot` suivie d'un point.

```js
assert.strictEqual(newNicknameGreeting, "I love my nickname but I wish people would call me awesomeTeacherBot.");
```

Tu devrais avoir une instruction `console` dans ton code.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 7);
```

Tu devrais afficher la valeur de `newNicknameGreeting` dans la console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*newNicknameGreeting\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
