---
id: 66adcf383276814776aba3ca
title: Étape 14
challengeType: 1
dashedName: step-14
---

# --description--

Ensuite, crée une variable appelée `favoriteSubjectSentence`.

Utilise la concaténation de chaînes pour assembler la chaîne `"My favorite subject is "` avec la variable `favoriteSubject` suivie d'un point.

Affecte le résultat à la variable `favoriteSubjectSentence`.

Ensuite, affiche la valeur de `favoriteSubjectSentence` dans la console.

# --hints--

Tu devrais avoir une variable appelée `favoriteSubjectSentence`.

```js
assert.isNotNull(favoriteSubjectSentence);
```

Ta variable `favoriteSubjectSentence` devrait être une chaîne.

```js
assert.isString(favoriteSubjectSentence);
```

Tu devrais utiliser la concaténation de chaînes pour assembler la chaîne `"My favorite subject is "` avec la variable `favoriteSubject` suivie d'un point. Fais attention aux espaces.

```js
assert.strictEqual(favoriteSubjectSentence, "My favorite subject is Computer Science.");
```

Tu devrais avoir une instruction `console` dans ton code.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 8);
```

Tu devrais afficher la valeur de `favoriteSubjectSentence` dans la console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*favoriteSubjectSentence\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

const newNicknameGreeting = "I love my nickname but I wish people would call me " + bot + ".";
console.log(newNicknameGreeting);

const favoriteSubject = "Computer Science";

--fcc-editable-region--

--fcc-editable-region--
```
