---
id: 66ad9c633fa26d3c1475eae3
title: Étape 6
challengeType: 1
dashedName: step-6
---

# --description--

Maintenant, il est temps d'ajouter un autre message du bot.

Ajoute une autre instruction `console` au code qui affiche le message `"Allow me to introduce myself."`.

# --hints--

Tu devrais avoir une troisième instruction `console.log()` dans ton code.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 3);
```

Ton instruction `console` devrait contenir le message `"Allow me to introduce myself."`. N'oublie pas les guillemets.

```js
assert.match(__helpers.removeJSComments(code), /("|')Allow\s+me\s+to\s+introduce\s+myself.\1/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

--fcc-editable-region--

--fcc-editable-region--
```
