---
id: 66ad8294a0ad902f1b31b612
title: Étape 1
challengeType: 1
dashedName: step-1
---

# --description--

Dans ce workshop, tu apprendras à manipuler les fondamentaux de JavaScript en construisant un bot de salutation.

Dans cette première étape, tu vas afficher un message dans la console depuis le bot de salutation.

Souviens-toi que tu as appris `console.log()` et les chaînes dans les leçons précédentes.

Voici un rappel de comment utiliser `console.log()` avec les chaînes :

```js
console.log("Hello, World!");
```

Ajoute une instruction `console.log()` qui affiche la chaîne `"Hi there!"` dans la console. N'oublie pas les guillemets autour du message !


# --hints--

Tu devrais avoir une instruction `console.log()` dans ton code.

```js
assert.match(__helpers.removeJSComments(code), /console\.log(.*)/);
```

Ton instruction `console` devrait contenir le message `"Hi there!"`. N'oublie pas les guillemets.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\((['"])Hi\s+there!\1\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
