---
id: 66adb844118ba74107ce771f
title: Étape 8
challengeType: 1
dashedName: step-8
---

# --description--

Le prochain message du bot concernera l'emplacement du bot.

Crée une variable appelée `botLocationSentence`.

Ensuite, utilise la concaténation de chaînes avec l'opérateur `+` pour assembler la chaîne `"I live in "` avec la variable `botLocation` suivie d'un point (`.`).

Affecte cette valeur à la variable `botLocationSentence`.

Ensuite, affiche la valeur de `botLocationSentence` dans la console.

# --hints--

Tu devrais avoir une variable `botLocationSentence`.

```js
assert.isNotNull(botLocationSentence);
```

Ta variable `botLocationSentence` devrait contenir une valeur de type chaîne.

```js
assert.isString(botLocationSentence);
```

Tu devrais utiliser la concaténation de chaînes avec l'opérateur `+` pour assembler la chaîne `"I live in "` avec la variable `botLocation` suivie d'un point (`.`). Fais attention aux espaces.

```js
assert.equal(botLocationSentence, "I live in the universe.");
```

Tu devrais avoir une instruction `console` dans ton code.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 5);
```

Ton instruction `console` devrait afficher la variable `botLocationSentence`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*botLocationSentence\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

--fcc-editable-region--

--fcc-editable-region--
```
