---
id: 66ad8d150264db3926eccfeb
title: Étape 3
challengeType: 1
dashedName: step-3
---

# --description--

Dans les leçons précédentes, tu as appris le mot-clé `let` et comment déclarer des variables en JavaScript.

Voici un rappel de comment déclarer une variable à l'aide du mot-clé `let` :

```js
let greeting;
```

Utilise le mot-clé `let` pour déclarer une variable appelée `bot`.

**REMARQUE** : Tu utilises `let` ici parce que, plus loin dans le workshop, tu changeras la valeur de la variable `bot`.

# --hints--

Tu devrais avoir le mot-clé `let` dans ton code.

```js
assert.match(__helpers.removeJSComments(code), /let/g);
```

Tu devrais déclarer une variable appelée `bot`.

```js
assert.isNotNull(bot);
```

Tu devrais utiliser le mot-clé `let` pour déclarer la variable.

```js
assert.match(__helpers.removeJSComments(code), /let\s+bot;?/g);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
--fcc-editable-region--

--fcc-editable-region--
```
