---
id: 66adc42868cab843ccee87d9
title: Étape 10
challengeType: 1
dashedName: step-10
---

# --description--

Maintenant, il est temps de voir la nouvelle valeur de `bot`.

Commence par créer une nouvelle variable appelée `nicknameIntroduction`.

Utilise la concaténation de chaînes pour assembler la chaîne `"My nickname is "` avec la variable `bot` suivie d'un point (`.`).

Affecte la chaîne obtenue à la variable `nicknameIntroduction`.

Ensuite, affiche la valeur de `nicknameIntroduction` dans la console.

# --hints--

Tu devrais avoir une variable appelée `nicknameIntroduction`.

```js
assert.isNotNull(nicknameIntroduction);
```

Ta variable `nicknameIntroduction` devrait être une chaîne.

```js
assert.isString(nicknameIntroduction);
```

Tu devrais utiliser la concaténation de chaînes pour assembler la chaîne `"My nickname is "` avec la variable `bot` suivie d'un point (`.`). Fais attention aux espaces.

```js
assert.strictEqual(nicknameIntroduction, "My nickname is professorBot.");
```

Tu devrais avoir une instruction `console` dans ton code.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 6);
```

Tu devrais afficher la valeur de `nicknameIntroduction` dans la console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*nicknameIntroduction\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

--fcc-editable-region--

--fcc-editable-region--
```
