---
id: 66ad8ab945d266383f318cbf
title: Étape 2
challengeType: 1
dashedName: step-2
---

# --description--

Maintenant, tu devrais voir le premier message du bot dans la console.

Il est temps d'ajouter un deuxième message du bot.

Ajoute une autre instruction `console.log` pour afficher le message `"I am excited to talk to you."` dans la console.

# --hints--

Tu devrais avoir une deuxième instruction `console.log()` dans ton code.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 2);
```

Ton instruction `console` devrait contenir le message `"I am excited to talk to you."`. N'oublie pas les guillemets.

```js
assert.match(__helpers.removeJSComments(code), /("|')I\s+am\s+excited\s+to\s+talk\s+to\s+you.\1/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
--fcc-editable-region--

--fcc-editable-region--
```
