---
id: 66adbc8a9793d64250f7e609
title: Étape 9
challengeType: 1
dashedName: step-9
---

# --description--

Dans les leçons précédentes, tu as appris à réaffecter des valeurs aux variables comme ceci :

```js
let name = "John";
console.log(name); // "John"

name = "Jane";
console.log(name); // "Jane"
```

En utilisant la réaffectation, affecte la chaîne `"professorBot"` à la variable `bot`.

# --hints--

Tu devrais affecter la chaîne `"professorBot"` à la variable `bot`.

```js
assert.strictEqual(bot, "professorBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

--fcc-editable-region--

--fcc-editable-region--
```
