---
id: 66ad998285df023b23dacdd3
title: Étape 5
challengeType: 1
dashedName: step-5
---

# --description--

Maintenant, il est temps d'initialiser la variable `botLocation`.

Quand tu as besoin de déclarer des variables composées de plusieurs mots, tu peux utiliser la convention de nommage <dfn>camelCase</dfn>.

Avec le `camelCase`, le premier mot est entièrement en minuscules et la première lettre de chaque mot suivant est en majuscule.

Voici un exemple :

```js
let thisIsCamelCase = "Alice the camel is actually a horse.";
```

Déclare et affecte la chaîne `"the universe"` à la variable `botLocation` sur la même ligne à l'aide du mot-clé `let`.

# --hints--

Utilise le mot-clé `let` pour déclarer une variable.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/let/g), 2);
```

Ta variable devrait s'appeler `botLocation`.

```js
assert.isNotNull(botLocation);
```

Tu devrais affecter une chaîne à ta variable `botLocation`.

```js
assert.isString(botLocation);
```

Ta variable `botLocation` devrait recevoir la chaîne `"the universe"`.

```js
assert.equal(botLocation, "the universe");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
