---
id: 66ada3f46945763dd97f43f8
title: Étape 7
challengeType: 1
dashedName: step-7
---

# --description--

Plus tôt, tu as créé les variables `bot` et `botLocation`. Maintenant, tu vas les utiliser pour afficher de nouveaux messages dans la console.

Dans les leçons précédentes, tu as appris à travailler avec la concaténation de chaînes en utilisant l'opérateur `+` pour assembler des chaînes comme ceci :

```js
let firstName = "John";

console.log("Hello, my name is " + firstName + ".");
// result: "Hello, my name is John."
```

Souviens-toi que tu dois faire attention aux espaces quand tu concatènes des chaînes avec des variables.

Crée une variable appelée `botIntroduction`.

Ensuite, utilise la concaténation de chaînes avec l'opérateur `+` pour assembler la chaîne `"My name is "` suivie de la variable `bot` suivie d'un point (`.`).

Affecte cette valeur à la variable `botIntroduction`.

Ensuite, affiche la variable `botIntroduction` dans la console.

# --hints--

Tu devrais avoir une variable appelée `botIntroduction`.

```js
assert.isNotNull(botIntroduction);
```

Ta variable `botIntroduction` devrait contenir une valeur de type chaîne.

```js
assert.isString(botIntroduction);
```

Tu devrais utiliser la concaténation de chaînes avec l'opérateur `+` pour assembler la chaîne `"My name is "` avec la variable `bot` suivie d'un point (`.`). Fais attention aux espaces. Affecte cette chaîne concaténée à la variable `botIntroduction`.

```js
assert.equal(botIntroduction, "My name is teacherBot.");
```

Tu devrais avoir une quatrième instruction `console.log()` dans ton code.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 4);
```

Ton instruction `console` devrait afficher la variable `botIntroduction`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*botIntroduction\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

--fcc-editable-region--

--fcc-editable-region--
```
