---
id: 68c343692285523aef0fc6bf
title: Étape 5
challengeType: 0
dashedName: step-5
---

# --description--

L’attribut `scope` sert à indiquer si une cellule d’en-tête est l’en-tête d’une ligne, d’une colonne, ou d’un groupe de lignes ou de colonnes. Voici un exemple :

```html
<th scope="col">Example Header</th>
```

Cela aide les lecteurs d’écran à comprendre la relation entre les cellules d’en-tête et les cellules de données.

Pour tous les éléments `th`, ajoute un attribut `scope` avec la valeur `col`.

# --hints--

Chaque élément `th` devrait avoir un attribut `scope` avec la valeur `"col"`.

```js
const tableHeaders = document.querySelectorAll("th");
tableHeaders.forEach(header => {
  assert.equal(header?.getAttribute("scope"), "col");
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
      --fcc-editable-region--
        <th>Time</th>
        <th>Track A</th>
        <th>Track B</th>
        <th>Track C</th>
      --fcc-editable-region--
      </tr>
    </thead>
  </table>
</body>
</html>
```
