---
id: 68c34ace730ca43e5fa7b8d6
title: Étape 9
challengeType: 0
dashedName: step-9
---

# --description--

Il est maintenant temps d’ajouter une autre ligne au tableau.

Commence par ajouter un autre élément `tr`. Dans cet élément `tr`, ajoute un élément `th` avec un attribut `scope` défini sur `"row"` et le contenu textuel `10:00 AM`.

Ensuite, ajoute trois éléments `td` avec les contenus textuels suivants :

- `Accessibility Deep Dive`
- `CSS for Beginners`
- `Inclusive Design Principles`

# --hints--

Tu devrais avoir un deuxième élément `tr` dans l’élément `tbody`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 2);
```

Tu devrais avoir un élément `th` dans le deuxième élément `tr`.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[1].querySelector("th"));
```

Ton élément `th` devrait avoir le contenu textuel `10:00 AM`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelector("th")?.textContent.trim(), "10:00 AM");
```

Ton élément `th` devrait avoir un attribut `scope` défini sur `"row"`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelector("th")?.getAttribute("scope"), "row");
```

Ton deuxième élément `tr` devrait avoir trois éléments `td`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr")[1].querySelectorAll("td"), 3);
```

Ton premier élément `td` devrait avoir le contenu textuel `Accessibility Deep Dive`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[0]?.textContent.trim(), "Accessibility Deep Dive");
```

Ton deuxième élément `td` devrait avoir le contenu textuel `CSS for Beginners`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[1]?.textContent.trim(), "CSS for Beginners");
```   

Ton troisième élément `td` devrait avoir le contenu textuel `Inclusive Design Principles`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[2]?.textContent.trim(), "Inclusive Design Principles");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
