---
id: 68c34329d430583a6cd5f7f1
title: Étape 4
challengeType: 0
dashedName: step-4
---

# --description--

Dans ton élément `tr`, ajoute quatre éléments `th`. Le premier élément `th` devrait avoir le texte `Time`, le deuxième le texte `Track A`, le troisième le texte `Track B`, et le quatrième le texte `Track C`.

# --hints--

Ton élément `tr` devrait avoir quatre éléments `th`.

```js
assert.lengthOf(document.querySelectorAll("tr > th"), 4);
```

Ton premier élément `th` devrait avoir le texte `Time`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(1)")?.textContent.trim(), "Time");
```

Ton deuxième élément `th` devrait avoir le texte `Track A`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(2)")?.textContent.trim(), "Track A");
```

Ton troisième élément `th` devrait avoir le texte `Track B`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(3)")?.textContent.trim(), "Track B");
```

Ton quatrième élément `th` devrait avoir le texte `Track C`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(4)")?.textContent.trim(), "Track C");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

--fcc-editable-region--
    <thead>
      <tr>
        
      </tr>
    </thead>
--fcc-editable-region--
  </table>
</body>
</html>
```
