---
id: 68c341332d98fe38fae733af
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

Il est maintenant temps de commencer à construire le tableau.

Ajoute un élément `table` sous ton élément `h1`. Dans l’élément `table`, ajoute un élément `caption` avec le texte `Schedule by Track and Time`.

# --hints--

Tu devrais avoir un élément `table`.

```js
assert.exists(document.querySelector("table"));
```

Ton élément `table` devrait avoir un élément `caption`.

```js
assert.exists(document.querySelector("table caption"));
```

Ton élément `caption` devrait avoir le texte `Schedule by Track and Time`.

```js
assert.equal(document.querySelector("table caption")?.textContent.trim(), "Schedule by Track and Time");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

--fcc-editable-region--
    
--fcc-editable-region--
</body>
</html>
```
