---
id: 68c34d57ed874e4004f7b5d4
title: Étape 11
challengeType: 0
dashedName: step-11
---

# --description--

Pour l’instant, l’élément `td` avec le contenu textuel `Break` ne s’étend que sur une seule colonne. Mais ce serait mieux s’il s’étendait sur les trois colonnes.

Comme tu l’as vu dans des ateliers et leçons précédents, tu peux utiliser l’attribut `colspan` pour qu’une cellule de tableau s’étende sur plusieurs colonnes.

```html
<tr>
  <td colspan="3">Total Points</td>
</tr>
```

Ajoute un attribut `colspan` à l’élément `td` et donne-lui la valeur `3`.

# --hints--

Ton élément `td` avec le contenu textuel `Break` devrait avoir un attribut `colspan` défini sur `3`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("td")?.getAttribute("colspan"), "3");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        --fcc-editable-region--
        <td>Break</td>
        --fcc-editable-region--
      </tr>
    </tbody>
    
  </table>
</body>
</html>
```
