---
id: 68c33e502cd97f2fd52cb63b
title: Étape 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Dans cet atelier, tu vas apprendre à créer des éléments de tableau accessibles en construisant le programme d’une conférence tech.

Commence par ajouter un élément `h1` avec le texte `Tech Conference 2025 Schedule`.

# --hints--

Tu devrais avoir un élément `h1`.

```js
const h1 = document.querySelector("h1");
assert.exists(h1);
```

Ton élément `h1` devrait avoir le texte `Tech Conference 2025 Schedule`.

```js
const h1 = document.querySelector("h1");
assert.strictEqual(h1?.textContent.trim(), "Tech Conference 2025 Schedule");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
    --fcc-editable-region--
    
    --fcc-editable-region--
</body>
</html>
```
