---
id: 68c3420cb0176439a95af1b8
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

L’étape suivante consiste à ajouter l’en-tête du tableau.

Ajoute un élément `thead` sous l’élément `caption`. Dans l’élément `thead`, ajoute un élément `tr`.

# --hints--

Tu devrais avoir un élément `thead`.

```js
assert.exists(document.querySelector("table thead"));
```

Tu devrais avoir un élément `tr` dans l’élément `thead`.

```js
assert.exists(document.querySelector("table thead tr"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    --fcc-editable-region--
        
    --fcc-editable-region--
  </table>
</body>
</html>
```
