---
id: 68c3482278a14e3cf225e897
title: Étape 7
challengeType: 0
dashedName: step-7
---

# --description--

Dans ton élément `tr`, ajoute un élément `th` avec le texte `9:00 AM`. Puis, sous cet élément `th`, ajoute trois éléments `td` avec les textes `Keynote: Tech Future`, `Intro to Web Dev` et `UX for All`.

# --hints--

Tu devrais avoir un élément `th` dans ton élément `tr`.

```js
assert.exists(document.querySelector("tbody tr th"));
```

Ton élément `th` devrait avoir le contenu textuel `9:00 AM`.

```js
assert.equal(document.querySelector("tbody tr th")?.textContent.trim(), "9:00 AM");
```

Tu devrais avoir un total de trois éléments `td` sous ton élément `th`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.lengthOf(tdElements, 3);
```

Ton premier élément `td` devrait avoir le contenu textuel `Keynote: Tech Future`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[0]?.textContent.trim(), "Keynote: Tech Future");
```

Ton deuxième élément `td` devrait avoir le contenu textuel `Intro to Web Dev`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[1]?.textContent.trim(), "Intro to Web Dev");
``` 

Ton troisième élément `td` devrait avoir le contenu textuel `UX for All`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[2]?.textContent.trim(), "UX for All");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </tr>
    </tbody>
    
  </table>
</body>
</html>
```
