---
id: 68c346d03709343bff132707
title: Étape 6
challengeType: 0
dashedName: step-6
---

# --description--

Dans les prochaines étapes, tu vas construire le corps du tableau.

Commence par ajouter un élément `tbody` sous l’élément `thead`. Puis, dans l’élément `tbody`, ajoute un élément `tr`.

# --hints--

Tu devrais avoir un élément `tbody` sous l’élément `thead`.

```js
assert.exists(document.querySelector("tbody"));
```

Tu devrais avoir un élément `tr` dans l’élément `tbody`.

```js
assert.exists(document.querySelector("tbody tr"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>
    
    --fcc-editable-region--
    
    --fcc-editable-region--
  </table>
</body>
</html>
```
