---
id: 68c34c608cf9e53f5d36b1d5
title: Étape 10
challengeType: 0
dashedName: step-10
---

# --description--

Ensuite, ajoute une troisième ligne au tableau. Commence par ajouter un autre élément `tr`. Dans cet élément `tr`, ajoute un élément `th` avec un attribut `scope` défini sur `"row"` et le contenu textuel `11:00 AM`. Puis, sous cet élément `th`, ajoute un élément `td` avec le contenu textuel `Break`.

# --hints--

Tu devrais avoir un troisième élément `tr` dans l’élément `tbody`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 3);
```

Tu devrais avoir un élément `th` dans le troisième élément `tr`.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[2].querySelector("th"));
```

Ton élément `th` devrait avoir le contenu textuel `11:00 AM`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("th")?.textContent.trim(), "11:00 AM");
```

Ton élément `th` devrait avoir un attribut `scope` défini sur `"row"`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("th")?.getAttribute("scope"), "row");
```

Tu devrais avoir un élément `td` dans le troisième élément `tr`.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[2].querySelector("td"));
```

Ton élément `td` devrait avoir le contenu textuel `Break`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("td")?.textContent.trim(), "Break");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

--fcc-editable-region--
      
--fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
