---
id: 579e2a2c335b9d72dd32e05c
title: Implémente l'algorithme slice and splice
challengeType: 26
dashedName: implement-the-slice-and-splice-algorithm
---

# --description--

Dans ce lab, tu dois créer un algorithme pour fusionner deux tableaux.

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Crée une fonction `frankenSplice` qui accepte deux tableaux et un indice.
1. Copie chaque élément du premier tableau dans le second tableau, dans l'ordre, en commençant à l'indice indiqué, et renvoie le tableau résultant.
1. Les tableaux d'entrée doivent rester inchangés après l'exécution de la fonction.

# --hints--

`frankenSplice([1, 2, 3], [4, 5], 1)` doit renvoyer `[4, 1, 2, 3, 5]`.

```js
assert.deepEqual(frankenSplice([1, 2, 3], [4, 5], 1), [4, 1, 2, 3, 5]);
```

`frankenSplice([1, 2], ["a", "b"], 1)` doit renvoyer `["a", 1, 2, "b"]`.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
assert.deepEqual(frankenSplice(testArr1, testArr2, 1), ['a', 1, 2, 'b']);
```

`frankenSplice(["claw", "tentacle"], ["head", "shoulders", "knees", "toes"], 2)` doit renvoyer `["head", "shoulders", "claw", "tentacle", "knees", "toes"]`.

```js
assert.deepEqual(
  frankenSplice(
    ['claw', 'tentacle'],
    ['head', 'shoulders', 'knees', 'toes'],
    2
  ),
  ['head', 'shoulders', 'claw', 'tentacle', 'knees', 'toes']
);
```

Tous les éléments du premier tableau doivent être ajoutés au second tableau dans leur ordre d'origine. `frankenSplice([1, 2, 3, 4], [], 0)` doit renvoyer `[1, 2, 3, 4]`.

```js
assert.deepEqual(frankenSplice([1, 2, 3, 4], [], 0), [1, 2, 3, 4]);
```

Le premier tableau doit rester inchangé après l'exécution de la fonction.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr1, [1, 2]);
```

Le second tableau doit rester inchangé après l'exécution de la fonction.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr2, ['a', 'b']);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function frankenSplice(arr1, arr2, n) {
  let result = arr2.slice();
  for (let i = 0; i < arr1.length; i++) {
    result.splice(n + i, 0, arr1[i]);
  }
  return result;
}
```
