---
id: 69899123889dcf5b19be1847
title: Implémenter l'algorithme de parcours en profondeur
challengeType: 26
dashedName: implement-depth-first-search-algorithm
---

# --description--

Dans ce labo, tu vas implémenter un algorithme de parcours de graphe appelé <dfn>parcours en profondeur</dfn> (DFS).

Alors qu'un <dfn>parcours en largeur</dfn> (BFS) explore les voisins niveau par niveau, le <dfn>parcours en profondeur</dfn> plonge aussi loin que possible le long d'un seul chemin d'arêtes avant de revenir en arrière.

Une fois que l'algorithme atteint la fin d'un chemin (un nœud sans voisins non visités), il revient au nœud le plus récent qui a encore des arêtes inexplorées et continue la recherche depuis là.

Pour implémenter cet algorithme de façon itérative, tu voudras utiliser une **pile**. En JavaScript, un `Array` standard sert parfaitement de pile avec les méthodes `.push()` et `.pop()`, suivant le principe **Last-In-First-Out (LIFO)**. Cela garantit que le voisin ajouté le plus récemment à la pile est le prochain à être exploré.

Une sortie simple de cet algorithme est une liste de nœuds atteignables depuis un nœud donné. Tu voudras donc aussi suivre les nœuds que tu visites.

**Objectif** : Remplis les user stories ci-dessous et fais passer tous les tests pour terminer le labo.

**User Stories :**

1. Tu dois définir une fonction nommée `dfs`.
1. La fonction `dfs` doit prendre deux arguments :

    - `graph` : Une matrice d'adjacence (un tableau de tableaux).
    - `root` : Un label de nœud numérique (l'indice de départ) qui est la valeur numérique du nœud entre `0` et `n - 1`, où `n` est le nombre total de nœuds dans le graphe.

1. La fonction doit renvoyer un tableau contenant tous les labels de nœuds atteignables depuis le nœud de départ.

# --hints--

Tu dois avoir une fonction nommée `dfs` qui prend deux arguments.

```js
assert.isFunction(dfs);
assert.lengthOf(dfs, 2);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1)` doit renvoyer une liste avec `1`, `2`, `3`, et `0`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1), [1, 2, 3, 0]);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3)` doit renvoyer une liste avec `1`, `2`, `3`, et `0`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3), [1, 2, 3, 0]);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3)` doit renvoyer `[3]`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3), [3]);
```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3)` doit renvoyer une liste avec `3` et `2`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3), [3, 2]);

```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0)` doit renvoyer une liste avec `0` et `1`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0), [0, 1]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dfs(graph, root) {
  let stack = [root];
  let visited = [];

  while (stack.length > 0) {
    let current = stack.pop();

    if (!visited.includes(current)) {
      visited.push(current);
      let neighbors = graph[current];

      for (let i = 0; i < neighbors.length; i++) {
        if (neighbors[i] === 1) {
          stack.push(i);
        }
      }
    }
  }
  return visited;
}
```
