---
id: 68e4c81577243e4db1676491
title: Étape 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Dans cet atelier, tu vas t'entraîner à utiliser le HTML sémantique en construisant une page web qui contient quelques conseils de Quincy Larson pour décrocher un emploi de développeur. Le boilerplate HTML de base a été préparé pour toi.

Commence par créer un élément `h1` avec le texte `Quincy's Tips for Getting a Developer Job`.

# --hints--

Tu devrais avoir un élément `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Ton élément `h1` devrait avoir le texte `Quincy's Tips for Getting a Developer Job`.

```js
assert.equal(document.querySelector('h1')?.textContent.trim(), "Quincy's Tips for Getting a Developer Job")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
