---
id: 68e65ec5348e40a739903e6d
title: Étape 2
challengeType: 0
dashedName: step-2
---

# --description--

Maintenant, crée un élément paragraphe sous l'élément `h1`. À l'intérieur de ce paragraphe, ajoute le texte `Learning to code is hard, but as Quincy Larson says, You can become a developer.`

# --hints--

Tu devrais avoir un élément `p`.

```js
assert.exists(document.querySelector('p'));
```

Ton paragraphe devrait avoir le texte `Learning to code is hard, but as Quincy Larson says, You can become a developer.`. Vérifie ton orthographe.

```js
assert.equal(document.querySelector('p')?.textContent.trim(), "Learning to code is hard, but as Quincy Larson says, You can become a developer.");
```

Ton paragraphe devrait être sous l'élément `h1`.

```js
assert.exists(document.querySelector('h1 + p'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
