---
id: 68e90335bd53a901c41dfc13
title: Étape 3
challengeType: 0
dashedName: step-3
---

# --description--

Dans une leçon précédente, tu as appris que l'élément `q` est utilisé pour distinguer du texte cité du contenu environnant.

Voici un exemple :

```html
<p>
  Nancy said, <q>Learning is fun!</q>
</p>
```

La plupart des navigateurs modernes ajoutent automatiquement des guillemets autour d'une citation en ligne quand tu utilises l'élément `q`.

À l'intérieur du paragraphe, entoure `You can become a developer.` d'un élément de citation en ligne, en gardant le reste du paragraphe inchangé.

# --hints--

Tu devrais avoir un élément `q`.

```js
assert.exists(document.querySelector('q'));
```

Ton élément `q` devrait être imbriqué dans le paragraphe qui est sous l'élément `h1`.

```js
assert.exists(document.querySelector('h1 + p > q'));
```

Ton élément `q` devrait avoir le texte `You can become a developer.`

```js
assert.equal(document.querySelector('h1 + p > q')?.textContent.trim(), 'You can become a developer.');
```

À l'intérieur de ton élément `q`, le texte ne doit pas avoir d'espace au début ou à la fin (pas d'espaces, de tabulations ou de retours à la ligne avant le premier caractère ou après le dernier).

```js
assert.equal(document.querySelector('h1 + p > q')?.textContent, 'You can become a developer.');
```

Tu devrais seulement ajouter des balises `q` autour de la partie appropriée de la phrase sans changer le texte original. L'élément `p` devrait toujours se lire : `Learning to code is hard, but as Quincy Larson says, You can become a developer.`

```js
const pElCleanedText = document.querySelector('h1 + p')?.innerText.trim().replaceAll(/ {2,}/g, ' ');
assert.equal(pElCleanedText, 'Learning to code is hard, but as Quincy Larson says, You can become a developer.');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
--fcc-editable-region--
    <p>Learning to code is hard, but as Quincy Larson says, You can become a developer.</p>
--fcc-editable-region--
  </body>
</html>
```
