---
id: 68e964b84f954c322376e592
title: Étape 4
challengeType: 0
dashedName: step-4
---

# --description--

Si la source d'une citation est un site web, tu peux le citer avec l'attribut `cite`. La valeur de cet attribut doit être une URL valide. Bien que cet attribut ne change pas la présentation de la citation, il est très utile pour donner aux lecteurs d'écran et aux moteurs de recherche plus d'informations sur la citation.

Voici un exemple d'élément de citation en ligne avec un attribut `cite` :

```html
<p>
  Nancy said,
  <q cite="https://example.com">Learning is fun!</q>
</p>
```

Ajoute l'attribut `cite` à l'élément de citation en ligne avec cette URL :

`https://www.freecodecamp.org/news/learn-to-code-book/`

# --hints--

Ton élément `q` devrait avoir un attribut `cite`.

```js
const qEl = document.querySelector('h1 + p > q');
assert.exists(qEl?.getAttribute('cite'));
```

L'attribut `cite` de ton élément `q` devrait avoir la valeur `https://www.freecodecamp.org/news/learn-to-code-book/`.

```js
const qEl = document.querySelector('h1 + p > q');
assert.equal(qEl?.getAttribute('cite'), 'https://www.freecodecamp.org/news/learn-to-code-book/');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
--fcc-editable-region--
      <q>You can become a developer.</q>
--fcc-editable-region--
    </p>
  </body>
</html>
```
