---
id: 68eaa9ef8a0867071f195f59
title: Étape 10
challengeType: 0
dashedName: step-10
---

# --description--

Jusqu'ici, tu as utilisé l'attribut `cite` pour attribuer la source de la citation, mais cet attribut ne montre pas vraiment la source à l'utilisateur.

Si tu veux attribuer la source de manière visible, tu peux ajouter un élément citation, `cite`, à l'extérieur de l'élément de citation en bloc. L'élément citation est un élément HTML que tu peux utiliser pour baliser le titre d'une œuvre créative référencée, comme un livre, un article, une chanson, un film, un site web ou un article de recherche.

Voici un exemple :

```html
<div>
  <blockquote cite="https://www.freecodecamp.org/news/is-college-worth-it/">
    The first thing you should consider about education is this is an economic decision.
  </blockquote>
  <p>&mdash;Quincy Larson, <cite>Is College Still Worth it? Tips from my 20 Years in Adult Education</cite></p>
</div>
```

À l'intérieur de l'élément `p` sous l'élément de citation en bloc, entoure `How to Learn to Code and Get a Developer Job [Full Book]` d'un élément `cite`.

# --hints--

Tu devrais avoir un élément `cite` à l'intérieur de l'élément `p` sous l'élément `blockquote`.

```js
assert.exists(document.querySelector('blockquote + p > cite'));
```

Ton élément `cite` devrait avoir le texte `How to Learn to Code and Get a Developer Job [Full Book]`.

```js
const citeEl = document.querySelector('blockquote + p > cite');
assert.equal(citeEl?.innerText.trim(), 'How to Learn to Code and Get a Developer Job [Full Book]');
```

En ignorant l'élément `cite` que tu viens d'ajouter, le texte de l'élément `p` devrait toujours être `—Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]`.

```js
const pElCleanedText = document.querySelector('blockquote + p')?.innerText.trim().replaceAll(/ {2,}/g, ' ');
assert.equal(pElCleanedText, '—Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?
        </blockquote>
        <p>
      --fcc-editable-region--
          &mdash;Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]
      --fcc-editable-region--
        </p>
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
