---
id: 68ea7e7b916e8270483689b6
title: Étape 7
challengeType: 0
dashedName: step-7
---

# --description--

Sous l'élément `h2`, tu vas ajouter une autre citation de Quincy. Cette fois-ci, la citation ne fera pas partie d'un paragraphe plus large. Au lieu de ça, tout le paragraphe sera une citation. Pour distinguer du texte cité comme celui-ci, tu devrais utiliser l'élément de citation en bloc : `blockquote`. Dans le navigateur, tu verras que le texte est légèrement en retrait.

Voici un exemple d'élément de citation en bloc avec du texte cité :

```html
<blockquote>
  The first thing you should consider about education is this is an economic decision.
</blockquote>
```

Maintenant, à l'intérieur de la première section, ajoute un élément de citation en bloc sous l'élément `h2` avec le texte `Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?`.

# --hints--

À l'intérieur du premier élément `section`, tu devrais avoir un élément `blockquote` sous l'élément `h2`.

```js
assert.exists(document.querySelector('section:first-of-type > h2 + blockquote'));
```

Ton élément `blockquote` devrait avoir le texte `Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?`.

```js
const firstBlockquoteEl = document.querySelector('section:first-of-type > h2 + blockquote');
assert.equal(firstBlockquoteEl?.innerText.trim(), 'Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
