---
id: 66c06fad3475cd92421b9ac2
title: Crée un calculateur d'année bissextile
challengeType: 26
dashedName: build-a-leap-year-calculator
---

# --description--

Une année bissextile est une année divisible par `4`, sauf pour les années divisibles par `100` et non divisibles par `400`. Par exemple, `2000` est une année bissextile, mais `1900` ne l'est pas. De plus, une année bissextile a un jour supplémentaire en février, qui est le 29e jour du mois.

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Définis une fonction appelée `isLeapYear` qui prend un nombre comme argument.
2. En dehors de la fonction, déclare une variable `year` qui stocke la valeur de l'année que tu veux vérifier.
3. À l'intérieur de la fonction, utilise une instruction `if/ else` ou un opérateur ternaire pour vérifier si l'année est bissextile.
4. Pour vérifier si l'année est bissextile, respecte les conditions suivantes :

   - Si l'année est divisible par `4`, alors c'est une année bissextile.
   - Sauf si l'année est aussi divisible par `100`, alors ce n'est pas une année bissextile.
   - Sauf si l'année est aussi divisible par `400`, alors c'est une année bissextile.

5. Si l'année est bissextile, retourne `[year] is a leap year.`. Sinon, retourne `[year] is not a leap year.`. Tu remplaceras `[year]` par le paramètre défini dans la fonction `isLeapYear`.
6. Tu dois appeler la fonction `isLeapYear` avec `year` comme argument et affecter le résultat à une variable nommée `result`.
7. Tu dois afficher la variable `result` dans la console en utilisant `console.log()`.

# --hints--

Tu dois définir une fonction nommée `isLeapYear`.

```js
assert.isFunction(isLeapYear);
```

La fonction `isLeapYear` doit prendre un nombre comme argument.

```js
assert.match(isLeapYear.toString(), /\s*isLeapYear\(\s*\w+\s*\)/);
```

Tu dois déclarer une variable `year` et lui affecter une valeur pour vérifier si c'est une année bissextile.

```js
assert.isDefined(year);
assert.isNumber(year);
assert.isAtLeast(year, 0);
assert.isTrue(Number.isInteger(year));
```

La variable `year` ne doit pas être vide.

```js
assert.isNotNull(year);
```

Avec `2024` comme valeur de la variable `year`, le `result` doit être `2024 is a leap year.`

```js
assert.strictEqual(isLeapYear(2024), '2024 is a leap year.');
```

Avec `2000` comme valeur de la variable `year`, le `result` doit être `2000 is a leap year.`

```js
assert.strictEqual(isLeapYear(2000), '2000 is a leap year.');
```

Avec `1900` comme valeur de la variable `year`, le `result` doit être `1900 is not a leap year.`

```js
assert.strictEqual(isLeapYear(1900), '1900 is not a leap year.');

```

Tu dois appeler la fonction `isLeapYear` et passer `year` comme paramètre.

```js
assert.match(__helpers.removeJSComments(code), /isLeapYear\(\s*year\s*\)/);
```

Tu dois déclarer une variable `result`.

```js
assert.isDefined(result);
```

Tu dois stocker le résultat de l'appel de la fonction `isLeapYear` dans une variable nommée `result`.

```js
assert.match(__helpers.removeJSComments(code), /result\s*=\s*isLeapYear\(\s*year\s*\)/);
```

Tu dois afficher le `result` dans la console en utilisant `console.log()`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*result\s*\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function isLeapYear(year) {
  if ((year % 4 === 0 && year % 100 !== 0) || year % 400 === 0) {
    return(year + " is a leap year.");
  } else {
    return(year + " is not a leap year.");
  }
}

const year = 2024;
const result = isLeapYear(year);
console.log(result);
```
