---
id: 6690e10ebe2181212abc9652
title: Concevoir une carte de visite
challengeType: 25
dashedName: design-a-business-card
demoType: onClick
---

# --description--

**Objectif :** Remplir les user stories ci-dessous et faire passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu devrais avoir un `div` avec un attribut `class` ayant la valeur `business-card` qui contiendra tous les autres éléments.
1. À l'intérieur de l'élément `.business-card`, il devrait y avoir un élément `img` avec un attribut `class` ayant la valeur `profile-image`. Tu peux définir la source de l'image à `https://cdn.freecodecamp.org/curriculum/labs/flower.jpg` si tu veux. Il devrait y avoir un `alt` avec une description significative.
1. À l'intérieur de l'élément `.business-card`, après l'élément `img`, tu devrais avoir trois éléments `p` avec un attribut `class` qui a, respectivement, les valeurs `full-name`, `designation` et `company`.
1. Le premier élément `p` devrait contenir ton nom.
1. Le deuxième élément `p` devrait contenir ta désignation.
1. Le troisième élément `p` devrait avoir le nom de ton entreprise.
1. Il devrait y avoir un élément `hr` en dessous du troisième élément `p`.
1. Après l'élément `hr`, il devrait y avoir deux éléments `p`. Dans le premier élément `p`, tu devrais avoir une adresse email comme texte. Dans le deuxième élément `p`, le texte devrait être un numéro de téléphone.
1. Après les deux éléments `p`, il devrait y avoir un élément `a` avec une classe `portfolio-link`, le texte `Portfolio`, et il devrait lier à une URL valide.
1. Il devrait y avoir un élément `hr` après l'élément `a` contenant le lien du portfolio.
1. Tu devrais avoir un autre élément `div` avec un attribut `class` ayant la valeur `social-media`. Dans cet élément, il devrait y avoir un élément `h2` avec le texte `Connect with me`.
1. À l'intérieur des éléments `.social-media`, il devrait y avoir trois éléments `a`. Les éléments `a` devraient avoir respectivement le texte `Twitter`, `LinkedIn` et `GitHub` avec des liens vers des sites web valides.
1. Tu devrais lier la feuille de style `styles.css` en utilisant la balise `link` dans la section `head` du fichier HTML.
1. L'arrière-plan de ta page devrait être `rosybrown`. La police globale devrait être `Arial`, avec une police de secours `sans-serif`.
1. Le sélecteur `.business-card` devrait avoir des propriétés pour définir l'élément à une largeur de `300px` et une couleur d'arrière-plan de ton choix. Tu devrais aussi définir un `padding` de `20px` tout autour et une `margin` du haut de `100px`. Le texte devrait être aligné au centre et la taille de la police devrait être `16px`. Les margins gauche et droite devraient être définies à `auto`.
1. Le sélecteur `.profile-image` devrait avoir une propriété `max-width` avec une valeur de `100%`.
1. Tous les éléments `p` devraient avoir une `margin` du haut et du bas de `5px`.
1. Tous les liens de la page ne devraient pas être soulignés.

# --hints--

Tu devrais avoir un `div` principal avec un attribut `class` ayant la valeur `business-card`.

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('class'), 'business-card');
```

À l'intérieur de l'élément `.business-card`, il devrait y avoir un élément `img` avec un attribut `class` ayant la valeur `profile-image`.

```js
assert.exists(document.querySelector('img.profile-image'));
```

La source de l'image devrait être définie à une image valide.

```js
assert.isAbove(document.querySelector('.profile-image')?.getAttribute('src')?.length, 0);
```

L'attribut `alt` de l'image devrait être défini à une description significative.

```js
assert.isAbove(document.querySelector('img')?.alt.length, 0);
```

À l'intérieur de l'élément `.business-card`, après l'élément `img`, il devrait y avoir un élément `p` avec un attribut `class` ayant la valeur `full-name`.

```js
assert.exists(document.querySelector('div.business-card p.full-name'));
```

Après l'élément `.full-name`, il devrait y avoir un élément `p` avec un attribut `class` ayant la valeur `designation`.

```js
assert.exists(document.querySelector('div.business-card p.designation'));
```

Après l'élément `.designation`, il devrait y avoir un élément `p` avec un attribut `class` ayant la valeur `company`.

```js
assert.exists(document.querySelector('div.business-card p.company'));
```

Le premier élément `p`, celui avec la classe `full-name`, devrait contenir ton nom.

```js
assert.isNotEmpty(document.querySelector(".full-name")?.innerText.trim());
```

Le deuxième élément `p`, celui avec un attribut `class` de `designation`, devrait contenir ta désignation.

```js
assert.isNotEmpty(document.querySelector(".designation")?.innerText.trim());
```

Le troisième élément `p`, celui avec l'attribut `class` ayant la valeur `company`, devrait contenir le nom de ton entreprise.

```js
assert.isNotEmpty(document.querySelector(".company")?.innerText.trim());
```

Après le troisième élément `p`, celui avec l'attribut `class` ayant la valeur `company`, il devrait y avoir un élément `hr`.

```js
assert.exists(document.querySelector('p.company + hr'));
```

Après le premier élément `hr`, il devrait y avoir un élément `p` avec une adresse email comme texte.

```js
assert.isNotEmpty(document.querySelectorAll('.business-card hr + p')[0].innerText.trim());
```

Après l'élément `p` de l'email, il devrait y avoir un autre élément `p` avec un numéro de téléphone comme texte.

```js
assert.match(document.querySelectorAll(".business-card p")[4].textContent,/.+/);
```

Après l'élément `p` du numéro de téléphone, il devrait y avoir un élément `a` avec la classe `portfolio-link`.

```js
assert.exists(document.querySelector('a.portfolio-link'));
```

L'élément `.portfolio-link` devrait avoir le texte `Portfolio`.

```js
assert.equal(document.querySelector('a.portfolio-link').innerText.trim(), 'Portfolio');
```

Ton élément `.portfolio-link` devrait avoir un `href` valide.

```js
assert.isAbove(document.querySelector('.portfolio-link').getAttribute('href').length, 0);
```

Après l'élément `a` contenant ton portfolio, il devrait y avoir un autre élément `hr` comme séparateur.

```js
assert.exists(document.querySelector('a + hr'));
```

Après le deuxième élément `hr`, il devrait y avoir un élément `div` avec un attribut `class` ayant la valeur `social-media`.

```js
assert.exists(document.querySelector('div.business-card div.social-media'));
```

À l'intérieur de l'élément `.social-media`, il devrait y avoir un élément `h2` avec le texte `Connect with me`.

```js
assert.equal(document.querySelector('.social-media h2')?.textContent.trim(), 'Connect with me');
```

À l'intérieur de l'élément `.social-media`, il devrait y avoir trois éléments `a`.

```js
assert.lengthOf(document.querySelectorAll('.social-media a'), 3);
```

Le premier élément `a` devrait avoir le texte `Twitter`.

```js
assert.equal(document.querySelectorAll(".social-media a")[0]?.textContent.trim(), 'Twitter');
```

Le `href` du premier élément `a` devrait pointer vers un lien valide.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[0]?.getAttribute('href')?.length, 0);
```

Le deuxième élément `a` devrait avoir le texte `LinkedIn`.

```js

assert.equal(document.querySelectorAll(".social-media a")[1]?.textContent.trim(), 'LinkedIn');
```

Le `href` du deuxième élément `a` devrait pointer vers un lien valide.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[1]?.getAttribute('href')?.length, 0);
```

Le troisième élément `a` devrait avoir le texte `GitHub`.

```js
assert.equal(document.querySelectorAll(".social-media a")[2]?.textContent.trim(), 'GitHub');
```

Le `href` du troisième élément `a` devrait pointer vers un lien valide.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[2]?.getAttribute('href')?.length, 0);
```

Ton élément `link` devrait être dans ton élément `head`.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

Ton élément `link` devrait avoir un attribut `rel` avec la valeur `stylesheet`.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

Ton élément `link` devrait avoir un attribut `href` avec la valeur `styles.css`.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');
```

Tu devrais définir la couleur d'arrière-plan de la page dans le sélecteur `body` à `rosybrown`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('background-color'), 'rosybrown');
```

Tu devrais définir la police de la page à `Arial` dans l'élément `body` avec une police de secours `sans-serif`.

```js
const el = new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('font-family');
const condition1= el === 'Arial, sans-serif';
const condition2= el === '"Arial", sans-serif';
assert.isTrue(condition1 || condition2);
```

Ton sélecteur `.business-card` devrait avoir une propriété `width` avec une valeur de `300px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('width'), '300px');
```

Ton sélecteur `.business-card` devrait avoir une propriété `background-color` avec une valeur de couleur valide.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.business-card')?.backgroundColor);
```

Ton sélecteur `.business-card` devrait avoir une propriété `padding` avec une valeur de `20px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('padding'), '20px');
```

Ton sélecteur `.business-card` devrait avoir une margin du haut de `100px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-top'), '100px');
```

Le texte dans l'élément `.business-card` devrait être aligné au centre.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('text-align'), 'center');
```

La taille de la police du texte dans l'élément `.business-card` devrait être `16px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('font-size'), '16px');
```

Les margins gauche et droite de l'élément `.business-card` devraient être définies à `auto`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-left'), 'auto');
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-right'), 'auto');
```

Ton sélecteur `.profile-image` devrait avoir une propriété `max-width` avec une valeur de `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.profile-image')?.getPropVal('max-width'), '100%');
```

Tes éléments `p` devraient avoir une `margin` du haut et du bas de `5px`.

```js
const marginTop = new __helpers.CSSHelp(document).getStyle('p')?.getPropVal('margin-top');
const marginBottom = new __helpers.CSSHelp(document).getStyle('p')?.getPropVal('margin-bottom');

assert.equal(marginTop, '5px');
assert.equal(marginBottom, '5px');
```

Les liens de la page ne devraient pas avoir de soulignement.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('a')?.getPropVal('text-decoration'), 'none');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Business Card</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Digital Business Card</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="business-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/flower.jpg" alt="A vibrant flower in full bloom" class="profile-image">
        <p class="full-name">Your Name</p>
        <p class="designation">Software Developer</p>
        <p class="company" >@freeCodeCamp</p>
        <hr>
        <p>Email: your-email@example.com</p>
        <p>Phone: (123) 456-7890</p>
        <a class="portfolio-link" href="http://example.com">Portfolio</a>
        <hr>
        <div class="social-media">
            <h2>Connect with me</h2>
            <a href="http://twitter.com/">Twitter</a>
            <a href="http://linkedin.com/">LinkedIn</a>
            <a href="http://github.com/">GitHub</a>
        </div>
    </div>
</body>

</html>

```

```css
body {
    background-color: rosybrown;
    font-family: 'Arial', sans-serif;
}

.business-card {
    background-color: lavenderblush;
    border: 3px solid darkslategray;
    padding: 20px;
    text-align: center;
    font-size: 16px;
    width: 300px;
    margin-left: auto;
    margin-right: auto;
    margin-top: 100px;
}

.profile-image {
    max-width: 100%;
    width: 150px;
    height: 150px;
    margin-bottom: 10px;
    border: 2px solid lightgray;
}

.full-name {
    font-size: 24px;
    color: darkslategray;
}

.designation {
    font-size: 18px;
    color: gray;
}

.company {
    font-size: 16px;
    color: darkgray;
}

p {
    margin: 5px 0;
}

a {
    text-decoration: none;
}

.social-media {
    margin-top: 10px;
    font-size: 14px;
}

.social-media a {
    margin: 10px;
    color: maroon
}

hr {
    margin: 20px 0;
}
```
