---
id: 5cc0bd7a49b71cb96132e54c
title: Construire un générateur de plage de nombres
challengeType: 26
dashedName: build-a-range-of-numbers-generator
---

# --description--

Dans ce lab, tu vas construire une fonction `rangeOfNumbers` qui génère un tableau de nombres dans une plage donnée.

Exemples :

- `rangeOfNumbers(3, 9)` renvoie `[3, 4, 5, 6, 7, 8, 9]`
- `rangeOfNumbers(5, 5)` renvoie `[5]`

Exigences :

- Utilise la récursion (la fonction doit s'appeler elle-même)
- Pas de boucles ni de méthodes de tableau (`for`, `while`, `Array.from`, `forEach`, `map`, `filter`, `reduce`)

**Objectif :** Remplir les user stories ci-dessous et faire passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu dois créer une fonction nommée `rangeOfNumbers` qui prend deux paramètres : `startNum` et `endNum`.
1. La fonction doit renvoyer un tableau d'entiers qui commence par le nombre représenté par le paramètre `startNum` et se termine par le nombre représenté par le paramètre `endNum` (inclus).
1. Le `startNum` sera toujours inférieur ou égal au `endNum`.
1. Ta fonction doit utiliser la récursion en s'appelant elle-même. Elle ne doit utiliser aucune syntaxe de boucle (`for`, `while`, `Array.from()`, ni de fonctions d'ordre supérieur telles que `forEach`, `map`, `filter` ou `reduce`).
1. La fonction doit gérer le cas de base où `startNum` est égal à `endNum` en renvoyant un tableau contenant uniquement ce seul nombre.
1. Pour le cas récursif, la fonction doit s'appeler avec des paramètres modifiés pour construire le tableau, puis ajouter le nombre courant au résultat.
1. La fonction ne doit pas s'appuyer sur des variables globales pour mettre en cache ou construire le tableau.

# --hints--

Tu dois avoir une fonction nommée `rangeOfNumbers`.

```js
assert.isFunction(rangeOfNumbers);
```

La fonction `rangeOfNumbers` doit prendre deux paramètres.

```js
assert.match(rangeOfNumbers.toString(), /rangeOfNumbers\s*\(\s*\w+\s*,\s*\w+\s*\)/);
```

Ta fonction doit renvoyer un tableau.

```js
assert(Array.isArray(rangeOfNumbers(5, 10)));
```

Ton code ne doit utiliser aucune syntaxe de boucle (`for`, `while`, `Array.from()` ou des fonctions d'ordre supérieur telles que `forEach`, `map`, `filter` ou `reduce`).

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce|Array\.from/g)
);
```

`rangeOfNumbers` doit utiliser la récursion (s'appeler elle-même) pour résoudre ce défi.

```js
assert.isAtLeast(
  __helpers.removeJSComments(rangeOfNumbers.toString()).match(/rangeOfNumbers\s*\(.+\)/g)?.length,
  2
);
```

`rangeOfNumbers(1, 5)` devrait renvoyer `[1, 2, 3, 4, 5]`.

```js
assert.deepStrictEqual(rangeOfNumbers(1, 5), [1, 2, 3, 4, 5]);
```

`rangeOfNumbers(6, 9)` devrait renvoyer `[6, 7, 8, 9]`.

```js
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

`rangeOfNumbers(4, 4)` devrait renvoyer `[4]`.

```js
assert.deepStrictEqual(rangeOfNumbers(4, 4), [4]);
```

`rangeOfNumbers(10, 15)` devrait renvoyer `[10, 11, 12, 13, 14, 15]`.

```js
assert.deepStrictEqual(rangeOfNumbers(10, 15), [10, 11, 12, 13, 14, 15]);
```

`rangeOfNumbers(2, 8)` devrait renvoyer `[2, 3, 4, 5, 6, 7, 8]`.

```js
assert.deepStrictEqual(rangeOfNumbers(2, 8), [2, 3, 4, 5, 6, 7, 8]);
```

Les variables globales ne doivent pas être utilisées pour mettre en cache le tableau.

```js
rangeOfNumbers(1, 3);
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function rangeOfNumbers(startNum, endNum) {
  if (endNum - startNum === 0) {
    return [startNum];
  } else {
    const numbers = rangeOfNumbers(startNum, endNum - 1);
    numbers.push(endNum);
    return numbers;
  }
}
```
