---
id: 66e2f376df6f315ee81de81a
title: Étape 8
challengeType: 1
dashedName: step-8
---

# --description--

Teste ta fonction `getWordCount` en créant une variable `wordCount` définie sur l'appel de la fonction `getWordCount` avec la phrase `"I love freeCodeCamp"`.

Ensuite, affiche ce qui suit dans la console : `"Word Count: [Word count goes here]"`. Remplace `[Word count goes here]` par le nom réel de la variable. Tu peux choisir d'utiliser ici des chaînes de gabarit ou la concaténation de chaînes avec l'opérateur `+`.

Avec ça, ton projet d'analyseur de phrase est terminé !

# --hints--

Tu dois créer une variable `wordCount`.

```js
assert.isNotNull(wordCount)
```

Ta variable `wordCount` doit être définie sur le résultat de `getWordCount("I love freeCodeCamp")`.

```js
assert.equal(wordCount, getWordCount("I love freeCodeCamp"));
```

Tu dois afficher la variable `wordCount` dans la console avec le paramètre `"Word Count: ${wordCount}"`.

```js
assert.match(code, /console\.log\((?:('|"|`)Word\s+Count:\s+('|"|`)\s+\+\s+wordCount|`Word\s+Count:\s+\${wordCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

const punctuationCount = getPunctuationCount("WHAT?!?!?!?!?");
console.log(`Punctuation Count: ${punctuationCount}`);

function getWordCount(sentence) {
  if (sentence.trim() === '') {
    return 0;
  }

  const words = sentence.trim().split(' ');
  let count = 0;

  for (const word of words) {
    if (word !== '') {
      count++;
    }
  }

  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

const punctuationCount = getPunctuationCount("WHAT?!?!?!?!?");
console.log(`Punctuation Count: ${punctuationCount}`);

function getWordCount(sentence) {
  if (sentence.trim() === '') {
    return 0;
  }

  const words = sentence.trim().split(' ');
  let count = 0;

  for (const word of words) {
    if (word !== '') {
      count++;
    }
  }

  return count;
}

const wordCount = getWordCount("I love freeCodeCamp");
console.log(`Word Count: ${wordCount}`);
```
