---
id: 66e2e99caa0cb3570f4822b4
title: Étape 4
challengeType: 1
dashedName: step-4
---

# --description--

Maintenant, il est temps de tester ta fonction `getConsonantCount`.

Crée une variable `consonantCount` et affecte-lui le résultat de l'appel de la fonction `getConsonantCount` avec l'argument `"Coding is fun"`

Ensuite, affiche ce qui suit dans la console : `"Consonant Count: [Consonant count goes here]"`. Remplace `[Consonant count goes here]` par le nom réel de la variable. Tu peux choisir d'utiliser ici des chaînes de gabarit ou la concaténation de chaînes avec l'opérateur `+`.

# --hints--

Tu dois créer une variable `consonantCount`.

```js
assert.isNotNull(consonantCount)
```

Ta variable `consonantCount` doit être définie sur le résultat de `getConsonantCount("Coding is fun")`.

```js
assert.equal(consonantCount, getConsonantCount("Coding is fun"));
```

Tu dois afficher la variable `consonantCount` dans la console avec le paramètre `"Consonant Count: ${consonantCount}"`.

```js
assert.match(code, /console\.log\((?:('|"|`)Consonant\s+Count:\s+('|"|`)\s+\+\s+consonantCount|`Consonant\s+Count:\s+\${consonantCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
