---
id: 66e2edb23013ca596f8d9568
title: Étape 6
challengeType: 1
dashedName: step-6
---

# --description--

Maintenant, il est temps de tester ta fonction `getPunctuationCount`.

Crée une variable `punctuationCount` et affecte-lui le résultat de l'appel de la fonction `getPunctuationCount` avec l'argument `"WHAT?!?!?!?!?"`

Ensuite, affiche ce qui suit dans la console : `"Punctuation Count: [Punctuation count goes here]"`. Remplace `[Punctuation count goes here]` par le nom réel de la variable. Tu peux choisir d'utiliser ici des chaînes de gabarit ou la concaténation de chaînes avec l'opérateur `+`.

# --hints--

Tu dois créer une variable `punctuationCount`.

```js
assert.isNotNull(punctuationCount)
```

Ta variable `punctuationCount` doit être définie sur le résultat de `getPunctuationCount("WHAT?!?!?!?!?")`.

```js
assert.equal(punctuationCount, getPunctuationCount("WHAT?!?!?!?!?"));
```

Tu dois afficher ce qui suit dans la console : `"Punctuation Count: [Punctuation count goes here]"`. Remplace `[Punctuation count goes here]` par le nom réel de la variable. Veille à utiliser ici une syntaxe de concaténation de chaînes correcte.

```js
assert.match(code, /console\.log\((?:('|"|`)Punctuation\s+Count:\s+('|"|`)\s+\+\s+punctuationCount|`Punctuation\s+Count:\s+\${punctuationCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
