---
id: 66e2d680e129e1423116a541
title: Étape 1
challengeType: 1
dashedName: step-1
---

# --description--

Dans ce workshop, tu vas construire un analyseur de phrase qui prendra une phrase et obtiendra le décompte du nombre de mots, de voyelles, de consonnes et de signes de ponctuation.

Pour commencer, crée une fonction `getVowelCount` avec un paramètre appelé `sentence`. Ta fonction doit retourner le nombre total de voyelles dans une phrase.

# --hints--

Tu dois créer une fonction `getVowelCount`.

```js
assert.isFunction(getVowelCount);
```

Ta fonction `getVowelCount` doit avoir un paramètre `sentence`.

```js
assert.match(getVowelCount.toString(), /sentence/);
```

Ta fonction `getVowelCount` doit retourner un nombre.

```js
assert.isNumber(getVowelCount("Apples are tasty fruits"))
```

Lorsque la phrase est `"Apples are tasty fruits"`, la fonction `getVowelCount` doit retourner `7`.

```js
assert.strictEqual(getVowelCount("Apples are tasty fruits"), 7);
```

Lorsque la phrase est `"Hello, World!"`, la fonction doit retourner `3`.

```js
assert.strictEqual(getVowelCount("Hello, World!"), 3);
```

Ton décompte de voyelles doit être insensible à la casse.

```js
assert.strictEqual(getVowelCount("Apples are tasty fruits"), 7);
assert.strictEqual(getVowelCount("freeCodeCamp is awesome"), 10);
```

Ta fonction `getVowelCount` doit retourner le décompte de voyelles correct pour n'importe quelle phrase.

```js
assert.strictEqual(getVowelCount("I went to the store"), 6);
assert.strictEqual(getVowelCount("The quick brown fox jumps over the lazy dog"), 11);
assert.strictEqual(getVowelCount("The cat in the hat"), 5);
assert.strictEqual(getVowelCount(""), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
