---
id: 66e2f06b191e305c00574e4d
title: Étape 7
challengeType: 1
dashedName: step-7
---

# --description--

Enfin, compte le nombre de mots en créant une fonction `getWordCount` avec un paramètre `sentence`. La fonction doit retourner le nombre total de mots dans la phrase passée lors de son appel.

# --hints--

Tu dois créer une fonction `getWordCount`.

```js
assert.isFunction(getWordCount);
```

Ta fonction `getWordCount` doit avoir un paramètre `sentence`.

```js
assert.match(getWordCount.toString(), /sentence/);
```

Ta fonction `getWordCount` doit retourner un nombre.

```js
assert.isNumber(getWordCount("I love freeCodeCamp"))
```

Lorsque la phrase est `"When are you gonna start learning to code?"`, la fonction `getWordCount` doit retourner `8`.

```js
assert.strictEqual(getWordCount("When are you gonna start learning to code?"), 8);
```

Lorsque la phrase est `"What's going on?"`, la fonction `getWordCount` doit retourner `3`.

```js
assert.strictEqual(getWordCount("What's going on?"), 3);
```

Ton décompte de mots doit être insensible à la casse.

```js
assert.strictEqual(getWordCount("freeCodeCamp offers free coding tutorials online"), 6);
assert.strictEqual(getWordCount("You can learn HTML, CSS, JavaScript, and more on freeCodeCamp"), 10);
```

Ta fonction `getWordCount` doit retourner le décompte de mots correct pour n'importe quelle phrase.

```js
assert.strictEqual(getWordCount("freeCodeCamp has a great community of kind people"), 8);
assert.strictEqual(getWordCount("The freeCodeCamp curriculum is constantly updated"), 6);
assert.strictEqual(getWordCount("freeCodeCamp teaches both front-end and back-end development"), 7);
```

Ta fonction `getWordCount` doit retourner le décompte de mots correct pour une chaîne vide, ou une chaîne contenant uniquement des espaces.

```js
assert.strictEqual(getWordCount(""), 0);
assert.strictEqual(getWordCount("   "), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

const punctuationCount = getPunctuationCount("WHAT?!?!?!?!?");
console.log(`Punctuation Count: ${punctuationCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
