---
id: 66e2df8c5676fe51cead8ec8
title: Étape 2
challengeType: 1
dashedName: step-2
---

# --description--

Maintenant, il est temps de tester ta fonction `getVowelCount`.

Crée une variable `vowelCount` et affecte-lui le résultat de l'appel de la fonction `getVowelCount` avec l'argument `"Apples are tasty fruits"`

Ensuite, affiche ce qui suit dans la console : `"Vowel Count: [vowel count goes here]"`. Remplace `[vowel count goes here]` par le nom réel de la variable. Tu peux choisir d'utiliser ici des chaînes de gabarit ou la concaténation de chaînes avec l'opérateur `+`.

# --hints--

Tu dois créer une variable `vowelCount`.

```js
assert.isNotNull(vowelCount)
```

Ta variable `vowelCount` doit être définie sur le résultat de `getVowelCount("Apples are tasty fruits")`.

```js
assert.equal(vowelCount, getVowelCount("Apples are tasty fruits"));
```

Tu dois afficher la variable `vowelCount` dans la console avec le paramètre `"Vowel Count: ${vowelCount}"`.

```js
assert.match(code, /console\.log\((?:('|"|`)Vowel\s+Count:\s+('|"|`)\s+\+\s+vowelCount|`Vowel\s+Count:\s+\${vowelCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
