---
id: 66e2eab8a5638f57b637b7cc
title: Étape 5
challengeType: 1
dashedName: step-5
---

# --description--

Tu dois maintenant compter le nombre de signes de ponctuation.

Crée une fonction `getPunctuationCount` avec un paramètre `sentence`.

À l'intérieur de la fonction, crée une boucle pour compter le nombre de caractères de ponctuation dans la `sentence` qui sera passée à la fonction lors de son appel. Pour nos besoins, un caractère de ponctuation est tout caractère qui n'est ni un espace (`" "`) ni une lettre.

Ta fonction `getPunctuationCount` doit retourner un nombre.

# --hints--

Tu dois créer une fonction `getPunctuationCount`.

```js
assert.isFunction(getPunctuationCount);
```

Ta fonction `getPunctuationCount` doit avoir un paramètre `sentence`.

```js
assert.match(getPunctuationCount.toString(), /sentence/);
```

Ta fonction `getPunctuationCount` doit retourner un nombre.

```js
assert.isNumber(getPunctuationCount("Coding is fun!"))
```

Lorsque la phrase est `"What's going on here?"`, la fonction `getPunctuationCount` doit retourner `2`.

```js
assert.strictEqual(getPunctuationCount("What's going on here?"), 2);
```

Lorsque la phrase est `"What????!"`, la fonction `getPunctuationCount` doit retourner `5`.

```js
assert.strictEqual(getPunctuationCount("What????!"), 5);
```

Ta fonction `getPunctuationCount` doit retourner le décompte de ponctuation correct pour n'importe quelle phrase.

```js
assert.strictEqual(getPunctuationCount("Be quick, sign up! freeCodeCamp awaits, friend!!!"), 6);
assert.strictEqual(getPunctuationCount("Guess what? freeCodeCamp is launching a new cert soon!"), 2);
assert.strictEqual(getPunctuationCount("freeCodeCamp, again? It's incredible!"), 4);
assert.strictEqual(getPunctuationCount(""), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
