---
id: 66e2e5d890ac8a54c9d703bd
title: Étape 3
challengeType: 1
dashedName: step-3
---

# --description--

Il est temps de compter les consonnes. Crée une fonction `getConsonantCount` avec un paramètre `sentence`.

À l'intérieur de la fonction, utilise une boucle pour compter le nombre de consonnes dans la `sentence` qui sera passée à la fonction lors de son appel. Une consonne est toute lettre qui n'est pas l'un des caractères suivants : `"aeiou"`.

Ta fonction `getConsonantCount` doit retourner un nombre.

# --hints--

Tu dois créer une fonction `getConsonantCount`.

```js
assert.isFunction(getConsonantCount);
```

Ta fonction `getConsonantCount` doit avoir un paramètre `sentence`.

```js
assert.match(getConsonantCount.toString(), /sentence/);
```

Ta fonction `getConsonantCount` doit retourner un nombre.

```js
assert.isNumber(getConsonantCount("Coding is fun"))
```

Lorsque la phrase est `"Coding is fun"`, la fonction `getConsonantCount` doit retourner `7`.

```js
assert.strictEqual(getConsonantCount("Coding is fun"), 7);
```

Lorsque la phrase est `"Hello, World!"`, la fonction `getConsonantCount` doit retourner `7`.

```js
assert.strictEqual(getConsonantCount("Hello, World!"), 7);
```

Ton décompte de consonnes doit être insensible à la casse.

```js
assert.strictEqual(getConsonantCount("Apples are tasty fruits"), 13);
assert.strictEqual(getConsonantCount("freeCodeCamp is awesome"), 11);
```

Ta fonction `getConsonantCount` doit retourner le décompte de consonnes correct pour n'importe quelle phrase.

```js
assert.strictEqual(getConsonantCount("I went to the store"), 9);
assert.strictEqual(getConsonantCount("The quick brown fox jumps over the lazy dog"), 24);
assert.strictEqual(getConsonantCount("The cat in the hat"), 9);
assert.strictEqual(getConsonantCount(""), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
