---
id: 673b567e3ba535dda140d278
title: Construire un système de vote
challengeType: 26
dashedName: build-a-voting-system
saveSubmissionToDB: true
---

# --description--

Dans ce lab, tu vas construire un système de vote qui utilise `Map` pour créer un sondage et `Set` pour empêcher les votes en double.

**Objectif :** Remplir les user stories ci-dessous et faire passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu dois initialiser une variable `poll` avec un nouvel objet `Map`.

2. Tu dois avoir une fonction `addOption` qui accepte un paramètre `option`.

3. Dans la fonction `addOption` :

   - Si l'`option` n'existe pas déjà dans le sondage, elle doit y être ajoutée avec un `Set` vide comme valeur pour suivre les votants. Tu dois aussi renvoyer le message `Option "<option>" added to the poll.`

   - Si l'`option` existe déjà, elle doit renvoyer le message `Option "<option>" already exists.`

   - Si tu essaies d'ajouter une option vide, la fonction doit renvoyer le message `Option cannot be empty.`

4. Tu dois avoir une fonction `vote` qui accepte deux paramètres, `option` (l'option pour laquelle voter) et `voterId` (un ID unique pour le votant).

5. Dans la fonction `vote` :

   - Si l'`option` n'existe pas dans le sondage, la fonction doit renvoyer le message `Option "<option>" does not exist.`

   - Si l'`option` existe, la fonction doit vérifier si le `voterId` a déjà voté pour cette `option`.

   - Si le votant a déjà voté, la fonction doit renvoyer le message `Voter <voterId> has already voted for "<option>".`

   - Si le votant n'a pas encore voté, `voterId` doit être ajouté au `Set` de `voters` pour cette option. La fonction doit renvoyer le message `Voter <voterId> voted for "<option>".`

6. Tu dois avoir au moins trois options dans ton `poll`.

7. Ton `poll` doit avoir au moins trois votes.

8. Tu dois avoir une fonction `displayResults` qui renvoie les résultats du sondage au format suivant :

```js
Poll Results:
OptionA: N votes
OptionB: N votes
.
.

/*
sample output

Poll Results:
Turkey: 2 votes
Morocco: 1 votes
*/
```

# --hints--

Tu dois avoir une variable `poll` initialisée avec un nouvel objet `Map`.

```js
assert.exists(poll);
assert.instanceOf(poll, Map);
```

Tu dois définir une fonction `addOption` qui accepte un paramètre `option`.

```js
assert.isFunction(addOption);
assert.match(
  addOption.toString(),
  /^function\s+addOption\s*\(\s*option\s*\)\s*\{|\(?\s*option\s*\)?\s*=>/
);
```

Tu dois définir une fonction `vote` qui accepte deux paramètres, `option` et `voterId`.

```js
assert.isFunction(vote);
assert.match(vote.toString(), /\(\s*option\s*,\s*voterId\s*\)/);
```

Tu dois définir une fonction `displayResults` pour afficher les résultats du sondage.

```js
assert.isFunction(displayResults);
```

Tu dois avoir au moins trois options dans ton `poll`.

```js
assert.isAtLeast(poll.size, 3);
```

Ton `poll` doit avoir au moins trois votes.

```js
let totalVotes = 0;
for (let voters of poll.values()) {
  totalVotes += voters.size;
}
assert.isAtLeast(totalVotes, 3);
```

Tu dois t'assurer que chaque option de vote correspond à un objet `Set`.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();

  addOption('Turkey');
  assert.instanceOf(poll.get('Turkey'), Set);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`addOption("Egypt")` devrait renvoyer `Option "Egypt" added to the poll.`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  assert.equal(addOption('Egypt'), 'Option "Egypt" added to the poll.');
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

L'ajout d'une option vide devrait renvoyer `"Option cannot be empty."`

```js
assert.equal(addOption(''), 'Option cannot be empty.');
```

Quand `Turkey` est déjà ajouté, `addOption("Turkey")` devrait renvoyer `Option "Turkey" already exists.`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Turkey');
  assert.equal(addOption('Turkey'), 'Option "Turkey" already exists.');
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Quand `Malaysia` existe dans les options de vote, `vote("Malaysia", "traveler1")` devrait renvoyer `Voter traveler1 voted for "Malaysia".`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Malaysia');
  assert.equal(
    vote('Malaysia', 'traveler1'),
    'Voter traveler1 voted for "Malaysia".'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`vote` devrait mettre à jour le `Set` des votants pour une option.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Alaska');
  vote('Alaska', 'seal');
  assert.isTrue(poll.get('Alaska').has('seal'));
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Quand `traveler1` essaie de voter à nouveau pour `Algeria`, `vote("Algeria", "traveler1")` devrait renvoyer `Voter traveler1 has already voted for "Algeria".`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Algeria');
  vote('Algeria', 'traveler1');
  assert.equal(
    vote('Algeria', 'traveler1'),
    'Voter traveler1 has already voted for "Algeria".'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Les votes en double ne devraient pas augmenter la taille du `Set`.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  vote('Turkey', 'traveler1');
  const voterSet = poll.get('Turkey');
  const initialSize = voterSet.size;

  vote('Turkey', 'traveler1');
  assert.equal(voterSet.size, initialSize);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Quand `Nigeria` n'est pas dans les options de vote, `vote("Nigeria", "traveler2")` devrait renvoyer `Option "Nigeria" does not exist.`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  assert.equal(
    vote('Nigeria', 'traveler2'),
    'Option "Nigeria" does not exist.'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Une option unique devrait pouvoir recevoir plusieurs votes.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Bali');
  assert.equal(vote('Bali', 'traveler1'), 'Voter traveler1 voted for "Bali".');
  assert.equal(vote('Bali', 'traveler2'), 'Voter traveler2 voted for "Bali".');
  assert.equal(poll.get('Bali').size, 2);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Quand l'option `Turkey` obtient deux votes, l'option `Morocco` un vote, et l'option `Spain` aucun vote, `displayResults()` devrait renvoyer `"Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes\nSpain: 0 votes"`.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Morocco');
  addOption('Spain');

  vote('Turkey', 'traveler1');
  vote('Turkey', 'traveler2');
  vote('Morocco', 'traveler3');

  assert.equal(
    displayResults(),
    'Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes\nSpain: 0 votes'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`displayResults()` devrait renvoyer les résultats au format correct.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Morocco');

  vote('Turkey', 'traveler1');
  vote('Turkey', 'traveler2');
  vote('Morocco', 'traveler3');

  assert.equal(
    displayResults(),
    'Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const poll = new Map();

function addOption(option) {
  if (!option || option.trim() === '') {
    return 'Option cannot be empty.';
  }
  if (!poll.has(option)) {
    poll.set(option, new Set());
    return `Option "${option}" added to the poll.`;
  } else {
    return `Option "${option}" already exists.`;
  }
}

function vote(option, voterId) {
  if (!poll.has(option)) {
    return `Option "${option}" does not exist.`;
  }
  const voters = poll.get(option);
  if (voters.has(voterId)) {
    return `Voter ${voterId} has already voted for "${option}".`;
  } else {
    voters.add(voterId);
    return `Voter ${voterId} voted for "${option}".`;
  }
}

function displayResults() {
  let results = 'Poll Results:\n';
  for (let [option, voters] of poll.entries()) {
    results += `${option}: ${voters.size} votes\n`;
  }
  return results.trim();
}

addOption('Turkey');
addOption('Morocco');
addOption('Spain');

vote('Turkey', 'traveler1');
vote('Turkey', 'traveler2');
vote('Morocco', 'traveler3');
```
