---
id: 66bba6fff611169359d9d36a
title: Créer une page d'e-mail confidentiel
challengeType: 25
dashedName: build-a-confidential-email-page
demoType: onClick
---

# --description--

**Objectif :** Remplis les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu devrais avoir un élément `main` avec un `id` de `email`.
1. Ton élément `#email` devrait avoir un `padding` de `50px`, une margin supérieure de `50px`, une `width` de `500px` et une `border` de `2px`.
1. La largeur totale de ton élément `#email`, y compris les paddings et les bordures, devrait être de `500px`.
1. Tu devrais avoir deux éléments `div`, l'un avec un id de `confidential` et l'autre avec un id de `top-secret`, dans ton élément `#email`.
1. Tes éléments `#confidential` et `#top-secret` devraient avoir un `display` de `inline-block`.
1. Tes éléments `#confidential` et `#top-secret` devraient avoir un `padding`, une margin gauche et une `border`.
1. L'élément `#confidential` devrait avoir le texte `CONFIDENTIAL`.
1. L'élément `#top-secret` devrait avoir le texte `TOP SECRET`.
1. Tes éléments `#confidential` et `#top-secret` devraient être pivotés avec une transformation CSS.
1. Tu devrais avoir au moins trois éléments paragraphe dans ton élément `#email`.
1. Tu devrais avoir au moins trois éléments `span` avec une classe `blurred`, dans tes éléments paragraphe.
1. Tu devrais avoir un sélecteur `.blurred` qui floute l'élément de `3px` avec un filtre CSS.

# --hints--

Tu devrais avoir un élément `main` avec un `id` de `email`.

```js
assert.exists(document.querySelector('main#email'));
```

Tu devrais avoir un sélecteur `#email` qui définit le `padding` de ses éléments sur tous les côtés à `50px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('padding'), '50px');
```

Tu devrais avoir un sélecteur `#email` qui définit le `margin-top` de ses éléments à `50px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('margin-top'), '50px');
```

Tu devrais avoir un sélecteur `#email` qui définit la `width` de ses éléments à `500px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('width'), '500px');
```

Ton élément `#email` devrait avoir une bordure de `2px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('border-width'), '2px');
```

Ton élément `#email` devrait avoir un `box-sizing` de `border-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.boxSizing, 'border-box');
```

Tu devrais avoir un élément `div` avec un `id` de `confidential` dans ton élément `#email`.

```js
assert.exists(document.querySelector('#email > div#confidential'));
```

Tu devrais avoir un élément `div` avec un `id` de `top-secret` dans ton élément `#email`.

```js
assert.exists(document.querySelector('#email > div#top-secret'));
```

Ton élément `#confidential` devrait avoir son `display` défini à `inline-block`.

```js
assert.equal(window.getComputedStyle(document.querySelector('#confidential'))?.display, 'inline-block');
```

Ton élément `#top-secret` devrait avoir son `display` défini à `inline-block`.

```js
assert.equal(window.getComputedStyle(document.querySelector('#top-secret'))?.display, 'inline-block');
```

Ton élément `#confidential` devrait avoir un `padding` sur tous les côtés.

```js
const style = window.getComputedStyle(document.querySelector('#confidential'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

Ton élément `#top-secret` devrait avoir un `padding` sur tous les côtés.

```js
const style = window.getComputedStyle(document.querySelector('#top-secret'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

Ton élément `#confidential` devrait avoir une margin gauche.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.marginLeft, '0px');
```

Ton élément `#top-secret` devrait avoir une margin gauche.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.marginLeft, '0px');
```

Ton élément `#confidential` devrait avoir une `border`.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.borderWidth, '0px');
```

Ton élément `#top-secret` devrait avoir une `border`.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.borderWidth, '0px');
```

Le texte de ton élément `#confidential` devrait être `CONFIDENTIAL`. Tu as soit omis le texte, soit fait une faute de frappe.

```js
assert.equal(document.querySelector('#confidential').innerText.toLowerCase(), 'confidential' );
```

Le texte de ton élément `#top-secret` devrait être `TOP SECRET`. Tu as soit omis le texte, soit fait une faute de frappe.

```js
assert.equal(document.querySelector('#top-secret').innerText.toLowerCase(), 'top secret' );
```

Ton élément `#confidential` devrait utiliser une `transform` CSS pour faire pivoter l'élément.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.transform, 'none');
```

Ton élément `#top-secret` devrait utiliser une `transform` CSS pour faire pivoter l'élément.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.transform, 'none');
```

Tu devrais avoir au moins trois éléments paragraphe dans ton élément `#email`.

```js
assert.isAtLeast(document.querySelectorAll('#email > p').length, 3);
```

Tu devrais avoir au moins trois éléments `span` avec une classe `blurred` dans tes paragraphes.

```js
assert.isAtLeast(document.querySelectorAll('p > span.blurred').length, 3);
```

Tes éléments `.blurred` ne devraient pas être vides.

```js
const els = document.querySelectorAll('.blurred');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.isAtLeast(el.innerText.length , 1))
```

Tu devrais avoir un sélecteur `.blurred` qui définit le `filter` de ses éléments à `blur(3px)`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.blurred')?.filter, 'blur(3px)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Confidential Email</title>

</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Secret Marshmallow Mission</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main id="email">
    <div id="confidential">CONFIDENTIAL</div>
    <p>Dear Agent <span class="blurred">S'more,</span></p>
    <p>We have an emergency. The secret formula for our <span class="blurred">Mega Marshmallow</span> has been
      compromised. This formula is what makes our marshmallows the fluffiest and most delicious.</p>
    <p>We suspect that <span class="blurred">Professor Puff</span> is behind this. He has taken the formula to his
      hidden laboratory. Your mission is to infiltrate the lab and secure the formula before it's too late.</p>
    <p>Be sure to keep the lab's location confidential. Any leak of this information could jeopardize the entire
      operation.</p>
    <div id="top-secret">TOP SECRET</div>
  </main>
</body>

</html>
```

```css
#email {
  background-color: SeaShell;
  margin: 50px auto;
  padding: 50px;
  width: 500px;
  border: 2px solid black;
  box-shadow: 5px 3px 3px gray;
  box-sizing: border-box;
}

#confidential {
  display: inline-block;
  margin-left: 100px;
  padding: 10px;
  text-align: center;
  font-size: 30px;
  transform: rotate(-20deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

#top-secret {
  margin-left: 200px;
  display: inline-block;
  padding: 10px;
  text-align: center;
  font-size: 20px;
  transform: rotate(15deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

.blurred {
  filter: blur(3px);
}
```
