---
id: 66e7ee20b79186306fc12da5
title: Construire un ensemble de cartes d'équipe de football
challengeType: 25 
dashedName: lab-football-team-cards
demoType: onClick
---

# --description--

Dans ce lab, tu vas construire un ensemble de cartes d'équipe de football. L'utilisateur devrait pouvoir utiliser le menu déroulant et filtrer les différents joueurs selon leurs postes.

**Objectif :** Remplis les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu devrais créer un objet `footballTeam` avec les propriétés suivantes : `team`, `year`, `headCoach`, `players`.
1. La propriété `team` devrait contenir le nom de l'équipe comme chaîne.
1. La propriété `year` devrait contenir l'année comme nombre.
1. La propriété `headCoach` devrait contenir le nom de l'entraîneur principal comme chaîne.
1. La propriété `players` devrait être un tableau contenant au moins quatre éléments.
1. Chaque élément du tableau `players` devrait être un objet avec les propriétés `name`, `position`, `isCaptain`.
1. La propriété `name` devrait contenir le nom du joueur comme chaîne.
1. La propriété `position` devrait avoir une des valeurs suivantes : `"forward"`, `"midfielder"`, `"defender"` ou `"goalkeeper"`.
1. La propriété `isCaptain` devrait avoir une valeur booléenne. Un des joueurs devrait avoir sa propriété `isCaptain` définie à `true`.
1. Tu devrais afficher les valeurs `headCoach`, `team` et `year` sur la page. Ces valeurs devraient être affichées dans les éléments HTML avec les `id` `head-coach`, `team` et `year`.
1. Tu devrais afficher les données des joueurs sur la page à l'intérieur de l'élément `#player-cards`, chaque joueur devrait être affiché dans un `div` avec la classe `player-card`, et imbriqués dedans, un `h2` contenant le nom du joueur, et `(Captain)` si le joueur est capitaine, et un `p` contenant `Position:` et le poste du joueur.

    Voici un exemple de ce à quoi le HTML devrait ressembler :

    ```html
    <div class="player-card">
      <h2>Sergio Batista</h2>
      <p>Position: midfielder</p>
    </div>
    <div class="player-card">
      <h2>(Captain) Diego Maradona</h2>
      <p>Position: midfielder</p>
    </div>
    ```

1. Quand le menu déroulant est utilisé pour sélectionner un des postes, seuls les éléments `.player-card` des joueurs de ce poste devraient être présents. Si l'option `"All Players"` est sélectionnée, alors tous les joueurs devraient s'afficher sur la page.

# --hints--

Tu devrais avoir une variable `footballTeam`.

```js
assert.exists(footballTeam);
```

La variable `footballTeam` devrait être un objet avec quatre propriétés : `team`, `year`, `headCoach` et `players`.

```js
assert.isObject(footballTeam);
assert.containsAllKeys(footballTeam, ['team', 'year', 'headCoach', 'players']);
```

La propriété `team` devrait être une chaîne.

```js
assert.isString(footballTeam.team);
```

La propriété `year` devrait être un nombre.

```js
assert.isNumber(footballTeam.year);
```

La propriété `headCoach` devrait être une chaîne.

```js
assert.isString(footballTeam.headCoach);
```

La propriété `players` devrait être un tableau d'au moins quatre objets, chaque objet devrait avoir les clés `name`, `position`, `isCaptain`.

```js
assert.isArray(footballTeam.players);
assert.isAtLeast(footballTeam.players.length, 4);
footballTeam.players.forEach(player => {
  assert.isObject(player);
  assert.containsAllKeys(player, ['name', 'position', 'isCaptain']);
})
```

La propriété `name` devrait avoir une valeur de type chaîne.

```js
footballTeam.players.forEach(({name}) => assert.isString(name));
```

La propriété `position` devrait avoir une des valeurs `"forward"`, `"midfielder"`, `"defender"` ou `"goalkeeper"`.

```js
footballTeam.players.forEach(({position}) => {
  assert.isString(position);
  assert.oneOf(position, ["forward", "midfielder", "defender", "goalkeeper"]);
});
```

La propriété `isCaptain` devrait avoir une valeur booléenne, et il ne devrait y avoir qu'un seul capitaine.

```js
footballTeam.players.forEach(({isCaptain}) => assert.isBoolean(isCaptain));
const listOfCaptains = footballTeam.players.filter(({isCaptain}) => isCaptain);
assert.lengthOf(listOfCaptains, 1);
```

Tu devrais afficher les valeurs `headCoach`, `team` et `year` de l'objet `footballTeam` dans les éléments HTML avec les `id` `head-coach`, `team` et `year`.

```js
const teamElement = document.querySelector('.team-stats #team');
const yearElement = document.querySelector('.team-stats #year');
const headCoachElement = document.querySelector('.team-stats #head-coach');
assert.equal(teamElement?.innerText.trim(), footballTeam.team);
assert.equal(yearElement?.innerText.trim(), footballTeam.year);
assert.equal(headCoachElement?.innerText.trim(), footballTeam.headCoach);
```

Quand l'option `All Players` est sélectionnée, tous les joueurs devraient être présents dans `#player-cards`.

```js
const select = document.querySelector('#players')
select.value = 'all';
select.dispatchEvent(new Event('change'))
const playerCards = document.querySelectorAll('.player-card');
const arrayFromPage = Array.from(playerCards).map(el => ({
    name: el.querySelector('h2').innerText.replace('(Captain)', '').trim(),
    position: el.querySelector('p').innerText.replace('Position:', '').trim(),
    isCaptain: /Captain/.test(el.querySelector('h2').innerText)
  }))

assert.sameDeepMembers(arrayFromPage, footballTeam.players);
```

Quand l'option `Position Forward` est sélectionnée, seuls les joueurs forward devraient être présents dans `#player-cards`.

```js
const forwards = footballTeam.players.filter(({position}) => position === 'forward')
const select = document.querySelector('#players')
select.value = 'forward';
select.dispatchEvent(new Event('change'))
const playerCards = document.querySelectorAll('.player-card');
const arrayFromPage = Array.from(playerCards).map(el => ({
    name: el.querySelector('h2').innerText.replace('(Captain)', '').trim(),
    position: el.querySelector('p').innerText.replace('Position:', '').trim(),
    isCaptain: /Captain/.test(el.querySelector('h2').innerText)
  }))

assert.sameDeepMembers(arrayFromPage, forwards);
```

Quand l'option `Position Midfielder` est sélectionnée, seuls les joueurs midfielder devraient être présents dans `#player-cards`.

```js
const midfielders = footballTeam.players.filter(({position}) => position === 'midfielder')
const select = document.querySelector('#players')
select.value = 'midfielder';
select.dispatchEvent(new Event('change'))
const playerCards = document.querySelectorAll('.player-card');
const arrayFromPage = Array.from(playerCards).map(el => ({
    name: el.querySelector('h2').innerText.replace('(Captain)', '').trim(),
    position: el.querySelector('p').innerText.replace('Position:', '').trim(),
    isCaptain: /Captain/.test(el.querySelector('h2').innerText)
  }))

assert.sameDeepMembers(arrayFromPage, midfielders);
```

Quand l'option `Position Defender` est sélectionnée, seuls les joueurs defender devraient être présents dans `#player-cards`.

```js
const defenders = footballTeam.players.filter(({position}) => position === 'defender')
const select = document.querySelector('#players')
select.value = 'defender';
select.dispatchEvent(new Event('change'))
const playerCards = document.querySelectorAll('.player-card');
const arrayFromPage = Array.from(playerCards).map(el => ({
    name: el.querySelector('h2').innerText.replace('(Captain)', '').trim(),
    position: el.querySelector('p').innerText.replace('Position:', '').trim(),
    isCaptain: /Captain/.test(el.querySelector('h2').innerText)
  }))

assert.sameDeepMembers(arrayFromPage, defenders);
```

Quand l'option `Position Goalkeeper` est sélectionnée, seuls les joueurs goalkeeper devraient être présents.

```js
const goalkeepers = footballTeam.players.filter(({position}) => position === 'goalkeeper')
const select = document.querySelector('#players')
select.value = 'goalkeeper';
select.dispatchEvent(new Event('change'))
const playerCards = document.querySelectorAll('.player-card');
const arrayFromPage = Array.from(playerCards).map(el => ({
    name: el.querySelector('h2').innerText.replace('(Captain)', '').trim(),
    position: el.querySelector('p').innerText.replace('Position:', '').trim(),
    isCaptain: /Captain/.test(el.querySelector('h2').innerText)
  }))

assert.sameDeepMembers(arrayFromPage, goalkeepers);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>
      Build a Set of Football Team Cards
    </title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1 class="title">Team stats</h1>
    <main>
      <div class="team-stats">
        <p>Team: <span id="team"></span></p>
        <p>Year: <span id="year"></span></p>
        <p>Head coach: <span id="head-coach"></span></p>
      </div>
      <label class="options-label" for="players">Filter Teammates:</label>
      <select name="players" id="players">
        <option value="all">All Players</option>
        <option value="forward">Position Forward</option>
        <option value="midfielder">Position Midfielder</option>
        <option value="defender">Position Defender</option>
        <option value="goalkeeper">Position Goalkeeper</option>
      </select>
      <div class="cards" id="player-cards"></div>
    </main>
    <footer>&copy; freeCodeCamp</footer>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --dark-grey: #0a0a23;
  --light-grey: #f5f6f7;
  --white: #ffffff;
  --black: #000;
}

body {
  background-color: var(--dark-grey);
  text-align: center;
  padding: 10px;
}

.title,
.options-label,
.team-stats,
footer {
  color: var(--white);
}

.title {
  margin: 1.3rem 0;
}

.team-stats {
  display: flex;
  justify-content: space-around;
  flex-wrap: wrap;
  font-size: 1.3rem;
  margin: 1.2rem 0;
}

.options-label {
  font-size: 1.2rem;
}

.cards {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
  align-items: center;
}

.player-card {
  background-color: var(--light-grey);
  padding: 1.3rem;
  margin: 1.2rem;
  width: 300px;
  border-radius: 15px;
}

@media (max-width: 768px) {
  .team-stats {
    flex-direction: column;
  }
}

```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>
      Build a Set of Football Team Cards
    </title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1 class="title">Team stats</h1>
    <main>
      <div class="team-stats">
        <p>Team: <span id="team"></span></p>
        <p>Year: <span id="year"></span></p>
        <p>Head coach: <span id="head-coach"></span></p>
      </div>
      <label class="options-label" for="players">Filter Teammates:</label>
      <select name="players" id="players">
        <option value="all">All Players</option>
        <option value="forward">Position Forward</option>
        <option value="midfielder">Position Midfielder</option>
        <option value="defender">Position Defender</option>
        <option value="goalkeeper">Position Goalkeeper</option>
      </select>
      <div class="cards" id="player-cards"></div>
    </main>
    <footer>&copy; freeCodeCamp</footer>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --dark-grey: #0a0a23;
  --light-grey: #f5f6f7;
  --white: #ffffff;
  --black: #000;
}

body {
  background-color: var(--dark-grey);
  text-align: center;
  padding: 10px;
}

.title,
.options-label,
.team-stats,
footer {
  color: var(--white);
}

.title {
  margin: 1.3rem 0;
}

.team-stats {
  display: flex;
  justify-content: space-around;
  flex-wrap: wrap;
  font-size: 1.3rem;
  margin: 1.2rem 0;
}

.options-label {
  font-size: 1.2rem;
}

.cards {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
  align-items: center;
}

.player-card {
  background-color: var(--light-grey);
  padding: 1.3rem;
  margin: 1.2rem;
  width: 300px;
  border-radius: 15px;
}

@media (max-width: 768px) {
  .team-stats {
    flex-direction: column;
  }
}

```

```js
const teamName = document.getElementById("team");
const worldCupYear = document.getElementById("year");
const headCoach = document.getElementById("head-coach");
const playerCards = document.getElementById("player-cards");
const playersDropdownList = document.getElementById("players");
const footballTeam = {
  team: "Argentina",
  year: 1986,
  headCoach: "Carlos Bilardo",
  players: [
    {
      name: "Sergio Almirón",
      position: "forward",
      isCaptain: false,
    },
    {
      name: "Sergio Batista",
      position: "midfielder",
      isCaptain: false,
    },
     {
      name: "Ricardo Bochini",
      position: "midfielder",
      isCaptain: false,
    },
    {
      name: "Claudio Borghi",
      position: "midfielder",
      isCaptain: false,
    },
    {
      name: "José Luis Brown",
      position: "defender",
      isCaptain: false,
    },
    {
      name: "Daniel Passarella",
      position: "defender",
      isCaptain: false,
    },
    {
      name: "Jorge Burruchaga",
      position: "forward",
      isCaptain: false,
    },
    {
      name: "Néstor Clausen",
      position: "defender",
      isCaptain: false,
    },
    {
      name: "José Luis Cuciuffo",
      position: "defender",
      isCaptain: false,
    },
    {
      name: "Diego Maradona",
      position: "midfielder",
      isCaptain: true,
    },
    {
      name: "Jorge Valdano",
      position: "forward",
      isCaptain: false,
    },
    {
      name: "Héctor Enrique",
      position: "midfielder",
      isCaptain: false,
    },
    {
      name: "Oscar Garré",
      position: "defender",
      isCaptain: false,
    },
    {
      name: "Ricardo Giusti",
      position: "midfielder",
      isCaptain: false,
    },
    {
      name: "Luis Islas",
      position: "goalkeeper",
      isCaptain: false,
    },
    {
      name: "Julio Olarticoechea",
      position: "defender",
      isCaptain: false,
    },
    {
      name: "Pedro Pasculli",
      position: "forward",
      isCaptain: false,
    },
    {
      name: "Nery Pumpido",
      position: "goalkeeper",
      isCaptain: false,
    },
    {
      name: "Oscar Ruggeri",
      position: "defender",
      isCaptain: false,
    },
    {
      name: "Carlos Tapia",
      position: "midfielder",
      isCaptain: false,
    },
    {
      name: "Marcelo Trobbiani",
      position: "midfielder",
      isCaptain: false,
    },
    {
      name: "Héctor Zelada",
      position: "goalkeeper",
      isCaptain: false,
    },
  ],
};

const { sport, team, year, players } = footballTeam;
const coachName = footballTeam.headCoach;

teamName.textContent = team;
worldCupYear.textContent = year;
headCoach.textContent = coachName;

const setPlayerCards = (arr = players) => {
  playerCards.innerHTML += arr
    .map(
      ({ name, position, isCaptain }) =>
        `
        <div class="player-card">
          <h2>${isCaptain ? "(Captain)" : ""} ${name}</h2>
          <p>Position: ${position}</p>
        </div>
      `
    )
    .join("");
};

playersDropdownList.addEventListener("change", (e) => {
  playerCards.innerHTML = "";

  switch (e.target.value) {
    case "forward":
      setPlayerCards(players.filter((player) => player.position === "forward"));
      break;
    case "midfielder":
      setPlayerCards(
        players.filter((player) => player.position === "midfielder")
      );
      break;
    case "defender":
      setPlayerCards(
        players.filter((player) => player.position === "defender")
      );
      break;
    case "goalkeeper":
      setPlayerCards(
        players.filter((player) => player.position === "goalkeeper")
      );
      break;
    default:
      setPlayerCards();
  }
});

setPlayerCards();

```
