---
id: 66e45c8140f9fda5c105ae26
title: Construire une page de flyer d'événement
challengeType: 25
dashedName: build-an-event-flyer-page
demoType: onClick
---

# --description--

**Objectif :** Remplir les user stories ci-dessous et faire passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu devrais avoir un élément `header` dans le body.
1. Ton élément `header` devrait avoir une image dedans pour ton événement, et un `h1` dedans avec le titre de ton événement, dans cet ordre. Tu peux utiliser cette image si tu veux : `https://cdn.freecodecamp.org/curriculum/labs/event.jpg`
1. Tu devrais avoir un élément `main` dans le body.
1. Ton élément `main` devrait avoir au moins deux éléments `section` dedans présentant les fonctionnalités de l'événement.
1. Tes éléments `section` devraient chacun avoir un `h2` dedans.
1. Ton body devrait avoir un padding haut et bas de `50px`.
1. Ton body devrait avoir une margin haute et basse de `0`, et une margin gauche et droite qui le centre.
1. Ton body devrait avoir une largeur définie relativement à la largeur du viewport.
1. Ton body devrait utiliser la fonction `calc` pour définir sa propriété `min-height` à 100% de la hauteur du viewport moins tout le padding appliqué en haut et en bas du body.
1. Tu devrais avoir au moins un `hr` dans ton flyer.
1. Tu devrais définir la `width` de tous les éléments `hr` et `section` à une valeur en pourcentage relative à leur parent.

**Note :** Assure-toi de lier ta feuille de style dans ton HTML et d'appliquer ton CSS.

# --hints--

Tu devrais avoir un élément `header` dans l'élément `body`.

```js
assert.exists(document.querySelector('body > header'));
```

Tu devrais avoir un élément `img` comme premier enfant dans le `header`.

```js
assert.equal(document.querySelector('header')?.children?.[0].tagName, 'IMG');
```

Tu devrais avoir un élément `h1` comme deuxième enfant dans le `header`.

```js
assert.equal(document.querySelector('header')?.children?.[1].tagName, 'H1');
```

Tu ne devrais avoir qu'un seul élément `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Ton `h1` devrait avoir le titre de ton événement.

```js
assert.isAtLeast(document.querySelector('h1').innerText.length, 1);
```

Tu devrais avoir un élément `main` dans l'élément `body`.

```js
assert.exists(document.querySelector('body > main'));
```

Tu devrais avoir au moins deux éléments `section` dans ton élément `main`.

```js
assert.isAtLeast(document.querySelectorAll('main > section').length, 2);
```

Tes éléments `section` devraient chacun avoir un `h2` dedans.

```js
const sections = document.querySelectorAll('main > section');
const headers = document.querySelectorAll('main > section > h2');
assert.isAtLeast(sections.length, 1);
assert.equal(sections.length, headers.length);
```

Tes éléments `h2` ne devraient pas être vides.

```js
const headers = document.querySelectorAll('main > section > h2');
assert.isAtLeast(headers.length, 1);
headers.forEach(header => assert.notEqual(header.innerHTML.trim(), ''));
```

Ton élément `body` devrait avoir un padding haut et bas de `50px`.

```js
const body = window.getComputedStyle(document.querySelector('body'));
assert.equal(body?.paddingTop, '50px');
assert.equal(body?.paddingBottom, '50px');
```

Ton élément `body` devrait avoir une margin haute et basse de `0`.

```js
const body = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(body?.marginTop, '0px');
assert.equal(body?.marginBottom, '0px');
```

Ton élément `body` devrait avoir une margin gauche et droite de `auto`.

```js
const body = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(body?.marginRight, 'auto');
assert.equal(body?.marginLeft, 'auto');
```

Ton élément `body` devrait définir sa largeur en utilisant `vw`.

```js
assert.isTrue(new __helpers.CSSHelp(document).getStyle('body')?.width?.endsWith('vw'));
```

Ton `body` devrait utiliser `calc` pour définir sa `min-height` à `100vh - 100px`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('body')?.minHeight, ['calc(-100px + 100vh)', 'calc(100vh - 100px)']);
```

Tu devrais avoir au moins un élément `hr`.

```js
assert.isAtLeast(document.querySelectorAll('hr')?.length, 1);
```

La largeur de tes éléments `hr` devrait être définie en utilisant une valeur en pourcentage.

```js
assert.isTrue([...new __helpers.CSSHelp(document).getCSSRules().values()].some(rule => 
  rule.selectorText?.split(',').some(selector => 
    selector.trim() === 'hr' && 
    rule.style.width?.endsWith('%')
  )
));
```

La largeur de tes éléments `section` devrait être définie en utilisant une valeur en pourcentage.

```js
assert.isTrue([...new __helpers.CSSHelp(document).getCSSRules().values()].some(rule => 
  rule.selectorText?.split(',').some(selector => 
    selector.trim() === 'section' && 
    rule.style.width?.endsWith('%')
  )
));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Build an Event Flyer Page</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Build an Event Flyer Page</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>

  <header>
    <img src="https://cdn.freecodecamp.org/curriculum/labs/rock.jpg">
    <h1>Pet Rock Festival</h1>
    <p>Join us for the first annual Pet Rock Festival!</p>
    <p><b>When:</b> <time datetime="2025-10-03">October 3rd</time> <b>Where:</b> Boulder Park</p>
  </header>

  <hr>

  <main>
    <section>
      <h2>Festival Highlights</h2>
      <ul>
        <li>Pet Rock Obedience Shows</li>
        <li>Best Dressed Pet Rock Contest</li>
        <li>Pet Rock Uphill Racing</li>
      </ul>
    </section>

    <section>
      <h2>Special Guests</h2>
      <ul>
        <li>John Gravel</li>
        <li>Wendy Cobblestone</li>
      </ul>
    </section>

    <section>
      <h2>Food Trucks</h2>
      <ul>
        <li>Rock 'n' Roll Tacos</li>
        <li>Granite Grill</li>
        <li>Stone Cold Smoothies</li>
      </ul>
    </section>
  </main>

  <hr>

<footer>
  <p>&copy; 2025 Pet Rock Festival</p>
</footer>

</body>

</html>
```

```css
html {
  background-color: gray;
}

body {
  padding: 50px 10px;
  margin: 0 auto;
  width: 90vw;
  min-height: calc(100vh - 100px);
  background-color: white;
  font-family: sans-serif;
  text-align: center;
}

hr {
  margin: 30px auto;
  width: 80%;
  background-color: black;
  height: 2px;
}

main {
  margin: 0 auto;
}

section {
  text-align: center;
  padding: 50px 5px;
  vertical-align: top;
  display: inline-block;
  min-height: 100px;
  width: 30%;
}

ul {
  text-align: left;
  list-style: inside;
  padding-left: 10px;
}
```
