---
id: 66c8ba975ee7230e29f6c4b2
title: Étape 7
challengeType: 1
dashedName: step-7
---

# --description--

Maintenant, il est temps de tester ta fonction `getLoanMessage`.

Utilise le tableau ci-dessous pour créer `4` variables et leurs valeurs :

| Nom de la variable | Valeur   |
| ----------- | ------- |
| `duplexLoanMsg`       | `getLoanMessage(85000, 850)` |
| `condoLoanMsg`   | `getLoanMessage(65000, 690)`   |
| `carLoanMsg`   | `getLoanMessage(45000, 660)`   |
| `noLoanMsg`   | `getLoanMessage(25000, 550)`   |

Ensuite, affiche chaque variable dans la console pour voir les messages.

Avec ça, ton projet de vérification d'éligibilité à un prêt est terminé !

# --hints--

Tu dois avoir une variable appelée `duplexLoanMsg`.

```js
assert.isNotNull(duplexLoanMsg);
```

Ta variable `duplexLoanMsg` doit être définie sur le résultat de `getLoanMessage(85000, 850)`.

```js
assert.equal(duplexLoanMsg, getLoanMessage(85000, 850));
```

Tu dois afficher la variable `duplexLoanMsg` dans la console.

```js
assert.match(code, /console\.log\(duplexLoanMsg\)/)
```

Tu dois avoir une variable appelée `condoLoanMsg`.

```js
assert.isNotNull(condoLoanMsg);
```

Ta variable `condoLoanMsg` doit être définie sur le résultat de `getLoanMessage(65000, 690)`.

```js
assert.equal(condoLoanMsg, getLoanMessage(65000, 690));
```

Tu dois afficher la variable `condoLoanMsg` dans la console.

```js
assert.match(code, /console\.log\(condoLoanMsg\)/)
```

Tu dois avoir une variable appelée `carLoanMsg`.

```js
assert.isNotNull(carLoanMsg);
```

Ta variable `carLoanMsg` doit être définie sur le résultat de `getLoanMessage(45000, 660)`.

```js
assert.equal(carLoanMsg, getLoanMessage(45000, 660));
```

Tu dois afficher la variable `carLoanMsg` dans la console.

```js
assert.match(code, /console\.log\(carLoanMsg\)/)
```

Tu dois avoir une variable appelée `noLoanMsg`.

```js
assert.isNotNull(noLoanMsg);
```

Ta variable `noLoanMsg` doit être définie sur le résultat de `getLoanMessage(25000, 550)`.

```js
assert.equal(noLoanMsg, getLoanMessage(25000, 550));
```

Tu dois afficher la variable `noLoanMsg` dans la console.

```js
assert.match(code, /console\.log\(noLoanMsg\)/)
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
  } else {
    return "You don't qualify for any loans."
  }
}

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
  } else {
    return "You don't qualify for any loans."
  }
}

const duplexLoanMsg = getLoanMessage(85000, 850)
console.log(duplexLoanMsg)

const condoLoanMsg = getLoanMessage(65000, 690)
console.log(condoLoanMsg)

const carLoanMsg = getLoanMessage(45000, 660)
console.log(carLoanMsg)

const noLoanMsg = getLoanMessage(25000, 550)
console.log(noLoanMsg)
```
