---
id: 66c8ba975ee7230e29f6c4af
title: Étape 4
challengeType: 1
dashedName: step-4
---

# --description--

Si le revenu annuel du demandeur est supérieur ou égal à `minIncomeForCondo`, ET si son score de crédit est supérieur ou égal à `minCreditScoreForCondo`, alors il est éligible à un prêt pour un appartement et une voiture.

Vérifie si c'est vrai dans la fonction `getLoanMessage`. Si c'est le cas, retourne la chaîne `"You qualify for a condo and car loan."`

# --hints--

Ta fonction `getLoanMessage` doit retourner une chaîne si le demandeur est éligible à un prêt pour un appartement et une voiture.

```js
assert.isString(getLoanMessage(45000, 680));
```

Ta fonction `getLoanMessage` doit retourner la chaîne `"You qualify for a condo and car loan."`.

```js
assert.strictEqual(getLoanMessage(45000, 680), "You qualify for a condo and car loan.");
```

Ta fonction `getLoanMessage` doit retourner `undefined` si le revenu annuel et le score de crédit du demandeur ne remplissent pas les conditions du prêt pour un appartement.

```js
assert.strictEqual(getLoanMessage(45000, 650), undefined);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
--fcc-editable-region--
  }

--fcc-editable-region--
}
```
