---
id: 66c8ba975ee7230e29f6c4ac
title: Étape 2
challengeType: 1
dashedName: step-2
---

# --description--

Lorsque l'utilisateur est éligible à un prêt, tu voudras lui afficher un message dans la console.

Pour cela, tu vas construire une fonction à l'intérieur de laquelle tu auras des vérifications qui retourneront le prêt auquel le demandeur est éligible.

Crée une fonction vide `getLoanMessage` avec les paramètres `annualIncome` et `creditScore`.

# --hints--

Tu dois créer une fonction `getLoanMessage`

```js
assert.isFunction(getLoanMessage)
```

Ta fonction `getLoanMessage` doit avoir `annualIncome` et `creditScore` comme paramètres.

```js
assert.match(getLoanMessage?.toString(), /\(\s*annualIncome,\s*creditScore\s*\)|\(\s*creditScore,\s*annualIncome\s*\)/);
```

`getLoanMessage` doit être une fonction vide.

```js
assert.match(getLoanMessage?.toString(), /\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

--fcc-editable-region--

--fcc-editable-region--
```
