---
id: 66c8ba975ee7230e29f6c4ad
title: Étape 3
challengeType: 1
dashedName: step-3
---

# --description--

Pour vérifier à quel prêt un utilisateur est éligible en fonction de `annualIncome` et `creditScore`, tu dois utiliser une instruction `if/else if` ou un ternaire directement à l'intérieur de la fonction `getLoanMessage`. Tu retourneras ensuite le message approprié dans le bloc de chaque condition.

En commençant par le prêt pour le duplex, vérifie si `annualIncome` est supérieur ou égal à `minIncomeForDuplex` ET si `creditScore` est supérieur ou égal à `minCreditScoreForDuplex`.

Si cette condition est vraie, alors le demandeur est éligible au prêt pour le duplex et aux autres prêts. Donc, à l'intérieur de la vérification, retourne la chaîne `"You qualify for a duplex, condo, and car loan."`

# --hints--

Ta fonction `getLoanMessage` doit retourner une chaîne.

```js
assert.isString(getLoanMessage(65000, 750));
assert.isString(getLoanMessage(60000, 750));
assert.isString(getLoanMessage(65000, 700));
assert.isString(getLoanMessage(60000, 700));
```

Ta fonction `getLoanMessage` doit retourner la chaîne `"You qualify for a duplex, condo, and car loan."`.

```js
assert.strictEqual(getLoanMessage(65000, 750), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(60000, 750), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(65000, 700), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(60000, 700), "You qualify for a duplex, condo, and car loan.");
```

Ta fonction `getLoanMessage` doit retourner `undefined` si le revenu annuel et le score de crédit du demandeur ne remplissent pas les conditions du prêt pour le duplex.

```js
assert.isUndefined(getLoanMessage(59000, 690));
assert.isUndefined(getLoanMessage(59000, 700));
assert.isUndefined(getLoanMessage(59000, 750));
assert.isUndefined(getLoanMessage(60000, 690));
assert.isUndefined(getLoanMessage(65000, 690));
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
--fcc-editable-region--

--fcc-editable-region--
}
```
