---
id: 66c8ba41a77db20a93f9d7a1
title: Étape 1
challengeType: 1
dashedName: step-1
---

# --description--

Dans ce workshop, tu vas réviser les conditionnels JavaScript en construisant une application de vérification d'éligibilité à un prêt.

L'application vérifiera si l'utilisateur est éligible à un prêt pour un duplex, une voiture et un appartement en fonction de son revenu annuel et de son score de crédit. Un score de crédit est un nombre qui représente ta capacité à gérer l'argent emprunté.

Pour commencer, crée les variables et valeurs suivantes.

| Nom de la variable | Valeur   |
| ----------- | ------- |
| `minIncomeForDuplex`       | `60000` |
| `minCreditScoreForDuplex`   | `700`   |
| `minIncomeForCondo`   | `45000`   |
| `minCreditScoreForCondo`    | `680`   |
| `minIncomeForCar`   | `30000`   |
| `minCreditScoreForCar`   | `650`   |

# --hints--

Tu dois créer une variable `minIncomeForDuplex` définie à `60000`.

```js
assert.equal(minIncomeForDuplex, 60000);
```

Tu dois créer une variable `minCreditScoreForDuplex` définie à `700`.

```js
assert.equal(minCreditScoreForDuplex, 700);
```

Tu dois créer une variable `minIncomeForCar` définie à `30000`.

```js
assert.equal(minIncomeForCar, 30000);
```

Tu dois créer une variable `minCreditScoreForCar` définie à `650`.

```js
assert.equal(minCreditScoreForCar, 650);
```

Tu dois créer une variable `minIncomeForCondo` définie à `45000`.

```js
assert.equal(minIncomeForCondo, 45000);
```

Tu dois créer une variable `minCreditScoreForCondo` définie à `680`.

```js
assert.equal(minCreditScoreForCondo, 680);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
