---
id: 66c8ba975ee7230e29f6c4b0
title: Étape 5
challengeType: 1
dashedName: step-5
---

# --description--

Maintenant, tu dois vérifier si le demandeur est éligible à un prêt pour une voiture uniquement. Si c'est le cas, retourne la chaîne `"You qualify for a car loan."`.

# --hints--

Ta fonction `getLoanMessage` doit retourner une chaîne si le demandeur est éligible à un prêt pour une voiture.

```js
assert.isString(getLoanMessage(30000, 650));
```

Ta fonction `getLoanMessage` doit retourner la chaîne `"You qualify for a car loan."`.

```js
assert.strictEqual(getLoanMessage(30000, 650), "You qualify for a car loan.");
```

Ta fonction `getLoanMessage` doit retourner `undefined` si le revenu annuel et le score de crédit du demandeur ne remplissent pas les conditions du prêt pour une voiture.

```js
assert.strictEqual(getLoanMessage(30000, 600), undefined);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
--fcc-editable-region--
  }

--fcc-editable-region--
}
```
