---
id: 67d83df6f82eda3868dd2a84
title: Crée un personnage de JDR
challengeType: 27
dashedName: build-an-rpg-character
---

# --description--

Dans ce laboratoire, tu vas t'entraîner aux bases de Python en construisant une petite appli qui crée un personnage pour une aventure de JDR.

**Objectif :** remplis les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu devrais avoir une fonction nommée `create_character`.
1. La fonction devrait accepter, dans l'ordre, un nom de personnage suivi de trois stats : force, intelligence et charisme.
1. Le nom du personnage devrait être validé :
   - Si le nom n'est pas une chaîne, la fonction devrait renvoyer `The character name should be a string`.
   - Si le nom est une chaîne vide, la fonction devrait renvoyer `The character should have a name`.
   - Si le nom fait plus de 10 caractères, la fonction devrait renvoyer `The character name is too long`.
   - Si le nom contient des espaces, la fonction devrait renvoyer `The character name should not contain spaces`.
1. Les stats devraient aussi être validées :
   - Si une ou plusieurs stats ne sont pas des entiers, la fonction devrait renvoyer `All stats should be integers`.
   - Si une ou plusieurs stats sont inférieures à 1, la fonction devrait renvoyer `All stats should be no less than 1`.
   - Si une ou plusieurs stats sont supérieures à 4, la fonction devrait renvoyer `All stats should be no more than 4`.
   - Si la somme de toutes les stats est différente de 7, la fonction devrait renvoyer `The character should start with 7 points`.
1. Si toutes les valeurs passent la vérification, la fonction devrait renvoyer une chaîne de quatre lignes :
   - la première ligne devrait contenir le nom du personnage
   - les lignes 2 à 4 devraient commencer par l'abréviation de la stat, `STR`, `INT` ou `CHA` (dans cet ordre), puis un espace, puis un nombre de points pleins (`●`) égal à la valeur de la stat, et un nombre de points vides (`○`) pour atteindre 10. Exemple : si la force vaut 3, il doit y avoir 3 points pleins suivis de 7 points vides. Les points sont donnés dans l'éditeur.

Voici la chaîne que devrait renvoyer `create_character('ren', 4, 2, 1)` :

   ```md
   ren
   STR ●●●●○○○○○○
   INT ●●○○○○○○○○
   CHA ●○○○○○○○○○
   ```

NOTE : même si `str` et `int` sont des abréviations courantes pour les stats, rappelle-toi que ce sont des mots-clés réservés en Python et qu'ils ne doivent pas servir de noms de variables.

# --hints--

Tu devrais avoir une fonction nommée `create_character`.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_function('create_character')
`
    )
})
```

Quand `create_character` est appelée avec un premier argument qui n'est **pas** une chaîne, elle devrait renvoyer `The character name should be a string`.

```js
({
    test: () => runPython(
        `
wrong_types = [3, 5, {'a': 3}, 3.1, [1, 2], (1, 2), {1, 2}]
for wrong_type in wrong_types:
    assert create_character(wrong_type, 1, 2, 4) == 'The character name should be a string'
`
    )
})
```

Quand `create_character` est appelée avec un premier argument qui est une chaîne, elle **ne devrait pas** renvoyer `The character name should be a string`.

```js
({
    test: () => runPython(
        `assert create_character('ren', 1, 2, 4) != 'The character name should be a string'`
    )
})
```

Quand `create_character` est appelée avec un premier argument qui est une chaîne vide, elle devrait renvoyer `The character should have a name`.

```js
({
    test: () => runPython(
        `
assert create_character('', 4, 2, 1) == 'The character should have a name'
`
    )
})
```

Quand `create_character` est appelée avec un premier argument qui n'est **pas** une chaîne vide, elle **ne devrait pas** renvoyer `The character should have a name`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'The character should have a name'
`
    )
})
```

Quand `create_character` est appelée avec un premier argument de plus de 10 caractères, elle devrait renvoyer `The character name is too long`.

```js
({
    test: () => runPython(
        `
for length in range(11, 27):
    assert create_character('abcdefghijklmnopqrstuvwxyz'[:length], 4, 2, 1) == 'The character name is too long'
`
    )
})
```

La fonction `create_character` **ne devrait pas** dire que le nom est trop long quand il ne dépasse pas 10 caractères.

```js
({
    test: () => runPython(
        `
for length in range(0, 11):
  assert create_character('aaabbcccdd'[:length], 4, 2, 1) != 'The character name is too long'
`
    )
})
```

Quand `create_character` est appelée avec un premier argument qui contient un espace, elle devrait renvoyer `The character name should not contain spaces`.

```js
({
    test: () => runPython(
        `
assert create_character('cha cha', 4, 1, 2) == 'The character name should not contain spaces'
`
    )
})
```

Quand `create_character` est appelée avec un premier argument qui *ne contient pas* d'espace, elle *ne devrait pas* renvoyer `The character name should not contain spaces`.

```js
({
    test: () => runPython(
        `
assert create_character('chacha', 4, 2, 1) != 'The character name should not contain spaces'
`
    )
})
```

Quand `create_character` est appelée avec un deuxième, troisième ou quatrième argument qui n'est **pas** un entier, elle devrait renvoyer `All stats should be integers`.

```js
({
    test: () => runPython(
        `
wrong_types = ['3', '5', {'a': 3}, 3.1, [1, 2], (1, 2), {1, 2}, 'str', 'friend']
for wrong_type in wrong_types:
    assert create_character('friend', wrong_type, 2, 1) == 'All stats should be integers'
    assert create_character('friend', 2, wrong_type, 1) == 'All stats should be integers'
    assert create_character('friend', 2, 1, wrong_type) == 'All stats should be integers'
`
    )
})
```

Quand `create_character` est appelée avec un deuxième, troisième et quatrième argument qui sont tous des entiers, elle **ne devrait pas** renvoyer `All stats should be integers`.

```js
({
    test: () => runPython(
        `
assert create_character('friend', 4, 2, 1) != 'All stats should be integers'
`
    )
})
```

Quand `create_character` est appelée avec un deuxième, troisième ou quatrième argument inférieur à `1`, elle devrait renvoyer `All stats should be no less than 1`.

```js
({
    test: () => runPython(
        `
expected = 'All stats should be no less than 1'

assert create_character('ren', 0, 4, 3) == expected
assert create_character('ren', 4, 0, 3) == expected
assert create_character('ren', 4, 3, 0) == expected
assert create_character('ren', -1, 4, 4) == expected
assert create_character('ren', 4, -1, 4) == expected
assert create_character('ren', 4, 4, -1) == expected
`
    )
})
```

Quand `create_character` est appelée avec un deuxième, troisième et quatrième argument tous au moins égaux à `1`, elle **ne devrait pas** renvoyer `All stats should be no less than 1`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'All stats should be no less than 1'
`
    )
})
```

Quand `create_character` est appelée avec un deuxième, troisième ou quatrième argument supérieur à `4`, elle devrait renvoyer `All stats should be no more than 4`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 1, 5, 1) == 'All stats should be no more than 4'
assert create_character('ren', 5, 1, 1) == 'All stats should be no more than 4'
assert create_character('ren', 1, 1, 5) == 'All stats should be no more than 4'
`
    )
})
```

Quand `create_character` est appelée avec un deuxième, troisième et quatrième argument tous au plus égaux à `4`, elle **ne devrait pas** renvoyer `All stats should be no more than 4`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'All stats should be no more than 4'
`
    )
})
```

Quand `create_character` est appelée avec un deuxième, troisième ou quatrième argument dont la somme **n'est pas** `7`, elle devrait renvoyer `The character should start with 7 points`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 4, 4) == 'The character should start with 7 points'
assert create_character('ren', 1, 1, 1) == 'The character should start with 7 points'
assert create_character('ren', 1, 2, 3) == 'The character should start with 7 points'
`
    )
})
```

Quand `create_character` est appelée avec un deuxième, troisième et quatrième argument dont la somme est `7`, elle **ne devrait pas** renvoyer `The character should start with 7 points`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'The character should start with 7 points'
`
    )
})
```

`create_character('ren', 4, 2, 1)` devrait renvoyer `ren\nSTR ●●●●○○○○○○\nINT ●●○○○○○○○○\nCHA ●○○○○○○○○○`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) == f'ren\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
    )
})
```

Quand `create_character` est appelée avec des valeurs valides, elle devrait afficher les stats du personnage comme demandé.

```js
({
    test: () => runPython(
        `
assert create_character('Bob', 1, 2, 4) == f'Bob\\nSTR ●○○○○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●●●●○○○○○○'
assert create_character('aaabbbcccd', 4, 2, 1) == f'aaabbbcccd\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
    )
})

```

# --seed--

## --seed-contents--

```py
full_dot = '●'
empty_dot = '○'


```

# --solutions--

```py
full_dot = '●'
empty_dot = '○'

def create_character(character_name, strength, intelligence, charisma):
    if type(character_name) != str:
        return 'The character name should be a string'
    if len(character_name) == 0:
        return 'The character should have a name'
    if len(character_name) > 10:
        return 'The character name is too long'
    if ' ' in character_name:
        return 'The character name should not contain spaces'

    if type(strength) != int or type(intelligence) != int or type(charisma) != int:
       return 'All stats should be integers'
    if strength < 1 or intelligence < 1 or charisma < 1:
        return 'All stats should be no less than 1'
    if strength > 4 or intelligence > 4 or charisma > 4:
        return 'All stats should be no more than 4'
    if strength + intelligence + charisma != 7:
        return 'The character should start with 7 points'

    character_sheet = f'{character_name}\n'
    character_sheet += f'STR {full_dot * strength}{empty_dot * (10 - strength)}\n'
    character_sheet += f'INT {full_dot * intelligence}{empty_dot * (10 - intelligence)}\n'
    character_sheet += f'CHA {full_dot * charisma}{empty_dot * (10 - charisma)}'
    return character_sheet
```
