---
id: 673b02b03134b04637bf7055
title: Créer un pied de page réutilisable
challengeType: 25
dashedName: build-a-reusable-footer
demoType: onClick
---

# --description--

**Objectif :** Remplis les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User Stories :**

1. Tu dois exporter un composant `Footer`.
2. Ton composant `Footer` doit renvoyer un élément `footer` qui contient tous les autres éléments.
3. Ton élément `footer` doit avoir :

   - Au moins trois listes non ordonnées, chacune avec au moins deux éléments de liste.
   - Au moins un paragraphe avec un symbole de copyright `©`.
   - Au moins trois liens avec la valeur `href` définie à `#` et le contenu du lien défini à une icône ou un texte de ton choix.

# --hints--

Tu dois exporter un composant `Footer`.

```js
assert.isFunction(window.index.Footer);
```

Ton composant `Footer` doit renvoyer un élément `footer`.

```js
const footer = document.getElementById('root').firstElementChild;
assert.equal(footer.tagName, 'FOOTER');
```

Ton composant `Footer` ne doit contenir qu'un seul élément `footer`.

```js
assert.lengthOf(document.querySelectorAll('footer'), 1);
```

Ton élément `footer` ne doit pas avoir de frères et sœurs.

```js
assert.lengthOf(document.getElementById('root').children, 1);
```

Ton élément `footer` doit contenir au moins trois listes non ordonnées.

```js
assert.isAtLeast(document.querySelectorAll('ul').length, 3);
```

Chacune de tes listes non ordonnées doit avoir au moins deux éléments de liste.

```js
const uls = [...document.querySelectorAll('ul')];
uls.forEach(ul => {
    assert.isAtLeast(ul.querySelectorAll('li').length, 2);
  }
);
```

Ton `footer` doit avoir au moins un élément paragraphe.

```js
assert.isAtLeast(document.querySelectorAll('p').length, 1);
```

Tu dois avoir un symbole de copyright (`©`) dans l'un de tes paragraphes.

```js
const texts = [...document.querySelectorAll('p')].map(p => p.textContent);
const hasCopyright = texts.some(text => text.includes('©'));
assert.isTrue(hasCopyright);
```

Ton `footer` doit avoir au moins trois liens avec la valeur `href` définie à `#`.

```js
const links = document.querySelectorAll('a[href="#"]');
assert.isAtLeast(links.length, 3);
```

Aucun de tes liens ne doit être vide.

```js
const links = [...document.querySelectorAll('a[href="#"]')];
links.forEach(link => {
    assert.isNotEmpty(link.textContent);
  }
);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Reusable Footer Component</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { Footer } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<Footer />);
    </script>
  </body>
</html>
```

```css

```

```jsx

```

# --solutions--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Reusable Footer Component</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { Footer } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<Footer />);
    </script>
  </body>
</html>
```

```css
.footer {
  background-color: #e3e1eb;
  padding: 20px;
  text-align: center;
  border-top: 1px solid #e5e7eb;
  font-family: Arial, sans-serif;
}

.footer-section {
  display: grid;
  grid-template-columns: repeat(3, 1fr);
  gap: 20px;
  justify-content: center;
  margin-bottom: 20px;
  align-items: center;
}

.footer ul {
  list-style: none;
  padding: 0;
  margin: 0;
}

.footer li {
  margin-bottom: 10px;
}

.footer a {
  text-decoration: none;
  color: #1f2937;
}

.footer a:hover {
  color: #6366f1;
}

.footer p {
  color: #6b7280;
  margin-top: 20px;
}

.footer-icons {
  margin-top: 20px;
  display: flex;
  justify-content: center;
}

.footer-icons a {
  margin: 0 10px;
  font-size: 24px;
  color: #6b7280;
  text-decoration: none;
}

.footer-icons a:hover {
  color: #6366f1;
}
```

```jsx
export const Footer = () => (
  <footer className='footer'>
    <div className='footer-section'>
      <ul>
        <li>
          <a href='#'>Fitness Dashboard</a>
        </li>
        <li>
          <a href='#'>Services</a>
        </li>
      </ul>
      <ul>
        <li>
          <a href='#'>Watch Videos</a>
        </li>
        <li>
          <a href='#'>Discord</a>
        </li>
      </ul>
      <ul>
        <li>
          <a href='#'>Privacy Policy</a>
        </li>
        <li>
          <a href='#'>Terms & Conditions</a>
        </li>
      </ul>
    </div>
    <p>© 2024 Fitness Dashboard. All Rights Reserved.</p>
    <div className='footer-icons'>
      <a href='#'>&#x1F3AE;</a>
      <a href='#'>&#x1F426;</a>
      <a href='#'>&#x1F4BB;</a>
      <a href='#'>&#x1F3C0;</a>
    </div>
  </footer>
);
```
