---
id: 69008059d43d6751385883c1
title: Déboguer des erreurs d'opérateurs d'incrémentation et de décrémentation dans une application boguée
challengeType: 26
dashedName: debug-increment-and-decrement-operator-errors-in-a-buggy-app
---

# --description--

Un ami t'a demandé de l'aider à déboguer son code. Il s'entraîne aux opérateurs d'incrémentation et de décrémentation mais obtient des résultats inattendus dans la console.

Ta tâche est de l'aider à identifier les erreurs et à mettre à jour le code pour que les sorties dans la console correspondent aux résultats attendus.

Tu peux te référer aux leçons précédentes sur les opérateurs d'incrémentation et de décrémentation ainsi que sur la notation préfixe et postfixe.

**User Stories :**

1. Tu devrais avoir une variable nommée `updatedCounter` qui produit correctement la valeur `11` à l'aide de l'opérateur d'incrémentation. Mets à jour l'expression pour utiliser la bonne notation préfixe ou postfixe au lieu de coder la valeur en dur.
2. Tu devrais avoir une variable nommée `finalScore` qui produit correctement la valeur `8` à l'aide de l'opérateur d'incrémentation. Mets à jour l'expression pour utiliser la bonne notation préfixe ou postfixe au lieu de coder la valeur en dur.
3. Tu devrais avoir une variable nommée `updatedCoins` qui produit correctement la valeur `2` à l'aide de l'opérateur de décrémentation. Mets à jour l'expression pour utiliser la bonne notation préfixe ou postfixe au lieu de coder la valeur en dur.
4. Tu devrais avoir une variable nommée `newHealth` qui produit correctement la valeur `7` à l'aide de l'opérateur de décrémentation. Mets à jour l'expression pour utiliser la bonne notation préfixe ou postfixe au lieu de coder la valeur en dur.

# --hints--

Tu devrais avoir une variable nommée `updatedCounter` qui utilise la notation d'incrémentation préfixe (`++counter`).

```js
assert.match(code, /updatedCounter\s*=\s*\+\+counter\s*;/);
assert.strictEqual(typeof updatedCounter, 'number');
```

`updatedCounter` devrait être égal à `11`.

```js
assert.strictEqual(updatedCounter, 11);
```

Tu devrais avoir une variable nommée `finalScore` qui utilise la notation d'incrémentation postfixe (`score++`).

```js
assert.match(code, /finalScore\s*=\s*score\+\+\s*;/);
assert.strictEqual(typeof finalScore, 'number');
```

`finalScore` devrait être égal à `8`.

```js
assert.strictEqual(finalScore, 8);
```

Tu devrais avoir une variable nommée `updatedCoins` qui utilise la notation de décrémentation préfixe (`--coins`).

```js
assert.match(code, /updatedCoins\s*=\s*--coins\s*;/);
assert.strictEqual(typeof updatedCoins, 'number');
```

`updatedCoins` devrait être égal à `2`.

```js
assert.strictEqual(updatedCoins, 2);
```

Tu devrais avoir une variable nommée `newHealth` qui utilise la notation de décrémentation postfixe (`health--`).

```js
assert.match(code, /newHealth\s*=\s*health--\s*;/);
assert.strictEqual(typeof newHealth, 'number');
```

`newHealth` devrait être égal à `7`.

```js
assert.strictEqual(newHealth, 7);
```

# --seed--

## --seed-contents--

```js
let counter = 10;
console.log(`Counter is currently ${counter}`);

let updatedCounter = counter++;
console.log(`Counter is now ${updatedCounter}`);

let score = 8;
console.log(`Score is currently ${score}`);

let finalScore = ++score;
console.log(`Score is still ${finalScore}`);

let coins = 3;
console.log(`Coins is currently ${coins}`);

let updatedCoins = coins--;
console.log(`Coins is now ${updatedCoins}`);

let health = 7;
console.log(`Health is currently ${health}`);

let newHealth = --health;
console.log(`Health is still ${newHealth}`);
```

# --solutions--

```js
let counter = 10;
console.log(`Counter is currently ${counter}`);

let updatedCounter = ++counter;
console.log(`Counter is now ${updatedCounter}`);

let score = 8;
console.log(`Score is currently ${score}`);

let finalScore = score++;
console.log(`Score is still ${finalScore}`);

let coins = 3;
console.log(`Coins is currently ${coins}`);

let updatedCoins = --coins;
console.log(`Coins is now ${updatedCoins}`);

let health = 7;
console.log(`Health is currently ${health}`);

let newHealth = health--;
console.log(`Health is still ${newHealth}`);
```
