---
id: afcc8d540bea9ea2669306b6
title: Crée une fonction de répétition de chaîne
challengeType: 26
dashedName: build-a-string-repeating-function
---

# --description--

Dans ce lab, tu vas créer une fonction qui répète une chaîne donnée un nombre précis de fois. Pour ce lab, n'utilise _pas_ la méthode intégrée `.repeat()`.

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories :**

1. Tu dois créer une fonction nommée `repeatStringNumTimes` qui prend deux paramètres : une chaîne et un nombre.
2. La fonction doit renvoyer la chaîne répétée le nombre de fois indiqué.
3. Si le nombre est inférieur ou égal à zéro, la fonction doit renvoyer une chaîne vide.

# --hints--

Tu dois créer une fonction nommée `repeatStringNumTimes`.

```js
assert.isFunction(repeatStringNumTimes);
```

`repeatStringNumTimes` doit prendre deux paramètres.

```js
assert.lengthOf(repeatStringNumTimes, 2);
```

La fonction `repeatStringNumTimes` doit toujours renvoyer une chaîne.

```js
assert.isString(repeatStringNumTimes('hello', 3));
```

`repeatStringNumTimes("*", 3)` doit renvoyer la chaîne `***`.

```js
assert.strictEqual(repeatStringNumTimes('*', 3), '***');
```

`repeatStringNumTimes("abc", 3)` doit renvoyer la chaîne `abcabcabc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 3), 'abcabcabc');
```

`repeatStringNumTimes("abc", 4)` doit renvoyer la chaîne `abcabcabcabc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 4), 'abcabcabcabc');
```

`repeatStringNumTimes("abc", 1)` doit renvoyer la chaîne `abc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 1), 'abc');
```

`repeatStringNumTimes("*", 8)` doit renvoyer la chaîne `********`.

```js
assert.strictEqual(repeatStringNumTimes('*', 8), '********');
```

`repeatStringNumTimes("abc", -2)` doit renvoyer une chaîne vide (`""`).

```js
assert.isEmpty(repeatStringNumTimes('abc', -2));
```

`repeatStringNumTimes("abc", 0)` doit renvoyer `""`.

```js
assert.isEmpty(repeatStringNumTimes('abc', 0));
```

La méthode intégrée `repeat()` ne doit pas être utilisée.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.repeat/g);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function repeatStringNumTimes(str, num) {
  if (num < 1) return '';
  return num === 1 ? str : str + repeatStringNumTimes(str, num - 1);
}

repeatStringNumTimes('abc', 3);
```
