---
id: 66ea89c25e1bd871688eba6a
title: Étape 6
challengeType: 1
dashedName: step-6
---

# --description--

Jusqu'à présent, tu n'as travaillé qu'avec des nombres à virgule flottante et le nombre `0`.

Pour la partie suivante du workshop, tu réviseras comment utiliser la méthode `Math.floor()`.

Commence par ajouter un `console.log()` qui affiche le message `"The Math.floor() method rounds the value down to the nearest whole integer."`

# --hints--

Tu devrais afficher le message `"The Math.floor() method rounds the value down to the nearest whole integer."`

```js
assert.match(code, /console\.log\((["'])The\s+Math\.floor\(\)\s+method\s+rounds\s+the\s+value\s+down\s+to\s+the\s+nearest\s+whole\s+integer\.\1\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);
--fcc-editable-region--

--fcc-editable-region--
```
