---
id: 66ea7fc048a6e26e33c598fe
title: Étape 3
challengeType: 1
dashedName: step-3
---

# --description--

Maintenant, il est temps de générer un nombre aléatoire à l'aide de la méthode `Math.random()`.

Crée une variable appelée `randomNum` et affecte-lui la valeur du résultat de l'appel de la méthode `Math.random()`.

Ensuite, affiche la variable `randomNum` dans la console.

Essaie d'ajouter un espace dans le code pour relancer le bot et voir différents nombres aléatoires générés.

# --hints--

Tu devrais avoir une variable appelée `randomNum`.

```js
assert.isNotNull(randomNum);
```

Tu devrais affecter le résultat de l'appel de la méthode `Math.random()` à la variable `randomNum`.

```js
assert.match(code, /randomNum\s*=\s*Math\.random\(\);?/);
```

Tu ne devrais pas coder en dur un nombre aléatoire pour la variable `randomNum`. Veille à utiliser la méthode `Math.random()` pour générer le nombre aléatoire.

```js
assert.notMatch(code, /randomNum\s*=\s*0\.\d+;?/);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomNum\s*=\s*\d+;?/);
```

Tu devrais afficher la variable `randomNum` dans la console.

```js
assert.match(code, /console\.log\(\s*randomNum\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");
--fcc-editable-region--

--fcc-editable-region--
```
