---
id: 66ea8239b9eeb76f12299030
title: Étape 4
challengeType: 1
dashedName: step-4
---

# --description--

La portion suivante du workshop consiste à réviser comment générer un nombre aléatoire entre deux valeurs.

Commence par ajouter un autre `console.log()` qui affiche le message `"Now, generate a random number between two values."` dans la console.

Ensuite, crée une variable appelée `min` et affecte-lui la valeur `1`, puis crée une variable appelée `max` et affecte-lui la valeur `100`.

À l'étape suivante, tu généreras un nombre aléatoire entre ces deux valeurs.

# --hints--

Tu devrais avoir un `console.log()` qui affiche le message `"Now, generate a random number between two values."`.

```js
assert.match(code, /console\.log\((`|'|")\s*Now,\s+generate\s+a\s+random\s+number\s+between\s+two\s+values\.\s*\1\);?/);
```

Tu devrais avoir une variable appelée `min`.

```js
assert.isNotNull(min);
```

Ta variable `min` devrait être un nombre.

```js
assert.isNumber(min);
```

Tu devrais affecter le nombre `1` à ta variable `min`.

```js
assert.equal(min, 1);
```

Tu devrais avoir une variable appelée `max`.

```js
assert.isNotNull(max);
```

Ta variable `max` devrait être un nombre.

```js
assert.isNumber(max);
```

Tu devrais affecter le nombre `100` à ta variable `max`.

```js
assert.equal(max, 100);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);
--fcc-editable-region--

--fcc-editable-region--
```
