---
id: 66ea845f54195c7020cc521b
title: Étape 5
challengeType: 1
dashedName: step-5
---

# --description--

La formule pour générer un nombre aléatoire entre deux valeurs est la suivante :

```js
Math.random() * (maximum - minimum) + minimum;
```

Cela produira un résultat qui est un nombre à virgule flottante entre deux valeurs.

Crée une variable appelée `randomNum2` et affecte-lui le résultat de la génération d'une valeur entre les valeurs `min` et `max`.

Ensuite, affiche la variable `randomNum2` pour voir le résultat. Essaie de rafraîchir la page pour voir différents résultats.

# --hints--

Tu devrais avoir une variable appelée `randomNum2`.

```js
assert.isNotNull(randomNum2);
```

Ta variable `randomNum2` devrait être un nombre.

```js
assert.isNumber(randomNum2);
```

Tu devrais affecter le résultat de la génération d'une valeur entre les valeurs `min` et `max`. Veille à ne pas coder une valeur en dur. Tu devrais utiliser la formule fournie ci-dessus.

```js
assert.isAtLeast(randomNum2, min);
assert.isAtMost(randomNum2, max);

assert.notMatch(code, /randomNum2\s*=\s*0\.\d+;?/);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomNum2\s*=\s*\d+;?/);
```

Tu devrais afficher la variable `randomNum2` dans la console.

```js
assert.match(code, /console\.log\(\s*randomNum2\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;
--fcc-editable-region--

--fcc-editable-region--
```
