---
id: 66ea7adae8053065a64f9002
title: Étape 1
challengeType: 1
dashedName: step-1
---

# --description--

Dans ce workshop, tu réviseras l'utilisation des différentes méthodes de l'objet `Math` en construisant un Mathbot. Ce Mathbot affichera quelques opérations mathématiques et messages dans la console dans le but de t'enseigner l'objet `Math` de JavaScript.

Pour cette première étape, commence par créer une variable appelée `botName` et affecte-lui la valeur de chaîne `"MathBot"`.

Ensuite, crée une autre variable appelée `greeting` et affecte-lui la phrase `"Hi there! My name is [botName goes here] and I am here to teach you about the Math object!"`. À la place de l'espace réservé `[botName goes here]`, utilise la variable `botName`.

Tu es libre d'utiliser des littéraux de modèle ou la concaténation de chaînes avec l'opérateur `+` pour y parvenir.

Enfin, affiche la variable `greeting` dans la console.

# --hints--

Tu devrais avoir une variable appelée `botName`.

```js
assert.isNotNull(botName);
```

Ta variable `botName` devrait être une chaîne.

```js
assert.isString(botName);
```

Tu devrais affecter la valeur de chaîne `"MathBot"` à la variable `botName`.

```js
assert.strictEqual(botName, "MathBot");
```

Tu devrais avoir une variable appelée `greeting`.

```js
assert.isNotNull(greeting);
```

Ta variable `greeting` devrait être une chaîne.

```js
assert.isString(greeting);
```

Tu devrais affecter la phrase suivante à la variable `greeting` : `"Hi there! My name is [MathBot goes here] and I am here to teach you about the Math object!"`. Veille à remplacer `[MathBot goes here]` par la variable `botName`.

```js
assert.strictEqual(greeting, "Hi there! My name is MathBot and I am here to teach you about the Math object!");
```

Tu devrais afficher la variable `greeting` dans la console.

```js
assert.match(code, /console\.log\(greeting\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
