---
id: 66eb3ec3b07dfe8f898646c0
title: Étape 7
challengeType: 1
dashedName: step-7
---

# --description--

Comme tu l'as appris dans les leçons précédentes, la méthode `Math.floor()` arrondit la valeur vers le bas à l'entier le plus proche.

```js
const price = 10.99;
Math.floor(price); // 10
```

Crée une variable appelée `numRoundedDown` et affecte-lui le résultat de l'arrondi vers le bas du nombre à virgule flottante `6.7` à l'entier le plus proche.

Ensuite, affiche la variable `numRoundedDown` dans la console pour voir le résultat.

# --hints--

Tu devrais avoir une variable `numRoundedDown`.

```js
assert.isNotNull(numRoundedDown);
```

La variable `numRoundedDown` devrait être un nombre.

```js
assert.isNumber(numRoundedDown);
```

Ta variable `numRoundedDown` devrait être égale au résultat de l'arrondi vers le bas du nombre à virgule flottante `6.7` à l'entier le plus proche.

```js
assert.equal(numRoundedDown, 6);
```

Tu ne devrais pas coder en dur la valeur `6` pour la variable `numRoundedDown`.

```js
assert.notMatch(code, /numRoundedDown\s*=\s*6/);
```

Tu devrais afficher la variable `numRoundedDown` dans la console.

```js
assert.match(code, /console\.log\(\s*numRoundedDown\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");
--fcc-editable-region--

--fcc-editable-region--
```
