---
id: 66eb488c0129e192542f0865
title: Étape 11
challengeType: 1
dashedName: step-11
---

# --description--

Dans les leçons précédentes, tu as appris à utiliser la méthode `Math.ceil()` comme ceci :

```js
const price = 10.01;
Math.ceil(price); // 11
```

La méthode `Math.ceil()` arrondit la valeur vers le haut à l'entier le plus proche.

Crée une variable appelée `numRoundedUp` et affecte-lui le résultat de l'arrondi vers le haut du nombre à virgule flottante `3.2` à l'entier le plus proche.

Ensuite, affiche la variable `numRoundedUp` dans la console pour voir le résultat.

# --hints--

Tu devrais avoir une variable `numRoundedUp`.

```js
assert.isNotNull(numRoundedUp);
```

La variable `numRoundedUp` devrait être un nombre.

```js
assert.isNumber(numRoundedUp);
```

Ta variable `numRoundedUp` devrait être égale au résultat de l'arrondi vers le haut du nombre à virgule flottante `3.2` à l'entier le plus proche.

```js
assert.equal(numRoundedUp, 4);
```

Tu ne devrais pas coder en dur la valeur `4` pour la variable `numRoundedUp`.

```js
assert.notMatch(code, /numRoundedUp\s*=\s*4/);
```

Tu devrais afficher la variable `numRoundedUp` dans la console.

```js
assert.match(code, /console\.log\(\s*numRoundedUp\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");

const randomInt = Math.floor(Math.random() * (max - min) + min);
console.log(randomInt);

console.log("The Math.ceil() method rounds the value up to the nearest whole integer.");

--fcc-editable-region--

--fcc-editable-region--
```
