---
id: 66ece147323c91f654bf8dd9
title: Étape 15
challengeType: 1
dashedName: step-15
---

# --description--

Dans les leçons précédentes, tu as appris à obtenir les valeurs maximale et minimale d'une plage de nombres comme ceci :

```js
Math.max(1, 2, 3, 4, 5); // 5
Math.min(1, 2, 3, 4, 5); // 1
```

Crée une variable appelée `maxNum` et affecte-lui le résultat de la recherche du nombre maximal parmi les nombres `3, 125, 55, 24`. Ensuite, affiche la valeur de `maxNum` dans la console.

En dessous, crée une variable appelée `minNum` et affecte-lui le résultat de la recherche du nombre minimal parmi les nombres `6, 90, 14, 90, 2`. Ensuite, affiche la valeur de `minNum` dans la console.

# --hints--

Tu devrais avoir une variable appelée `maxNum`.

```js
assert.isNotNull(maxNum);
```

Ta variable `maxNum` devrait recevoir le résultat de la recherche du nombre maximal parmi les nombres `3, 125, 55, 24`.

```js
assert.equal(maxNum, 125);
```

Tu ne devrais pas coder en dur la valeur `125` dans la variable `maxNum`. Veille à utiliser plutôt la méthode `Math.max()`.

```js
assert.notMatch(code, /maxNum\s*=\s*125/);
```

Tu devrais afficher la valeur de `maxNum` dans la console.

```js
assert.match(code, /console\.log\(\s*maxNum\s*\)/);
```

Tu devrais avoir une variable appelée `minNum`.

```js
assert.isNotNull(minNum);
```

Ta variable `minNum` devrait recevoir le résultat de la recherche du nombre minimal parmi les nombres `6, 90, 14, 90, 2`.

```js
assert.equal(minNum, 2);
```

Tu ne devrais pas coder en dur la valeur `2` dans la variable `minNum`. Veille à utiliser plutôt la méthode `Math.min()`.

```js
assert.notMatch(code, /minNum\s*=\s*2/);
```

Tu devrais afficher la valeur de `minNum` dans la console.

```js
assert.match(code, /console\.log\(\s*minNum\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");

const randomInt = Math.floor(Math.random() * (max - min) + min);
console.log(randomInt);

console.log("The Math.ceil() method rounds the value up to the nearest whole integer.");

const numRoundedUp = Math.ceil(3.2);
console.log(numRoundedUp);

console.log(
  "The Math.round() method rounds the value to the nearest whole integer."
);

const numRounded = Math.round(2.7);
console.log(numRounded);
const numRounded2 = Math.round(11.2);
console.log(numRounded2);

console.log("The Math.max() and Math.min() methods are used to get the maximum and minimum number from a range.");

--fcc-editable-region--

--fcc-editable-region--
```
