---
id: 69218e46f4490b33f2481a1a
title: Étape 9
challengeType: 1
dashedName: step-9
---

# --description--

La formule pour générer un entier aléatoire entre deux valeurs est la suivante :

```js
Math.floor(Math.random() * (maximum - minimum) + minimum);
```

Cela produira un résultat qui est un entier entre deux valeurs.

Crée une variable appelée `randomInt` et affecte-lui le résultat de la génération d'une valeur entre les valeurs `min` et `max`.

Ensuite, affiche la variable `randomInt` pour voir le résultat. Essaie de rafraîchir la page pour voir différents résultats.

# --hints--

Tu devrais avoir une variable appelée `randomInt`.

```js
assert.isNotNull(randomInt);
```

Ta variable `randomInt` devrait être un nombre.

```js
assert.isNumber(randomInt);
```

Tu devrais affecter le résultat de la génération d'une valeur entre les valeurs `min` et `max`. Veille à ne pas coder une valeur en dur. Tu devrais utiliser la formule fournie ci-dessus.

```js
assert.isAtLeast(randomInt, min);
assert.isAtMost(randomInt, max);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomInt\s*=\s*\d+;?/);
```

Tu devrais afficher la variable `randomInt` dans la console.

```js
assert.match(code, /console\.log\(\s*randomInt\s*\);?/);
```


# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");
--fcc-editable-region--

--fcc-editable-region--
```
