---
id: a5229172f011153519423690
title: Crée un calculateur de somme des Fibonacci impairs
challengeType: 26
dashedName: build-an-odd-fibonacci-sum-calculator
---

# --description--

Dans ce lab, tu vas créer un calculateur de somme des Fibonacci impairs qui calcule la somme de tous les nombres de Fibonacci impairs inférieurs ou égaux à un entier positif donné.

**Objectif :** réalise les user stories ci-dessous et fais passer tous les tests pour terminer le lab.

**User stories**

1. Tu dois avoir une fonction `sumFibs` qui accepte un nombre en argument.
1. La fonction `sumFibs` doit renvoyer la somme de tous les nombres de Fibonacci impairs inférieurs ou égaux au nombre donné.
1. La suite de Fibonacci commence par `0` et `1`, et chaque nombre suivant est la somme des deux précédents.
1. Seuls les nombres de Fibonacci impairs doivent être ajoutés à la somme.

# --hints--

Tu dois avoir une fonction `sumFibs`.

```js
assert.isFunction(sumFibs);
```

`sumFibs(1)` doit renvoyer un nombre.

```js
assert.isNumber(sumFibs(1));
```

`sumFibs(1000)` doit renvoyer `1785`.

```js
assert.strictEqual(sumFibs(1000), 1785);
```

`sumFibs(4000000)` doit renvoyer `4613732`.

```js
assert.strictEqual(sumFibs(4000000), 4613732);
```

`sumFibs(4)` doit renvoyer `5`.

```js
assert.strictEqual(sumFibs(4), 5);
```

`sumFibs(75024)` doit renvoyer `60696`.

```js
assert.strictEqual(sumFibs(75024), 60696);
```

`sumFibs(75025)` doit renvoyer `135721`.

```js
assert.strictEqual(sumFibs(75025), 135721);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sumFibs(num) {
  let a = 0, b = 1;
  let sum = 0;

  while (b <= num) {
    if (b % 2 !== 0) {
      sum += b;
    }
    [a, b] = [b, a + b];
  }

  return sum;
}
```
